[1566687602] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.93%
[1566687753] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566687832] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 19.39%
[1566687903] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.71%
[1566688031] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566688216] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566688402] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566688788] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566688939] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 63.39%
[1566691102] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566691260] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566691413] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566691463] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566691534] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.52%
[1566691614] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566691685] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.32%
[1566691765] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566695043] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 38.1%
[1566695396] SERVICE FLAPPING ALERT: ieatnetsimv5116-46;check netsim cpu;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1566695396] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 50.1%
[1566695668] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1566695668] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566696451] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.02%
[1566696601] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.69%
[1566698677] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566698830] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566698982] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.17%
[1566702628] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566702779] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566702930] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.47%
[1566703951] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.19%
[1566704102] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.72%
[1566706560] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 37.18%
[1566709836] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.84%
[1566709987] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.56%
[1566710138] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.4%
[1566710288] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.96%
[1566710439] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.19%
[1566717014] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566717165] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 61.02%
[1566719023] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566720869] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566721031] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566721279] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566722253] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.38%
[1566722404] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.53%
[1566722900] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.92%
[1566723051] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 23.72%
[1566723202] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.03%
[1566723270] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566723352] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 23.77%
[1566723445] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566723667] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566724518] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.68%
[1566724533] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1566724675] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566724831] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566726155] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.12%
[1566726305] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.18%
[1566728467] SERVICE FLAPPING ALERT: ieatnetsimv5116-68;check netsim cpu;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1566728467] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 75.7%
[1566728928] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.63%
[1566729079] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.09%
[1566734240] SERVICE ALERT: ieatnetsimv5116-61;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566734360] SERVICE ALERT: ieatnetsimv5116-61;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566734481] SERVICE ALERT: ieatnetsimv5116-61;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566742870] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.28%
[1566743021] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.21%
[1566746500] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566746657] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566746814] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.68%
[1566750445] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 45.72%
[1566751008] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566751183] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566751323] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566751431] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566752017] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;OK;HARD;3;OK- CPU IS USED 34.09%
[1566753282] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566753450] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566753771] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566754005] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.35%
[1566754156] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.11%
[1566754306] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.26%
[1566754457] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.26%
[1566756991] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.14%
[1566757142] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.4%
[1566757572] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.55%
[1566757723] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.04%
[1566757873] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.71%
[1566758024] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.88%
[1566759240] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566759400] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566759503] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.51%
[1566759561] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566759654] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 69.36%
[1566764134] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566764319] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566764505] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566766370] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566768235] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566768421] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566768607] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566769546] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 45.18%
[1566769695] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.48%
[1566769846] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 12.4%
[1566770169] SERVICE ALERT: ieatnetsimv5116-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.4%
[1566770319] SERVICE ALERT: ieatnetsimv5116-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 64.69%
[1566770711] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 79.17%
[1566771692] SERVICE FLAPPING ALERT: ieatnetsimv5116-46;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566772607] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.12%
[1566772758] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.08%
[1566773969] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1647) IS LOWER THAN 80% OF ENM SNMP NODES (2500)