[1570231783] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1570231783] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (65) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570235387] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (61) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570237354] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (180) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570237774] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2710) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570237980] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570239148] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.26%
[1570239299] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 49.69%
[1570239450] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.23%
[1570239601] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 64.32%
[1570240536] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1570240780] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;SOFT;2;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1570242234] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570242864] HOST ALERT: ieatnetsimv7035-64;DOWN;SOFT;1;PING NOK
[1570242926] HOST ALERT: ieatnetsimv7035-64;UP;SOFT;1;PING OK
[1570258304] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1570258432] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1570258560] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1570259777] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570260368] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1570260593] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570262106] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1570262176] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1570262274] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570262303] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1570262441] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1570262461] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570262648] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570262690] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570265943] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1570266307] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1570266637] SERVICE FLAPPING ALERT: ieatlms5735-1;check nas hc;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1570266637] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1570266984] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (1280) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570267110] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (1280) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570267237] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1280) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570267797] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (56) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570267921] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570268018] SERVICE ALERT: ieatnetsimv017-17;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.37%
[1570268046] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570268169] SERVICE ALERT: ieatnetsimv017-17;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 3.37%
[1570269895] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (75) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570269915] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570270020] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (76) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570270145] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (106) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570271343] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2738) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570272181] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570272744] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570273405] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570273749] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (120) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570274171] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570276849] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1570276977] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1570277104] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1570277356] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (110) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570277480] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (111) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570277604] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1570277604] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (111) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570278912] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1570278912] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1570283204] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (19327) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570284808] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570288413] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (112) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570291431] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.46%
[1570291582] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 60.88%
[1570295619] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (123) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570299223] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570301215] SERVICE ALERT: ieatnetsimv5116-40;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.74%
[1570301365] SERVICE ALERT: ieatnetsimv5116-40;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.68%
[1570313120] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570314186] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3999) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570314987] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)