[1565305309] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1565305818] SERVICE ALERT: ieatnetsimv7035-107;check netsim disk;OK;HARD;3;OK
[1565306440] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.82%
[1565306593] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 74.99%
[1565306914] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565307172] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1565307183] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565307308] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17286) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21430)
[1565307355] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1565307537] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1565307801] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1565307980] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1565308162] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1565308736] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (23003) IS HIGHER THAN 95% OF ENM NODES (21497)
[1565309400] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1565309534] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.23%
[1565309582] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1565309638] HOST ALERT: ieatnetsimv5116-61;DOWN;SOFT;1;PING NOK
[1565309669] HOST ALERT: ieatnetsimv5116-61;UP;SOFT;1;PING OK
[1565309685] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.99%
[1565309838] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.28%
[1565309989] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 74.4%
[1565310174] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2346) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565310184] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (23003) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1565310305] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2346) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565310445] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2346) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565310452] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH CM SUPERVISION (23003) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1565310641] HOST ALERT: ieatnetsimv5116-69;DOWN;SOFT;1;PING NOK
[1565310703] HOST ALERT: ieatnetsimv5116-69;UP;SOFT;1;PING OK
[1565310708] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (23003) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1565311199] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1565311199] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21430)
[1565311358] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565311788] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565312742] SERVICE ALERT: ieatnetsimv7035-102;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565313347] HOST ALERT: ieatnetsimv5116-104;DOWN;SOFT;1;PING NOK
[1565313409] HOST ALERT: ieatnetsimv5116-104;UP;SOFT;1;PING OK
[1565314950] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17166) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21430)
[1565315358] SERVICE FLAPPING ALERT: ieatnetsimv5116-17;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565316705] SERVICE FLAPPING ALERT: ieatnetsimv7035-34;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565316966] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1565317150] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1565319133] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1565319315] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1565319988] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.56%
[1565319995] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565320115] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565320138] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.23%
[1565320236] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565322076] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (6643) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1565322630] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.32%
[1565322632] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565322781] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.35%
[1565322846] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565322932] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.66%
[1565323068] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565323888] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.31%
[1565324039] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.5%
[1565324189] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.59%
[1565324340] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.57%
[1565324942] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565325037] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.73%
[1565325188] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.95%
[1565325339] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;HARD;3;OK- CPU IS USED 47.07%
[1565325823] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (16508) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21426)
[1565326405] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1565326405] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (5370) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1565326846] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565327055] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565327186] SERVICE ALERT: ieatnetsimv7035-11;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565327259] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565327306] SERVICE ALERT: ieatnetsimv7035-11;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565327427] SERVICE ALERT: ieatnetsimv7035-11;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565328031] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565328033] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565328034] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565328035] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565328045] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565328051] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565328139] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.15%
[1565328185] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1565328185] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1565328185] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1565328188] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1565328196] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.8%
[1565328202] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.54%
[1565328290] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.56%
[1565328335] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;OK;HARD;3;OK- CPU IS USED 14.37%
[1565328336] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 71.33%
[1565328336] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1565328339] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 60.87%
[1565328346] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;HARD;3;OK- CPU IS USED 68.37%
[1565328353] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 73.83%
[1565328436] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1565328968] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13050) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565329101] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565329150] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565329242] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13050) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565329367] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565329583] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565329895] SERVICE FLAPPING ALERT: ieatlms3901;check hw resources hc;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1565329895] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1565329906] SERVICE ALERT: ieatnetsimv7035-26;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565330046] SERVICE ALERT: ieatnetsimv7035-26;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565330228] SERVICE ALERT: ieatnetsimv7035-26;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565331267] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.9%
[1565331418] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1565331569] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.98%
[1565331769] SERVICE ALERT: ieatwlvm5116;check workload profiles started;OK;HARD;3;OK- ALL WORKLOAD PROFILES ARE RUNNING
[1565331967] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.03%
[1565332057] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1990) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1565333109] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17683) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21426)
[1565333608] SERVICE ALERT: ieatnetsimv7035-65;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565333743] SERVICE ALERT: ieatnetsimv7035-65;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565333877] SERVICE ALERT: ieatnetsimv7035-65;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565333902] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1565334089] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1565334256] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2443) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1565334277] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1565334277] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1565334628] SERVICE ALERT: ieatnetsimv7035-11;check netsim disk;OK;HARD;3;OK
[1565334637] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;OK;HARD;3;OK
[1565334764] SERVICE ALERT: ieatnetsimv7035-65;check netsim disk;OK;HARD;3;OK
[1565334838] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565334864] SERVICE ALERT: ieatnetsimv7035-109;check netsim disk;OK;HARD;3;OK
[1565334958] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565335048] SERVICE ALERT: ieatnetsimv7035-83;check netsim disk;OK;HARD;3;OK
[1565335079] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565335091] SERVICE ALERT: ieatnetsimv7035-102;check netsim disk;OK;HARD;3;OK
[1565335199] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;OK;HARD;3;OK- CPU IS USED 31.98%
[1565336464] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13872) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565337040] HOST ALERT: ieatnetsimv5116-78;DOWN;SOFT;1;PING NOK
[1565337102] HOST ALERT: ieatnetsimv5116-78;UP;SOFT;1;PING OK
[1565337260] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4080) IS HIGHER THAN 95% OF ENM CPP NODES (4255)
[1565337610] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565338439] SERVICE ALERT: ieatnetsimv5116-49;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565338439] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565338559] SERVICE ALERT: ieatnetsimv5116-49;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565338559] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565338680] SERVICE ALERT: ieatnetsimv5116-49;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565338680] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565339180] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565339330] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 19.93%
[1565339599] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565339982] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565340655] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565341545] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1565341575] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1565341593] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (6492) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1565342316] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565342495] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565342793] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565342952] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1565343042] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565343110] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1565345464] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (150) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1565345638] SERVICE ALERT: ieatnetsimv5116-32;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565345638] HOST ALERT: ieatnetsimv5116-32;DOWN;SOFT;1;PING NOK
[1565345700] HOST ALERT: ieatnetsimv5116-32;UP;SOFT;1;PING OK
[1565345758] SERVICE ALERT: ieatnetsimv5116-32;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565345879] SERVICE ALERT: ieatnetsimv5116-32;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565347601] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (17088) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21426)
[1565348271] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2096) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565348663] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2334) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1565348747] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.65%
[1565348790] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2336) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1565348898] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.69%
[1565348916] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2336) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1565349059] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565349242] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5759) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1565349413] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565349427] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (7063) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1565349710] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565350289] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565350373] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED %
[1565350449] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1565350610] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1565351763] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.59%
[1565351913] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 70.95%
[1565352064] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.33%
[1565352215] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 60.77%
[1565352376] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565352524] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1565353296] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1565353447] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (6642) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1565354216] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565354474] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565354604] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565354734] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565355375] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565355484] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2400) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1565355693] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565356451] SERVICE ALERT: ieatnetsimv5116-112;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565356453] SERVICE ALERT: ieatnetsimv5116-113;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565356453] SERVICE ALERT: ieatnetsimv5116-114;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565357020] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565357357] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565357394] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565357547] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1565357699] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1565357847] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565358165] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565360433] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565360446] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565360489] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1565360584] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1565360626] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565360670] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1565360735] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.78%
[1565360807] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565360851] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1565362504] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1565362834] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1565363166] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.76%
[1565363192] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1565363316] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.1%
[1565363748] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565366958] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1565366958] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2497) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1565366979] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565367067] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.33%
[1565367218] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.99%
[1565367246] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565367368] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.66%
[1565367519] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 75.16%
[1565367776] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565368302] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565369474] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565369885] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565370286] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565370412] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1565370624] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1565370648] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.57%
[1565370798] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.52%
[1565370810] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1565370949] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.26%
[1565370995] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1565371100] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.25%
[1565371641] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.47%
[1565371775] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.88%
[1565371791] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 73.69%
[1565371926] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.75%
[1565372076] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;HARD;3;OK- CPU IS USED 62.39%
[1565374305] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.12%
[1565374438] HOST ALERT: ieatnetsimv5116-79;DOWN;SOFT;1;PING NOK
[1565374456] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.03%
[1565374500] HOST ALERT: ieatnetsimv5116-79;UP;SOFT;1;PING OK
[1565374607] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.99%
[1565377393] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1565377574] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1565377660] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565377756] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1565378944] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565379125] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565379306] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565380260] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565380428] SERVICE FLAPPING ALERT: ieatlms3901;check hw resources hc;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1565381180] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1565381837] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;OK;HARD;3;OK- CPU IS USED 56.02%
[1565382070] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565382100] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565382221] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.94%
[1565383929] SERVICE FLAPPING ALERT: ieatnetsimv5116-70;check netsim nodes;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1565383929] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565384125] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565384689] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565385062] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565386529] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565386533] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11529) IS FROM 80% TO 95% OF ENM COMECIM NODES (12141)
[1565386734] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11529) IS FROM 80% TO 95% OF ENM COMECIM NODES (12141)
[1565386807] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565386888] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11529) IS FROM 80% TO 95% OF ENM COMECIM NODES (12141)
[1565387100] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565387116] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565387372] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565387393] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565387430] HOST ALERT: ieatnetsimv5116-89;DOWN;SOFT;1;PING NOK
[1565387430] HOST ALERT: ieatnetsimv5116-91;DOWN;SOFT;1;PING NOK
[1565387449] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1565387492] HOST ALERT: ieatnetsimv5116-89;UP;SOFT;1;PING OK
[1565387492] HOST ALERT: ieatnetsimv5116-91;UP;SOFT;1;PING OK
[1565388369] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565390129] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.8%