[1565393542] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.27%
[1565393692] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.06%
[1565396268] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.12%
[1565396418] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.15%
[1565396571] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.82%
[1565396722] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.58%
[1565403038] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565403159] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565403279] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565404096] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2351) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565404233] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2335) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565404365] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2320) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565405403] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1565405986] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565406177] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2400) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1565406335] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565407571] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.63%
[1565407722] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.47%
[1565407872] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.57%
[1565408023] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.43%
[1565410197] HOST ALERT: ieatnetsimv7035-04;DOWN;SOFT;1;PING NOK
[1565410259] HOST ALERT: ieatnetsimv7035-04;UP;SOFT;1;PING OK
[1565418297] HOST ALERT: ieatnetsimv7035-69;DOWN;SOFT;1;PING NOK
[1565418359] HOST ALERT: ieatnetsimv7035-69;UP;SOFT;1;PING OK
[1565421690] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.27%
[1565421841] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.34%
[1565422039] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565422442] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565423369] SERVICE FLAPPING ALERT: ieatnetsimv5116-70;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565423683] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2911) IS LOWER THAN 80% OF ENM CPP NODES (4255)
[1565423818] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4071) IS HIGHER THAN 95% OF ENM CPP NODES (4255)
[1565423868] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1565423912] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1565424572] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1565427048] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565429192] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.03%
[1565429343] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.96%
[1565429494] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.67%
[1565429644] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.63%
[1565432021] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565432309] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565433394] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.06%
[1565433545] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.72%
[1565433695] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.49%
[1565434672] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1565435032] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1565435365] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1565435638] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565435638] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565435758] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565435759] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565435879] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565435879] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565436014] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6362) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1565436157] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6127) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1565436300] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6014) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1565437967] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (20342) IS LOWER THAN 80% OF ENM NODES (25514)
[1565438121] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (4308) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1565439080] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565439400] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565441680] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565441952] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3707) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1565442090] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3710) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1565442223] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3710) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1565442601] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1565442838] HOST ALERT: ieatnetsimv5116-11;DOWN;SOFT;1;PING NOK
[1565442900] HOST ALERT: ieatnetsimv5116-11;UP;SOFT;1;PING OK
[1565443015] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565444170] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.38%
[1565444321] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.64%
[1565444526] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.42%
[1565444677] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.66%
[1565448428] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.24%
[1565448578] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.86%
[1565448729] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.93%
[1565448879] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.98%
[1565449215] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565450135] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.47 seconds)
[1565451055] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565453058] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565453879] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;OK;HARD;3;OK
[1565453879] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;OK;HARD;3;OK
[1565453880] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;OK;HARD;3;OK
[1565453880] SERVICE ALERT: ieatnetsimv5116-38;check netsim disk;OK;HARD;3;OK
[1565453880] SERVICE ALERT: ieatnetsimv5116-61;check netsim disk;OK;HARD;3;OK
[1565453880] SERVICE ALERT: ieatnetsimv5116-32;check netsim disk;OK;HARD;3;OK
[1565453882] SERVICE ALERT: ieatnetsimv5116-49;check netsim disk;OK;HARD;3;OK
[1565453882] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;OK;HARD;3;OK
[1565454539] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565454967] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1565455217] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1565455453] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1565455631] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1565455659] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1565455809] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1565455889] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1565456074] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1565456262] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1565456580] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1565457175] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1565457500] SERVICE FLAPPING ALERT: ieatnetsimv5116-70;check netsim nodes;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1565457500] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565457758] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1565457940] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1565458041] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12426) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1565458122] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1565458187] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13101) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1565458335] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13162) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1565458580] SERVICE ALERT: ieatnetsimv5116-73;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1565458849] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (20751) IS BETWEEN 80% AND 95% OF ENM NODES (25514)
[1565459313] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17903) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21422)
[1565460191] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2025) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565460324] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2097) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565460456] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2172) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565461208] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565461383] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565462269] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2400) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1565463360] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6435) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1565463837] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2442) IS LOWER THAN 80% OF ENM CPP NODES (4255)
[1565463984] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565464863] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1565464903] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565464990] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1565465118] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1565465384] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1565467034] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.57%
[1565467184] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.75%
[1565467452] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3494) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1565467683] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2352) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565467814] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2400) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1565468850] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565469313] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565469442] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.65%
[1565469593] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.01%
[1565470513] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565470645] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565470774] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565470844] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565470935] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.6%
[1565471053] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565471085] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.81%
[1565471236] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 91.66%
[1565471248] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565472459] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (6961) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1565473154] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565473347] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2165) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565473368] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565473479] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2239) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565473591] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565473613] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2298) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565473912] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1565474866] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.83%
[1565475424] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1565475424] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2400) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1565476640] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;HARD;3;OK- CPU IS USED 63.99%
[1565477645] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565477780] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565477915] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES