[1569279802] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized cpp;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1569279802] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1569286394] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1569286582] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1569286769] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1569287307] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1569287436] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1569287565] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1569288665] SERVICE FLAPPING ALERT: ieatnetsimv5116-96;check netsim disk;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569288972] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1569288972] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1569289892] SERVICE FLAPPING ALERT: ieatlms3901;check unsync nodes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569291572] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569291998] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1569292918] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569292949] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM CPP NODES (0)
[1569293537] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1569294457] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569295377] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1569296297] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1569297349] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569297964] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM CPP NODES (0)
[1569298799] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1569298884] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569299613] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM CPP NODES (0)
[1569299901] SERVICE ALERT: ieatnetsimv7035-96;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1569300039] SERVICE ALERT: ieatnetsimv7035-96;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1569300178] SERVICE ALERT: ieatnetsimv7035-96;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1569301822] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM COMECIM NODES (0)
[1569301974] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1569303036] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1569305947] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569306665] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;OK;HARD;3;OK
[1569307733] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569308022] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569308599] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569308631] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1569308653] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1569308816] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1569308942] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1569309001] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1569309048] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1569309573] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;(Service check timed out after 800.00 seconds)
[1569309845] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM COMECIM NODES (0)
[1569309968] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569310265] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1569310386] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1569310506] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1569310889] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1569311416] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1569311466] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1569311543] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1569311671] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1569312108] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1569312295] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1569312445] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569312482] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1569313365] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1569314285] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1569314650] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1569314737] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (1287) IS HIGHER THAN 95% OF ENM NODES (96)
[1569314931] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (20364) IS BETWEEN 80% AND 95% OF ENM NODES (24450)
[1569315191] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (19614) IS BETWEEN 80% AND 95% OF ENM NODES (24450)
[1569315279] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1569315742] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (22812) IS BETWEEN 80% AND 95% OF ENM NODES (24450)
[1569316041] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH CM SUPERVISION (22812) IS BETWEEN 80% AND 95% OF ENM NODES (24450)
[1569316125] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (14058) IS HIGHER THAN 95% OF ENM COMECIM NODES (14059)
[1569316154] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1569316154] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1569316331] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (22812) IS BETWEEN 80% AND 95% OF ENM NODES (24450)
[1569316489] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6585) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569317346] SERVICE ALERT: ieatlms3901;check number nodes;OK;HARD;3;OK- NUMBER OF NODES CONFIGURED IN ENM (24450) IS HIGHER THAN 95% OF NETSIM NODES (17154)
[1569318252] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1569318387] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1569318521] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1569318887] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1569319014] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1569319142] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1569322931] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1569323066] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569323192] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1569323432] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 801.02 seconds)
[1569323848] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1569323982] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (20364) IS BETWEEN 80% AND 95% OF ENM NODES (24450)
[1569324190] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (19614) IS BETWEEN 80% AND 95% OF ENM NODES (24450)
[1569324344] SERVICE ALERT: ieatwlvm5116;check workload profiles started;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1569324464] SERVICE FLAPPING ALERT: ieatlms3901;check cm supervision nodes;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1569324464] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (22812) IS BETWEEN 80% AND 95% OF ENM NODES (24450)
[1569324472] SERVICE ALERT: ieatwlvm5116;check workload profiles started;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1569324489] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569324534] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569324603] SERVICE ALERT: ieatwlvm5116;check workload profiles started;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1569325193] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6435) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569325229] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1569325413] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (1429) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1569325416] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1569325539] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (1425) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1569325601] SERVICE ALERT: ieatnetsimv7035-96;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1569325604] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1569325665] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1425) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1569326576] SERVICE ALERT: ieatnetsimv7035-96;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1569326697] SERVICE ALERT: ieatnetsimv7035-96;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1569326817] SERVICE ALERT: ieatnetsimv7035-96;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1569327421] SERVICE ALERT: ieatnetsimv7035-96;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1569327555] SERVICE ALERT: ieatnetsimv7035-96;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1569327693] SERVICE ALERT: ieatnetsimv7035-96;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1569328754] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4174) IS FROM 80% TO 95% OF ENM CPP NODES (4446)
[1569328887] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4174) IS FROM 80% TO 95% OF ENM CPP NODES (4440)
[1569329021] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4174) IS FROM 80% TO 95% OF ENM CPP NODES (4433)
[1569329357] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569329838] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569330569] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6585) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569330834] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4174) IS HIGHER THAN 95% OF ENM CPP NODES (4354)
[1569331200] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized cpp;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1569334384] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1569334435] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2736) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569334513] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1569334643] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1569334796] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569335130] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6585) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569336247] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4174) IS FROM 80% TO 95% OF ENM CPP NODES (4416)
[1569336382] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4174) IS FROM 80% TO 95% OF ENM CPP NODES (4410)
[1569336517] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4174) IS FROM 80% TO 95% OF ENM CPP NODES (4404)
[1569338331] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized cpp;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1569338331] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4174) IS HIGHER THAN 95% OF ENM CPP NODES (4378)
[1569341664] SERVICE FLAPPING ALERT: ieatlms3901;check cm supervision nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569345453] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1569347262] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1569347393] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1569347523] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1569352671] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1569365325] SERVICE ALERT: ieatnetsimv7035-99;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1569365462] SERVICE ALERT: ieatnetsimv7035-99;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1569365601] SERVICE ALERT: ieatnetsimv7035-99;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES