[1575505170] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1575505349] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1575505527] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1575505563] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575505626] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;OK;HARD;3;OK
[1575505699] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;SOFT;2;OK
[1575506007] SERVICE ALERT: ieatwlvm9034;check workload profiles started;OK;HARD;3;OK- ALL WORKLOAD PROFILES ARE RUNNING
[1575506283] SERVICE ALERT: ieatnetsimv7102-05;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.25%
[1575506434] SERVICE ALERT: ieatnetsimv7102-05;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.44%
[1575506549] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1575507499] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures RBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575508357] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1575508485] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1575508613] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1575509150] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures SBG-IS;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1575509229] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1575509229] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1575509501] SERVICE ALERT: ieatnetsimv7102-19;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.47%
[1575509651] SERVICE ALERT: ieatnetsimv7102-19;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 50.62%
[1575510106] SERVICE FLAPPING ALERT: ieatlms7435;check synchronized RadioNode;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575512174] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6274;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1575512825] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575513015] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575513204] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575513318] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1575513401] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575513413] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical DSC;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1575513508] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1575513697] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1575514049] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical EPG;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575515074] SERVICE ALERT: ieatlms7435;check node fs;OK;HARD;3;OK
[1575515567] SERVICE ALERT: ieatlms7435;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1575515771] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575516114] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-IS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575519950] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ESC;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575520183] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6672;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575520904] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures snmp;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575521156] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MTAS;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1575521192] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-BSP;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1575521356] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1575524144] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575524333] SERVICE ALERT: ieatlms7435;check node fs;OK;SOFT;2;OK
[1575524408] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1575524523] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575524585] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-Indoor;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575524637] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1575524712] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575524827] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1575524902] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575524985] HOST ALERT: ieatnetsimv5116-66;DOWN;SOFT;1;PING NOK
[1575525016] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1575525047] HOST ALERT: ieatnetsimv5116-66;UP;SOFT;1;PING OK
[1575525218] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures EPG;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575526290] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures ESC;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1575526428] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RBS;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1575526772] SERVICE FLAPPING ALERT: ieatlms7435;check node fs;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1575526772] SERVICE ALERT: ieatlms7435;check node fs;OK;HARD;3;OK
[1575526887] SERVICE FLAPPING ALERT: ieatlms7435;check node fs hc;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1575526887] SERVICE ALERT: ieatlms7435;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1575527025] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RNC;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1575529058] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575529082] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575529247] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures cpp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575529789] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1575530175] SERVICE ALERT: ieatnetsimv9034-64;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575530316] SERVICE ALERT: ieatnetsimv9034-64;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575530620] SERVICE ALERT: ieatwlvm9034;check workload profiles started;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1575530753] SERVICE ALERT: ieatwlvm9034;check workload profiles started;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1575530885] SERVICE ALERT: ieatwlvm9034;check workload profiles started;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1575531172] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575531280] SERVICE ALERT: ieatnetsimv7035-58;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575531416] SERVICE ALERT: ieatnetsimv7035-58;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575531642] SERVICE ALERT: ieatnetsimv7035-56;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575531787] SERVICE ALERT: ieatnetsimv7035-56;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575532082] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575532239] SERVICE ALERT: ieatnetsimv7035-84;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575532379] SERVICE ALERT: ieatnetsimv7035-84;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575532835] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575532945] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MSC-DB-BSP;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575532971] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575533108] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575533374] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures BSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575533601] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575533634] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575533914] SERVICE ALERT: ieatnetsimv7035-82;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575534042] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575534065] SERVICE ALERT: ieatnetsimv7035-82;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575534158] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1575534211] SERVICE ALERT: ieatnetsimv7035-82;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575534231] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575534347] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1575534369] SERVICE ALERT: ieatnetsimv7035-82;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575534421] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575534537] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1575534925] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;HARD;3;OK
[1575536433] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1575536741] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575536877] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;SOFT;2;OK
[1575537098] SERVICE ALERT: ieatlms5735-1;check fm hb failures ESC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (80) IS BETWEEN 5% AND 20% OF ENM NODES (1000)
[1575537232] SERVICE ALERT: ieatlms5735-1;check fm hb failures ESC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (80) IS BETWEEN 5% AND 20% OF ENM NODES (1000)
[1575537357] SERVICE ALERT: ieatlms5735-1;check fm hb failures ESC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (80) IS BETWEEN 5% AND 20% OF ENM NODES (1000)
[1575537392] SERVICE ALERT: ieatnetsimv9041-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575537423] SERVICE ALERT: ieatlms7435;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1575537548] SERVICE ALERT: ieatnetsimv9041-11;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1575537662] SERVICE ALERT: ieatnetsimv7102-01;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.98%
[1575537710] SERVICE ALERT: ieatnetsimv9041-11;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1575537813] SERVICE ALERT: ieatnetsimv7102-01;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.95%
[1575537830] SERVICE ALERT: ieatnetsimv7102-04;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.95%
[1575537875] SERVICE ALERT: ieatlms7435;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1575537964] SERVICE ALERT: ieatnetsimv7102-01;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 92.31%
[1575537981] SERVICE ALERT: ieatnetsimv7102-04;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.11%
[1575538131] SERVICE ALERT: ieatnetsimv7102-04;check netsim cpu;OK;HARD;3;OK- CPU IS USED 59.27%
[1575538294] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575538325] SERVICE ALERT: ieatlms7435;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1575538401] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2333) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1575538473] SERVICE FLAPPING ALERT: ieatlms7435;check node fs;STARTED; Service appears to have started flapping (25.7% change >= 20.0% threshold)
[1575538473] SERVICE FLAPPING ALERT: ieatlms7435;check node fs hc;STARTED; Service appears to have started flapping (26.0% change >= 20.0% threshold)
[1575538480] SERVICE ALERT: ieatlms3901;check list nodes - DEMO;CRITICAL;HARD;3;28
[1575538529] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2305) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1575538557] SERVICE FLAPPING ALERT: ieatlms7435;check node fs;STARTED; Service appears to have started flapping (25.7% change >= 20.0% threshold)
[1575538557] SERVICE FLAPPING ALERT: ieatlms7435;check node fs hc;STARTED; Service appears to have started flapping (26.0% change >= 20.0% threshold)
[1575538620] SERVICE FLAPPING ALERT: ieatlms7435;check node fs;STARTED; Service appears to have started flapping (25.7% change >= 20.0% threshold)
[1575538620] SERVICE FLAPPING ALERT: ieatlms7435;check node fs hc;STARTED; Service appears to have started flapping (26.0% change >= 20.0% threshold)
[1575538655] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2302) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1575538686] SERVICE FLAPPING ALERT: ieatlms7435;check node fs;STARTED; Service appears to have started flapping (25.7% change >= 20.0% threshold)
[1575538686] SERVICE FLAPPING ALERT: ieatlms7435;check node fs hc;STARTED; Service appears to have started flapping (26.0% change >= 20.0% threshold)
[1575538785] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.41%
[1575538785] SERVICE ALERT: ieatlms5735-1;check node fs;OK;SOFT;2;OK
[1575538935] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.21%
[1575539198] SERVICE ALERT: ieatnetsimv7102-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.97%
[1575539205] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1575539336] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (150)
[1575539349] SERVICE ALERT: ieatnetsimv7102-59;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 1%
[1575539465] SERVICE ALERT: ieatlms3901;check list nodes - DEMO;WARNING;HARD;3;WARNING - some  type nodes in management server (141.137.208.23) are different (1)
[1575539466] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN CISCO-ASR900 NODES !
[1575539535] SERVICE ALERT: ieatnetsimv9041-11;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1575539592] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (150)
[1575539613] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1575539733] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN CISCO-ASR900 NODES !
[1575539836] SERVICE ALERT: ieatnetsimv9041-11;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.91%
[1575540931] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13139) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1575540963] SERVICE ALERT: ieatlms5735-1;check fm hb failures ESC;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN ESC NODES !
[1575541080] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13687) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1575541368] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (2540) IS HIGHER THAN 95% OF ENM NODES (2560)
[1575541595] SERVICE ALERT: ieatnetsimv7102-01;check netsim cpu;OK;HARD;3;OK- CPU IS USED 55.02%
[1575541690] SERVICE ALERT: ieatlms7435;check node fs;OK;HARD;3;OK
[1575541807] SERVICE ALERT: ieatlms7435;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1575542085] HOST ALERT: ieatnetsimv5116-69;DOWN;SOFT;1;PING NOK
[1575542147] HOST ALERT: ieatnetsimv5116-69;UP;SOFT;1;PING OK
[1575542278] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2385) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1575542283] SERVICE ALERT: ieatlms7435;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1575542409] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2373) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1575542544] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (2470) IS HIGHER THAN 95% OF ENM NODES (2560)
[1575542635] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (51)
[1575542686] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.73%
[1575542763] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (51)
[1575542838] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575542866] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1575542887] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (51)
[1575542990] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575543047] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1575543227] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1575544211] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (275) IS HIGHER THAN 20% OF ENM NODES (1201)
[1575544366] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575544371] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (275) IS HIGHER THAN 20% OF ENM NODES (1201)
[1575544495] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (276) IS HIGHER THAN 20% OF ENM NODES (1201)
[1575544547] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575544727] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575544853] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13299) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1575544994] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13796) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1575546497] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1575546589] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1575546621] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;OK;HARD;3;OK- CPU IS USED 43.17%
[1575546888] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1575548085] HOST ALERT: ieatnetsimv017-10;DOWN;SOFT;1;PING NOK
[1575548100] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (109) IS BETWEEN 5% AND 20% OF ENM NODES (1201)
[1575548147] HOST ALERT: ieatnetsimv017-10;UP;SOFT;1;PING OK
[1575548449] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575548630] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575548749] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1575548811] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575548811] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575548930] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1575548959] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575549077] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1575549111] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1575549111] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1575549149] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575549267] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1575549339] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575549456] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1575550022] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1575550104] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (51)
[1575550229] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (51)
[1575551840] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (6235) IS FROM 80% TO 95% OF ENM  NODES (7500)
[1575551966] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (5485) IS LOWER THAN 80% OF ENM  NODES (7500)
[1575552091] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (4513) IS LOWER THAN 80% OF ENM  NODES (7500)
[1575552437] SERVICE ALERT: ieatnetsimv7102-82;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.37%
[1575552514] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575552588] SERVICE ALERT: ieatnetsimv7102-82;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.75%
[1575552642] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1575552741] SERVICE ALERT: ieatnetsimv7102-82;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575553443] SERVICE ALERT: ieatlms7435;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1575553645] HOST ALERT: ieatnetsimv9034-97;DOWN;SOFT;1;PING NOK
[1575553707] HOST ALERT: ieatnetsimv9034-97;UP;SOFT;1;PING OK
[1575553854] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575553867] SERVICE ALERT: ieatlms3901;check list nodes - DEMO;CRITICAL;HARD;3;CRITICAL - some  type nodes in management server (141.137.208.23) are different (27)
[1575553892] SERVICE ALERT: ieatlms7435;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1575553947] SERVICE ALERT: ieatnetsimv7102-39;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575554005] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575554098] SERVICE ALERT: ieatnetsimv7102-39;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575554157] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575554161] SERVICE ALERT: ieatnetsimv7102-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.51%
[1575554194] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1575554250] SERVICE ALERT: ieatnetsimv7102-39;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.78%
[1575554311] SERVICE ALERT: ieatlms7435;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1575554312] SERVICE ALERT: ieatnetsimv7102-08;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.51%
[1575554478] SERVICE ALERT: ieatnetsimv7102-08;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.05%
[1575554552] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1575554911] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1575555322] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (20) IS LOWER THAN 5% OF ENM NODES (1201)
[1575555697] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (7462) IS HIGHER THAN 95% OF ENM  NODES (7500)
[1575556373] SERVICE ALERT: ieatnetsimv7102-82;check netsim cpu;OK;HARD;3;OK- CPU IS USED 74.13%
[1575557167] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575557787] SERVICE FLAPPING ALERT: ieatnetsimv7102-07;check netsim cpu;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1575557787] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.64%
[1575557880] SERVICE ALERT: ieatnetsimv7102-39;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.93%
[1575558112] SERVICE ALERT: ieatnetsimv7102-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 76.64%
[1575558939] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (278) IS HIGHER THAN 20% OF ENM NODES (1201)
[1575559032] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575559072] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (202) IS BETWEEN 5% AND 20% OF ENM NODES (1201)
[1575559208] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1575559208] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (226) IS BETWEEN 5% AND 20% OF ENM NODES (1201)
[1575559217] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575559403] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575561154] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (51)
[1575561285] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (51)
[1575561409] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (51)
[1575561743] SERVICE ALERT: ieatnetsimv7102-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.16%
[1575561894] SERVICE ALERT: ieatnetsimv7102-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.76%
[1575562630] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1575562758] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1575562817] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (416) IS HIGHER THAN 20% OF ENM NODES (1201)
[1575562886] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1575564584] HOST ALERT: ieatnetsimv5116-96;DOWN;SOFT;1;PING NOK
[1575564646] HOST ALERT: ieatnetsimv5116-96;UP;SOFT;1;PING OK
[1575565879] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575566031] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575566423] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (199) IS BETWEEN 5% AND 20% OF ENM NODES (1201)
[1575566826] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1575566947] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1575567067] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1575568615] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1575568615] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1575570031] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (379) IS HIGHER THAN 20% OF ENM NODES (1201)
[1575572219] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (18) IS HIGHER THAN 20% OF ENM NODES (51)
[1575572348] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (17) IS HIGHER THAN 20% OF ENM NODES (51)
[1575572474] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1575573641] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1575573694] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1575573869] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575578209] SERVICE FLAPPING ALERT: ieatlms7435;check node fs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575578325] SERVICE FLAPPING ALERT: ieatlms7435;check node fs hc;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575583072] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575583374] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)