[1564702050] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.74%
[1564702200] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 53.98%
[1564702660] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564703174] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.05%
[1564703327] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.04%
[1564703481] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.36%
[1564703580] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.46 seconds)
[1564704442] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564707111] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.25%
[1564708842] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564709670] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564710590] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564711510] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1564712430] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1564714342] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.14%
[1564714492] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 75.95%
[1564714644] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.23%
[1564714795] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564714946] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.32%
[1564715631] SERVICE ALERT: ieatnetsimv5116-40;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1564715751] SERVICE ALERT: ieatnetsimv5116-40;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1564715872] SERVICE ALERT: ieatnetsimv5116-40;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1564716750] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.51%
[1564716901] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 38.26%
[1564717052] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.23%
[1564717203] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 71.37%
[1564723896] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564728296] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564729152] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564731739] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1564732534] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1564737133] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (150) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564738052] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564739381] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (19205) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21429)
[1564739695] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (19211) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21429)
[1564739892] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (19211) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21429)
[1564744188] SERVICE ALERT: ieatnetsimv5116-112;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1564744316] SERVICE ALERT: ieatnetsimv5116-112;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1564744444] SERVICE ALERT: ieatnetsimv5116-112;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1564746173] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1564746199] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.38%
[1564746350] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.77%
[1564746501] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.52%
[1564746687] SERVICE ALERT: ieatnetsimv7035-23;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564746690] SERVICE ALERT: ieatnetsimv7035-20;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564746838] SERVICE ALERT: ieatnetsimv7035-23;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 65.36%
[1564746840] SERVICE ALERT: ieatnetsimv7035-20;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 44.38%
[1564747093] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1564747707] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (150) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564748914] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.93%
[1564748918] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.97%
[1564748940] SERVICE ALERT: ieatnetsimv5116-76;check netsim disk;OK;HARD;3;OK
[1564749064] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.78%
[1564749069] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.91%
[1564749075] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564749215] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.48%
[1564749219] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.89%
[1564749223] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.88%
[1564749226] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564749267] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564749292] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564749300] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564749330] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564749331] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564749373] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.69%
[1564749376] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564749422] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564749443] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564749451] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564749481] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564749482] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564749524] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.88%
[1564749536] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564749578] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564749594] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564749594] SERVICE ALERT: ieatnetsimv5116-113;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1564749602] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564749633] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564749633] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564749634] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564749689] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564749721] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564749724] SERVICE ALERT: ieatnetsimv5116-113;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1564749753] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564749785] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564749843] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564749854] SERVICE ALERT: ieatnetsimv5116-113;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1564749872] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564749907] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564749937] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564750024] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564750024] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564750065] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564750131] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;HARD;3;OK- CPU IS USED 22.17%
[1564750176] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564750327] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564750522] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13101) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564750663] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13101) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564750805] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13101) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564751160] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1564751396] SERVICE ALERT: ieatnetsimv5116-114;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1564751526] SERVICE ALERT: ieatnetsimv5116-114;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1564751655] SERVICE ALERT: ieatnetsimv5116-114;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1564751873] SERVICE ALERT: ieatnetsimv5116-40;check netsim disk;OK;HARD;3;OK
[1564751875] SERVICE ALERT: ieatnetsimv5116-41;check netsim disk;OK;HARD;3;OK
[1564751985] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564752845] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;OK;HARD;3;OK- CPU IS USED 29.67%
[1564752850] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 26.42%
[1564754815] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564755352] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1564755496] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1564755634] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1564756442] SERVICE ALERT: ieatnetsimv5116-45;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1564756563] SERVICE ALERT: ieatnetsimv5116-45;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1564756684] SERVICE ALERT: ieatnetsimv5116-45;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1564756806] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.4%
[1564756808] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 19.86%
[1564756879] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1564756957] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 38.64%
[1564757004] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1564757073] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.71%
[1564758073] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (278) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1564758973] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1564759126] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1564759280] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1564760207] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;OK;HARD;3;OK- CPU IS USED 1.88%
[1564760484] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1564760611] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1564760669] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1564760740] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1564760847] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1564760865] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1564760896] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 81.39%
[1564760913] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1564760962] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.34%
[1564761107] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1564761112] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.37%
[1564761301] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1564762021] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2535) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564763176] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1564763586] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1564764016] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.51%
[1564764166] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 63.09%
[1564764385] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564767633] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.12%
[1564767664] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 53.19%
[1564767968] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.36%
[1564768361] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;OK;HARD;3;OK- CPU IS USED 55.32%
[1564768398] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (179) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564769286] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564769310] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564769413] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2535) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564769598] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564769783] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564769924] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2535) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564771043] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1564771171] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1564771264] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.88%
[1564771299] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1564771650] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564771654] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.99%
[1564771835] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564772021] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564773282] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6385) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1564773428] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (7147) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1564774495] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564774684] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564774784] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1564774873] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564774964] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1564775143] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1564775405] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6334) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1564775672] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (7142) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1564777430] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564778351] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1564778425] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.14%
[1564779038] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4740) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1564779271] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1564779308] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4776) IS HIGHER THAN 95% OF ENM SNMP NODES (5017)
[1564779423] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (22907) IS BETWEEN 80% AND 95% OF ENM NODES (25520)
[1564779424] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM CPP NODES (6973)
[1564779428] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (3969) IS LOWER THAN 80% OF ENM NODES (25520)
[1564779577] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (7147) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1564779674] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (25494) IS HIGHER THAN 95% OF ENM NODES (25520)
[1564780365] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564780451] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23812) IS BETWEEN 80% AND 95% OF ENM NODES (25520)
[1564780549] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1564780732] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564782307] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.07%
[1564782409] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564782458] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 30.75%
[1564782564] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;SOFT;2;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564782595] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564782777] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1564782961] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564783049] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4566) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1564783201] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4604) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1564783342] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564783343] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4644) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1564783531] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564783661] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564783735] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564784389] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564784551] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2535) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564784736] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564784824] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564784888] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2535) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564785007] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564785079] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564785263] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564785457] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564785799] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1564785980] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;SOFT;2;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1564786357] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564786482] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;SOFT;2;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)