[1564009207] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564009300] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1564009688] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1564009688] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2414) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1564012651] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564012867] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1564012961] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1564013624] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1564013828] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1564013956] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1564014084] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1564014728] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564014822] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1564014909] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564015002] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1564015090] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564015183] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1564016950] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1564018844] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1564019888] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.55%
[1564020039] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.08%
[1564020190] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.88%
[1564020335] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1564020337] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1564020503] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2363) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1564020635] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2305) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1564020635] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1564020766] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2291) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1564023062] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.22%
[1564023212] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.94%
[1564023363] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 19.61%
[1564023821] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;HARD;3;OK- CPU IS USED 25.29%
[1564028152] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (4515) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1564028297] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (7149) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1564029022] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1564031993] HOST ALERT: ieatnetsimv7035-44;DOWN;SOFT;1;PING NOK
[1564032055] HOST ALERT: ieatnetsimv7035-44;UP;SOFT;1;PING OK
[1564032092] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1564032622] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1564032995] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564033183] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564033372] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564035240] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1564035919] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;OK;HARD;3;OK
[1564036287] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564036473] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564036658] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564037367] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1564037496] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1564037499] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1564037626] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1564037627] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1564037754] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1564038908] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564039097] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564039284] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564042305] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1564042377] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564042487] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;SOFT;2;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1564042841] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5955) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1564042962] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1564042962] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1564043036] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1564043041] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (7149) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1564043125] SERVICE ALERT: ieatnetsimv7035-102;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1564043162] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1564043162] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1564043243] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7151) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1564044163] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564044792] SERVICE ALERT: ieatnetsimv7035-34;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1564045083] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1564045113] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1564045723] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1564045917] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1564045977] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1564046770] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1564046897] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1564047025] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1564047086] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564047234] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14035) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564047588] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564047773] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564047959] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1564047959] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564049136] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1564049138] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564049867] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1564050317] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564051211] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1564051235] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1564051378] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1564051885] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1564052155] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564052651] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.38%
[1564052802] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 16.92%
[1564052953] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.71%
[1564053049] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1564053103] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 21.41%
[1564053969] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564054889] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1564055695] HOST ALERT: ieatnetsimv7035-05;DOWN;SOFT;1;PING NOK
[1564055757] HOST ALERT: ieatnetsimv7035-05;UP;SOFT;1;PING OK
[1564055809] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1564056033] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1564057025] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1564057794] HOST ALERT: ieatnetsimv7035-39;DOWN;SOFT;1;PING NOK
[1564057856] HOST ALERT: ieatnetsimv7035-39;UP;SOFT;1;PING OK
[1564058410] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564058591] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564058772] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564058868] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1564059048] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1564059229] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1564059561] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.43%
[1564059715] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564059867] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.73%
[1564060168] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1564060530] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1564060691] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564060876] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564061062] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564061440] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1564061569] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1564061697] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1564063499] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.93%
[1564064187] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1564064257] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical DSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1564064408] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1564064662] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1564064727] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1564064911] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564065025] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1564065081] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1564065097] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564065263] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1564065279] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;SOFT;2;OK- NUMBER OF PM FILES (1510) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564065541] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1564065639] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1564065723] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;SOFT;2;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1564065821] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564067129] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.34%
[1564067137] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1564067140] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (100) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1564067281] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.81%
[1564068366] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1564070744] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1564075431] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1564076015] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1564076345] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1564076674] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1564082568] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564083408] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1564084503] HOST ALERT: ieatnetsimv7035-10;DOWN;SOFT;1;PING NOK
[1564084565] HOST ALERT: ieatnetsimv7035-10;UP;SOFT;1;PING OK
[1564089254] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.65%
[1564089405] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 69.55%
[1564089556] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.68%
[1564089706] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.1%
[1564089857] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 93.98%
[1564092305] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1564093487] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;HARD;3;OK- CPU IS USED 64.08%
[1564093528] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1564094995] HOST ALERT: ieatnetsimv7035-70;DOWN;SOFT;1;PING NOK
[1564095057] HOST ALERT: ieatnetsimv7035-70;UP;SOFT;1;PING OK
[1564095128] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)