[1571614704] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1571627183] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571627910] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571631920] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2560) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571632634] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571633158] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3999) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571633574] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571633812] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (30) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571644206] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571644400] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571644594] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571644834] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12929) IS FROM 80% TO 95% OF ENM COMECIM NODES (14127)
[1571645007] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12718) IS FROM 80% TO 95% OF ENM COMECIM NODES (14127)
[1571645173] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12581) IS FROM 80% TO 95% OF ENM COMECIM NODES (14127)
[1571645349] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6551) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571645866] SERVICE ALERT: ieatnetsimv7035-68;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.69%
[1571645877] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1571645889] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.54%
[1571646016] SERVICE ALERT: ieatnetsimv7035-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 23.66%
[1571646028] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 20.54%
[1571646040] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 15.55%
[1571646348] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (92) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (100)
[1571646469] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571646475] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (92) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (100)
[1571646603] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (100) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (100)
[1571646951] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13366) IS FROM 80% TO 95% OF ENM COMECIM NODES (14092)
[1571647024] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13886) IS HIGHER THAN 95% OF ENM COMECIM NODES (14127)
[1571647087] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.32%
[1571647098] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13129) IS FROM 80% TO 95% OF ENM COMECIM NODES (14092)
[1571647238] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 51.13%
[1571647240] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13168) IS FROM 80% TO 95% OF ENM COMECIM NODES (14092)
[1571649077] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571649309] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (30) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571649527] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (150) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571652078] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571652208] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1571652338] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1571652677] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13453) IS HIGHER THAN 95% OF ENM COMECIM NODES (14092)
[1571653430] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.56%
[1571653581] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 16%
[1571653736] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6439) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571657486] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (180) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571657753] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1571657904] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2560) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571658330] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571660097] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571660287] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571660472] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571661616] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571661803] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571661990] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571663619] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571663842] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571663860] SERVICE ALERT: ieatlms3901;check rpm diff;CRITICAL;HARD;3;(No output on stdout) stderr: /usr/local/nagios/libexec/s4/diff_rpm_iso_installed.sh: line 15: summary: command not found
[1571663979] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571663993] SERVICE ALERT: ieatlms3901;check rpm diff;WARNING;HARD;3;----------------------------
[1571664137] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571664190] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571664330] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571664504] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571664740] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571664891] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571665071] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571665152] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571665481] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571665670] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571665918] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1571668222] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571669481] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2349) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571669616] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2360) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571669804] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2338) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571670033] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (80) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571670086] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (2724) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571670159] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (143) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571670217] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;SOFT;2;OK- NUMBER OF PM FILES (3722) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571670555] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6559) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571671621] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2470) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1571672266] SERVICE ALERT: ieatnetsimv7035-75;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571672394] SERVICE ALERT: ieatnetsimv7035-75;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571672527] SERVICE ALERT: ieatnetsimv7035-63;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571672655] SERVICE ALERT: ieatnetsimv7035-63;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571674393] SERVICE ALERT: ieatnetsimv5116-41;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.19%
[1571674544] SERVICE ALERT: ieatnetsimv5116-41;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.76%
[1571681609] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (180) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571682345] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571686640] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2590) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571687109] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571687861] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571691569] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571692315] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571693905] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1571694025] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1571694146] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%