[1562108945] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.67%
[1562109095] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.15%
[1562109246] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.43%
[1562109492] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.92%
[1562109645] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.7%
[1562109675] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.17%
[1562109799] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.97%
[1562109826] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.09%
[1562109980] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.19%
[1562111731] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.57%
[1562111882] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 65.97%
[1562112039] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.91%
[1562112190] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 72.88%
[1562112877] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;OK;HARD;3;OK- CPU IS USED 73.5%
[1562113430] SERVICE FLAPPING ALERT: ieatnetsimv7035-22;check netsim cpu;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1562113430] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 73.91%
[1562113611] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;HARD;3;OK- CPU IS USED 69.34%
[1562130331] HOST ALERT: ieatnetsimv7035-34;DOWN;SOFT;1;PING NOK
[1562130393] HOST ALERT: ieatnetsimv7035-34;UP;SOFT;1;PING OK
[1562134015] SERVICE ALERT: ieatnetsimv5116-19;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1562134166] SERVICE ALERT: ieatnetsimv5116-19;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 54.22%
[1562134864] SERVICE ALERT: ieatnetsimv7035-110;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562135294] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7109)
[1562136059] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7109)
[1562136160] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.85%
[1562136248] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7109)
[1562136311] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 58.77%
[1562136462] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562136612] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.62%
[1562136636] SERVICE FLAPPING ALERT: ieatnetsimv7035-19;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562137110] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7109)
[1562138479] SERVICE ALERT: ieatnetsimv7035-110;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562138608] SERVICE ALERT: ieatnetsimv7035-110;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562140364] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.87%
[1562140515] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.48%
[1562140629] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17125) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (20185)
[1562140949] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (179) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7109)
[1562141870] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1562142353] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7109)
[1562143840] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562143994] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.19%
[1562144144] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.43%
[1562144314] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (17155) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21945)
[1562144740] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.97%
[1562144791] SERVICE ALERT: ieatnetsimv5116-57;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562144823] SERVICE ALERT: ieatnetsimv5116-26;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562144825] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562144869] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562144891] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 75.59%
[1562144943] SERVICE ALERT: ieatnetsimv5116-57;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562144975] SERVICE ALERT: ieatnetsimv5116-26;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562144977] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562145021] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562145094] SERVICE ALERT: ieatnetsimv5116-57;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562145125] SERVICE ALERT: ieatnetsimv5116-22;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562145126] SERVICE ALERT: ieatnetsimv5116-26;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562145127] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.39%
[1562145172] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562145177] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562145204] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562145276] SERVICE ALERT: ieatnetsimv5116-22;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562145331] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562145361] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562145427] SERVICE ALERT: ieatnetsimv5116-22;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562145484] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562145514] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562145860] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.68%
[1562146011] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.14%
[1562146164] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.57%
[1562146299] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562146484] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562146669] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562147776] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.21%
[1562147789] SERVICE ALERT: ieatnetsimv5116-29;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562147791] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562147940] SERVICE ALERT: ieatnetsimv5116-29;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562147941] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562148091] SERVICE ALERT: ieatnetsimv5116-29;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562148093] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562148151] SERVICE ALERT: ieatnetsimv5116-25;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562148151] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562148152] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562148152] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562148153] SERVICE ALERT: ieatnetsimv5116-58;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562148153] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562148154] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562148161] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562148178] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562148302] SERVICE ALERT: ieatnetsimv5116-25;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562148302] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562148303] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 49.46%
[1562148303] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562148304] SERVICE ALERT: ieatnetsimv5116-58;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562148304] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562148305] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562148314] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562148367] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562148454] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562148454] SERVICE ALERT: ieatnetsimv5116-25;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562148454] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562148455] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562148455] SERVICE ALERT: ieatnetsimv5116-58;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562148457] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562148466] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.39%
[1562148555] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562148724] SERVICE ALERT: ieatnetsimv5116-57;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.11%
[1562148802] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;OK;HARD;3;OK- CPU IS USED 20.81%
[1562148909] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.15%
[1562149058] SERVICE ALERT: ieatnetsimv5116-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 26.11%
[1562149062] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99%
[1562149212] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.4%
[1562149794] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.51%
[1562150344] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562151407] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.87%
[1562151561] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.49%
[1562151717] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562151722] SERVICE ALERT: ieatnetsimv5116-29;check netsim cpu;OK;HARD;3;OK- CPU IS USED 33.34%
[1562151723] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.12%
[1562152085] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.99%
[1562152086] SERVICE ALERT: ieatnetsimv5116-58;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.45%
[1562152086] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 17.89%
[1562152097] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;OK;HARD;3;OK- CPU IS USED 16.56%
[1562152243] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.77%
[1562152393] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.41%
[1562152544] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.22%
[1562152843] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.68%
[1562152993] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.05%
[1562153144] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 93.93%
[1562153424] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.34%
[1562153577] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.04%
[1562153729] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.87%
[1562155685] SERVICE ALERT: ieatnetsimv5116-25;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.24%
[1562155957] SERVICE ALERT: ieatnetsimv5116-26;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.82%
[1562156314] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.33%
[1562156775] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;OK;HARD;3;OK- CPU IS USED 78.04%
[1562157623] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562158228] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.08%
[1562158379] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.33%
[1562158531] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.6%
[1562159285] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;OK;HARD;3;OK- CPU IS USED 11.7%
[1562159287] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.37%
[1562159460] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1562159729] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;2;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (8308) IS LOWER THAN 80% OF ENM COMECIM NODES (11144)
[1562159775] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 67.1%
[1562159932] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1562159945] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 24.36%
[1562160027] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (6111) IS LOWER THAN 80% OF ENM COMECIM NODES (11144)
[1562160171] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (24328) IS HIGHER THAN 95% OF ENM NODES (24353)
[1562160406] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.39%
[1562160557] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.46%
[1562160709] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.35%
[1562160728] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (19870) IS BETWEEN 80% AND 95% OF ENM NODES (24353)
[1562161893] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (9575) IS FROM 80% TO 95% OF ENM COMECIM NODES (11144)
[1562162550] SERVICE FLAPPING ALERT: ieatnetsimv7035-19;check netsim cpu;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1562162550] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;OK;HARD;3;OK- CPU IS USED 81.67%
[1562164560] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.84%
[1562165528] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (10830) IS HIGHER THAN 95% OF ENM COMECIM NODES (11144)
[1562165762] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89%
[1562166928] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1562167358] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (7896) IS LOWER THAN 80% OF ENM COMECIM NODES (11144)
[1562167644] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;2;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (7756) IS LOWER THAN 80% OF ENM COMECIM NODES (11144)
[1562167908] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized comecim;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1562167908] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (7895) IS LOWER THAN 80% OF ENM COMECIM NODES (11144)
[1562167940] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.95%
[1562168028] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (19869) IS BETWEEN 80% AND 95% OF ENM NODES (24477)
[1562169745] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (11018) IS HIGHER THAN 95% OF ENM COMECIM NODES (11144)
[1562172514] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1562172682] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562172695] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1562172863] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562172875] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1562173043] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562174345] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.08%
[1562174496] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 10.4%
[1562175691] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562175877] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562175895] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1562176024] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1562176064] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562176154] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1562180583] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562180735] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.76%
[1562180889] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.94%
[1562182042] SERVICE FLAPPING ALERT: ieatnetsimv7035-29;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)