[1562626808] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;OK;HARD;3;OK
[1562626809] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;OK;HARD;3;OK
[1562627536] SERVICE ALERT: ieatnetsimv7035-54;check netsim disk;OK;HARD;3;OK
[1562627704] SERVICE FLAPPING ALERT: ieatlms3901;check node fs hc;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1562628190] SERVICE FLAPPING ALERT: ieatlms3901;check node fs;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1562628440] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.69%
[1562628591] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.17%
[1562628742] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.1%
[1562630410] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1562630531] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1562630651] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1562631136] SERVICE ALERT: ieatnetsimv7035-54;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1562631257] SERVICE ALERT: ieatnetsimv7035-54;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1562631377] SERVICE ALERT: ieatnetsimv7035-54;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1562632373] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.44%
[1562633970] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562634014] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1562634134] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1562634255] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1562636237] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562636304] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1562636425] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562636434] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1562636563] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1562636612] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562642782] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.78%
[1562642934] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.25%
[1562643084] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.13%
[1562646387] SERVICE ALERT: ieatnetsimv7035-100;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.56%
[1562646393] SERVICE ALERT: ieatnetsimv7035-102;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.08%
[1562646544] SERVICE ALERT: ieatnetsimv7035-102;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.97%
[1562646551] SERVICE ALERT: ieatnetsimv7035-100;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562646694] SERVICE ALERT: ieatnetsimv7035-102;check netsim cpu;OK;HARD;3;OK- CPU IS USED 76.86%
[1562646709] SERVICE ALERT: ieatnetsimv7035-100;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562649173] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1562649280] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562649304] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1562649374] SERVICE ALERT: ieatnetsimv7035-100;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562649517] SERVICE ALERT: ieatnetsimv7035-100;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562649651] SERVICE ALERT: ieatnetsimv7035-100;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562649790] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1562651271] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1562652251] SERVICE FLAPPING ALERT: ieatnetsimv5116-79;check netsim disk;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1562652251] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;OK;HARD;3;OK
[1562652256] SERVICE FLAPPING ALERT: ieatnetsimv5116-81;check netsim disk;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1562652256] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;OK;HARD;3;OK
[1562652681] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562653940] SERVICE ALERT: ieatnetsimv7035-100;check netsim cpu;OK;HARD;3;OK- CPU IS USED 72.13%
[1562654556] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562654781] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562654938] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim nodes;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1562654938] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562656160] SERVICE ALERT: ieatnetsimv7035-81;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1562656281] SERVICE ALERT: ieatnetsimv7035-81;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1562656401] SERVICE ALERT: ieatnetsimv7035-81;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1562656785] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1562657025] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1562657032] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1562657205] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1562657269] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1562657567] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1562657678] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1562657685] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562657711] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1562657845] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1562657850] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1562657953] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1562658070] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562658189] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1562658251] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562658255] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1562658528] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1562659440] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1562660116] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562660118] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1562660147] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562660334] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562660521] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562660867] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1562661782] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1562661912] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1562662041] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1562662230] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1562663527] SERVICE ALERT: ieatnetsimv7035-50;check netsim disk;OK;HARD;3;OK
[1562663538] SERVICE ALERT: ieatnetsimv7035-55;check netsim disk;OK;HARD;3;OK
[1562663557] SERVICE ALERT: ieatnetsimv7035-62;check netsim disk;OK;HARD;3;OK
[1562663601] SERVICE ALERT: ieatnetsimv7035-81;check netsim disk;OK;HARD;3;OK
[1562663612] SERVICE ALERT: ieatnetsimv7035-85;check netsim disk;OK;HARD;3;OK
[1562663777] SERVICE FLAPPING ALERT: ieatnetsimv7035-54;check netsim disk;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1562663777] SERVICE ALERT: ieatnetsimv7035-54;check netsim disk;OK;HARD;3;OK
[1562664063] SERVICE ALERT: ieatnetsimv7035-100;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562664528] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562664535] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562664538] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562664564] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562664567] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562664680] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562664686] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562664689] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562664715] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562664718] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562664831] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562664837] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562664841] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562664867] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562664869] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562664956] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562665108] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562665258] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562665259] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562665409] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562665561] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562665815] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562667375] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562667527] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562667669] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562667931] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (3998) IS HIGHER THAN 95% OF ENM CPP NODES (1460)
[1562668370] SERVICE ALERT: ieatnetsimv7035-100;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.27%
[1562668461] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;OK;HARD;3;OK- CPU IS USED 27.71%
[1562668468] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 40.21%
[1562668471] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 58.45%
[1562668497] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 68.66%
[1562668499] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 50.58%
[1562668521] SERVICE ALERT: ieatnetsimv7035-100;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 60.71%
[1562668524] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.85%
[1562668679] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562668757] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1562668817] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1562668830] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.4%
[1562668889] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;OK;HARD;3;OK- CPU IS USED 16.73%
[1562669181] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562669192] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.93%
[1562669201] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1562669288] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1562669327] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1562669350] SERVICE ALERT: ieatlms3901;check unsync nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1562669366] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562669430] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562669551] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562669565] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562669648] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;SOFT;1;(Service check timed out after 800.02 seconds)
[1562669677] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1562669701] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562670105] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1562670247] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1562670532] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1562670568] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1562670598] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1562671025] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1562671167] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1562671293] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562671453] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1562671488] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1562671525] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562671914] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.84%
[1562671945] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1562672373] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1562672865] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1562674165] SERVICE ALERT: ieatnetsimv7035-83;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1562674285] SERVICE ALERT: ieatnetsimv7035-83;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1562674406] SERVICE ALERT: ieatnetsimv7035-83;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1562674913] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM CPP NODES (0)
[1562675003] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562675016] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1562675016] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562676223] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.58%
[1562676374] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.34%
[1562676451] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1562676541] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4999) IS HIGHER THAN 95% OF ENM SNMP NODES (5017)
[1562676749] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (12064) IS HIGHER THAN 95% OF ENM COMECIM NODES (12144)
[1562676821] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3990) IS LOWER THAN 80% OF ENM CPP NODES (7110)
[1562677072] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3980) IS LOWER THAN 80% OF ENM CPP NODES (7110)
[1562677166] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1562677267] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized cpp;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1562677267] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3980) IS LOWER THAN 80% OF ENM CPP NODES (7110)
[1562677439] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (25452) IS HIGHER THAN 95% OF ENM NODES (25477)
[1562677518] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23582) IS BETWEEN 80% AND 95% OF ENM NODES (25477)
[1562677946] SERVICE ALERT: ieatnetsimv5116-88;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.43%
[1562678006] SERVICE ALERT: ieatnetsimv7035-83;check netsim disk;OK;HARD;3;OK
[1562678097] SERVICE ALERT: ieatnetsimv5116-88;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 49.04%
[1562678585] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (9984) IS FROM 80% TO 95% OF ENM COMECIM NODES (12144)
[1562678669] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562678769] SERVICE ALERT: ieatlms3901;check number nodes;OK;HARD;3;OK- NUMBER OF NODES CONFIGURED IN ENM (25477) IS HIGHER THAN 95% OF NETSIM NODES (22495)
[1562678772] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;2;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (6271) IS LOWER THAN 80% OF ENM COMECIM NODES (12144)
[1562678857] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562678957] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562678968] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (3008) IS LOWER THAN 80% OF ENM COMECIM NODES (12144)
[1562679059] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1562679059] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562679107] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 57.93%
[1562680060] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1562680190] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1562680320] SERVICE FLAPPING ALERT: ieatlms3901;check vcs sg;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1562680320] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1562680805] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized comecim;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1562680805] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (9863) IS FROM 80% TO 95% OF ENM COMECIM NODES (12144)
[1562680892] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6439) IS FROM 80% TO 95% OF ENM CPP NODES (7111)
[1562682639] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (12136) IS HIGHER THAN 95% OF ENM COMECIM NODES (12144)
[1562682718] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (7010) IS HIGHER THAN 95% OF ENM CPP NODES (7111)
[1562684904] SERVICE ALERT: ieatnetsimv7035-58;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1562685025] SERVICE ALERT: ieatnetsimv7035-58;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1562685145] SERVICE ALERT: ieatnetsimv7035-58;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1562685497] SERVICE ALERT: ieatnetsimv7035-50;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562685499] SERVICE ALERT: ieatnetsimv7035-52;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562685628] SERVICE ALERT: ieatnetsimv7035-50;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562685629] SERVICE ALERT: ieatnetsimv7035-52;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562685759] SERVICE ALERT: ieatnetsimv7035-52;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562685777] SERVICE ALERT: ieatlms3901;check pm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (24577) IS HIGHER THAN 95% OF ENM NODES (25478)
[1562685888] SERVICE ALERT: ieatnetsimv7035-52;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562686137] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (25305) IS HIGHER THAN 95% OF ENM NODES (25478)
[1562686458] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.21%
[1562686608] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 72.3%
[1562687394] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (6408) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7108)
[1562690461] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.96%
[1562691607] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17769) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (22194)
[1562696166] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1562696347] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1562696528] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1562696652] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562696833] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562697014] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562702595] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1562702652] SERVICE FLAPPING ALERT: ieatnetsimv5116-79;check netsim disk;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1562703377] SERVICE FLAPPING ALERT: ieatnetsimv7035-54;check netsim disk;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1562705681] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562705749] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562706256] SERVICE FLAPPING ALERT: ieatnetsimv5116-81;check netsim disk;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1562706767] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1562706767] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (6408) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7108)
[1562711167] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1562711527] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562712730] SERVICE FLAPPING ALERT: ieatlms3901;check vcs sg;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)