[1562195806] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.18%
[1562195957] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.27%
[1562196442] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.97%
[1562196593] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.33%
[1562196747] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.99%
[1562196773] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.49%
[1562196924] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.41%
[1562196991] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.76%
[1562197079] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.86%
[1562197141] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.14%
[1562197292] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.21%
[1562198920] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.89%
[1562200377] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;HARD;3;OK- CPU IS USED 76.12%
[1562200709] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;OK;HARD;3;OK- CPU IS USED 78.38%
[1562200922] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.54%
[1562203772] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1562203774] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562203955] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1562203957] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562203978] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized comecim;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562204137] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1562204139] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562207046] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.73%
[1562207197] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 20.57%
[1562211542] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.41%
[1562211692] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.62%
[1562219365] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1562219811] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (23048) IS BETWEEN 80% AND 95% OF ENM NODES (24477)
[1562220053] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (24328) IS HIGHER THAN 95% OF ENM NODES (24477)
[1562221138] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562222199] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1562222200] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1562225135] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1562225303] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1562226242] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.58%
[1562226394] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.02%
[1562226544] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.36%
[1562227030] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (23128) IS BETWEEN 80% AND 95% OF ENM NODES (24601)
[1562227441] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11143) IS FROM 80% TO 95% OF ENM COMECIM NODES (12110)
[1562227768] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (24328) IS HIGHER THAN 95% OF ENM NODES (25403)
[1562227830] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11143) IS FROM 80% TO 95% OF ENM COMECIM NODES (12144)
[1562227925] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (19867) IS LOWER THAN 80% OF ENM NODES (25477)
[1562228032] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11143) IS FROM 80% TO 95% OF ENM COMECIM NODES (12144)
[1562230175] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1562230175] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.58%
[1562230205] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562230304] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1562230326] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.7%
[1562230405] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562230434] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1562230476] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.98%
[1562230592] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562232243] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1562232460] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562234327] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562234515] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562234702] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1562234702] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562236490] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562236498] SERVICE ALERT: ieatnetsimv5116-105;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562236554] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.18%
[1562236569] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.63%
[1562236617] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562236630] SERVICE ALERT: ieatnetsimv5116-105;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562236705] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 49.28%
[1562236720] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.97%
[1562236745] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562236761] SERVICE ALERT: ieatnetsimv5116-105;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562237789] SERVICE ALERT: ieatnetsimv5116-61;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.97%
[1562237939] SERVICE ALERT: ieatnetsimv5116-61;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 18.94%
[1562238551] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.43%
[1562238702] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.06%
[1562238854] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.58%
[1562239966] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.72%
[1562240117] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.78%
[1562240273] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.59%
[1562240553] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.57%
[1562240704] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.31%
[1562240855] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.22%
[1562242484] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;OK;HARD;3;OK- CPU IS USED 76.88%
[1562243903] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.86%
[1562244486] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.69%
[1562244907] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.06%
[1562245062] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.67%
[1562245213] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.42%
[1562245661] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562245661] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1562245753] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562245772] SERVICE ALERT: ieatnetsimv5116-105;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562245791] SERVICE ALERT: ieatnetsimv7035-93;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1562245843] SERVICE ALERT: ieatlms3901;check node fs;OK;SOFT;2;OK
[1562245843] SERVICE ALERT: ieatlms3901;check node fs hc;OK;SOFT;2;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1562245911] SERVICE ALERT: ieatnetsimv7035-93;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1562246032] SERVICE ALERT: ieatnetsimv7035-93;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1562246653] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1562246783] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1562246912] SERVICE FLAPPING ALERT: ieatlms3901;check vcs sg;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1562246912] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1562247824] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1562248006] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1562248187] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1562248722] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1562249625] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562249806] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562249987] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562250890] SERVICE ALERT: ieatnetsimv5116-103;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562251018] SERVICE ALERT: ieatnetsimv5116-103;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562251146] SERVICE ALERT: ieatnetsimv5116-103;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562251161] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562251288] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562251416] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562252694] SERVICE ALERT: ieatnetsimv5116-104;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562252822] SERVICE ALERT: ieatnetsimv5116-104;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562252949] SERVICE ALERT: ieatnetsimv5116-104;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562252982] SERVICE ALERT: ieatnetsimv5116-105;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562253112] SERVICE ALERT: ieatnetsimv5116-105;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562253241] SERVICE ALERT: ieatnetsimv5116-105;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562254132] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1562254274] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1562254405] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1562256559] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (1543) IS LOWER THAN 80% OF ENM CPP NODES (7110)
[1562256677] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23048) IS BETWEEN 80% AND 95% OF ENM NODES (25477)
[1562256702] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (1543) IS LOWER THAN 80% OF ENM CPP NODES (7110)
[1562256813] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (19113) IS LOWER THAN 80% OF ENM NODES (25477)
[1562256879] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (1543) IS LOWER THAN 80% OF ENM CPP NODES (7110)
[1562256965] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23048) IS BETWEEN 80% AND 95% OF ENM NODES (25477)
[1562257048] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (19113) IS LOWER THAN 80% OF ENM NODES (25477)
[1562257187] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23048) IS BETWEEN 80% AND 95% OF ENM NODES (25477)
[1562257450] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (19113) IS LOWER THAN 80% OF ENM NODES (25477)
[1562258916] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.7%
[1562259070] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.27%
[1562259221] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.66%
[1562266452] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 82.95%
[1562268969] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562277428] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562277430] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562277433] SERVICE ALERT: ieatnetsimv5116-15;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562277556] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562277557] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562277566] SERVICE ALERT: ieatnetsimv5116-15;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562277634] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562277636] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562277643] SERVICE ALERT: ieatnetsimv5116-19;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562277649] SERVICE ALERT: ieatnetsimv5116-12;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562277697] SERVICE ALERT: ieatnetsimv5116-15;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562277767] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562277771] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562277778] SERVICE ALERT: ieatnetsimv5116-19;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562277781] SERVICE ALERT: ieatnetsimv5116-12;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562277845] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562277854] SERVICE ALERT: ieatnetsimv5116-19;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1562277900] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562277902] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562277913] SERVICE ALERT: ieatnetsimv5116-19;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562277965] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562277975] SERVICE ALERT: ieatnetsimv5116-19;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1562278086] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562278096] SERVICE ALERT: ieatnetsimv5116-19;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1562278339] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562278460] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562278581] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!