[1562454041] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.74%
[1562454192] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79.05%
[1562454342] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.97%
[1562454493] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.06%
[1562454644] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 93.97%
[1562458275] SERVICE FLAPPING ALERT: ieatnetsimv7035-18;check netsim cpu;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1562458275] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;HARD;3;OK- CPU IS USED 68.91%
[1562458655] SERVICE FLAPPING ALERT: ieatnetsimv7035-29;check netsim cpu;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1562458655] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.86%
[1562480678] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.61%
[1562480829] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 12.77%
[1562495754] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.27%
[1562495905] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.08%
[1562495949] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.26%
[1562496100] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.45%
[1562498287] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.52%
[1562498443] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.91%
[1562498595] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.25%
[1562502225] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;HARD;3;OK- CPU IS USED 72.29%
[1562510650] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.72%
[1562510801] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.67%
[1562510952] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 90.02%
[1562511549] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562511681] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562511813] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562514582] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 79.53%
[1562526706] SERVICE FLAPPING ALERT: ieatnetsimv7035-18;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562538648] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11154) IS FROM 80% TO 95% OF ENM COMECIM NODES (12144)
[1562538807] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (11896) IS HIGHER THAN 95% OF ENM COMECIM NODES (12144)