[1562973350] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1562974104] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6603) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1562976634] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562976771] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562978625] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 801.02 seconds)
[1562979545] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1562980327] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562980420] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6603) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1562981953] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (4143) IS LOWER THAN 80% OF ENM CPP NODES (6764)
[1562987381] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5435) IS FROM 80% TO 95% OF ENM CPP NODES (6764)
[1562991004] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (5393) IS LOWER THAN 80% OF ENM CPP NODES (6764)
[1562992345] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562992499] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562993332] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562994626] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5420) IS FROM 80% TO 95% OF ENM CPP NODES (6764)
[1562996449] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (5386) IS LOWER THAN 80% OF ENM CPP NODES (6764)
[1563001357] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized cpp;STARTED; Service appears to have started flapping (43.5% change >= 20.0% threshold)
[1563002732] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1563002884] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1563003049] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1563003193] SERVICE ALERT: ieatnetsimv5116-13;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563003314] SERVICE ALERT: ieatnetsimv5116-13;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563003435] SERVICE ALERT: ieatnetsimv5116-13;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563005835] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563005955] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563006076] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563008049] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized cpp;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1563008464] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2404) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563008590] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2404) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563008718] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2404) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563009043] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized cpp;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1563010247] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563010372] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563010496] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563012082] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized cpp;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1563012245] SERVICE ALERT: ieatnetsimv5116-20;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563012365] SERVICE ALERT: ieatnetsimv5116-20;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563012486] SERVICE ALERT: ieatnetsimv5116-20;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563012494] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563012499] SERVICE ALERT: ieatnetsimv5116-19;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563012614] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563012619] SERVICE ALERT: ieatnetsimv5116-19;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563012686] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563012735] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563012740] SERVICE ALERT: ieatnetsimv5116-19;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563012811] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563012935] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563013289] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563013413] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563013539] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563015849] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563015854] SERVICE ALERT: ieatnetsimv5116-14;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563015970] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563015972] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563015975] SERVICE ALERT: ieatnetsimv5116-14;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563016055] SERVICE ALERT: ieatnetsimv5116-08;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563016090] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563016092] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563016093] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563016095] SERVICE ALERT: ieatnetsimv5116-14;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563016176] SERVICE ALERT: ieatnetsimv5116-08;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563016213] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563016213] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563016296] SERVICE ALERT: ieatnetsimv5116-08;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563016333] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563022562] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1563024019] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563024023] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563024144] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563024147] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563024270] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563024272] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563024401] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1563024555] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1563024708] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1563026886] SERVICE ALERT: ieatnetsimv5116-20;check netsim disk;OK;HARD;3;OK
[1563026890] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;OK;HARD;3;OK
[1563026895] SERVICE ALERT: ieatnetsimv5116-14;check netsim disk;OK;HARD;3;OK
[1563027013] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;OK;HARD;3;OK
[1563027096] SERVICE ALERT: ieatnetsimv5116-08;check netsim disk;OK;HARD;3;OK
[1563027133] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;OK;HARD;3;OK
[1563027135] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;OK;HARD;3;OK
[1563027140] SERVICE ALERT: ieatnetsimv5116-19;check netsim disk;OK;HARD;3;OK
[1563030270] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1563033188] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563033903] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6392) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563036786] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.42%
[1563036937] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 4.08%
[1563042023] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563042690] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6392) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563058011] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)