[1574985870] SERVICE ALERT: ieatnetsimv7035-93;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574986812] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574987719] SERVICE ALERT: ieatnetsimv7035-93;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574987904] SERVICE ALERT: ieatnetsimv7035-93;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574988060] SERVICE ALERT: ieatnetsimv7035-93;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574988384] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1574990416] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574990539] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574990665] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1575006618] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575006685] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.86%
[1575006836] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 76.13%
[1575008839] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575008922] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575008974] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1575009051] SERVICE ALERT: ieatnetsimv7102-56;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.95%
[1575009058] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575009076] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-BSP;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (12)
[1575009129] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1575009195] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575009201] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-BSP;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (12)
[1575009202] SERVICE ALERT: ieatnetsimv7102-56;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.99%
[1575009313] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (968) IS LOWER THAN 80% OF ENM  NODES (1750)
[1575009328] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-BSP;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (12)
[1575009353] SERVICE ALERT: ieatnetsimv7102-56;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.82%
[1575009437] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1004) IS LOWER THAN 80% OF ENM  NODES (1750)
[1575009579] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1054) IS LOWER THAN 80% OF ENM  NODES (1750)
[1575010963] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1575011384] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1475) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1575011639] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575011825] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575012043] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575012592] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized Router6672;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575012641] SERVICE ALERT: ieatnetsimv7102-41;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.36%
[1575012792] SERVICE ALERT: ieatnetsimv7102-41;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 71.67%
[1575012983] SERVICE ALERT: ieatnetsimv7102-56;check netsim cpu;OK;HARD;3;OK- CPU IS USED 1.59%
[1575013047] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1575013171] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1575013215] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1) IS HIGHER THAN 20% OF ENM NODES (4)
[1575013298] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1575013341] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1) IS HIGHER THAN 20% OF ENM NODES (4)
[1575013466] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1) IS HIGHER THAN 20% OF ENM NODES (4)
[1575013490] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1575013492] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1575013587] SERVICE ALERT: ieatnetsimv9041-07;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575013620] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (2) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1575013621] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1575013751] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (2) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1575013796] SERVICE ALERT: ieatnetsimv9041-07;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1575013909] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575014006] SERVICE ALERT: ieatnetsimv9041-07;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1575016161] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (41) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1575016284] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (42) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1575016408] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (43) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1575016903] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1575016907] SERVICE ALERT: ieatnetsimv7035-93;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1575017077] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN RNC NODES !
[1575017213] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1575017334] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1575017357] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (3) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: 4)
[1575017453] SERVICE ALERT: ieatnetsimv9041-07;check netsim cpu;OK;HARD;3;OK- CPU IS USED 45.4%
[1575017454] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1575017640] SERVICE ALERT: ieatnetsimv9041-07;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1575018050] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;OK;HARD;3;OK- NUMBER OF PM FILES (50) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1575018053] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (100) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1575018116] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (743) IS HIGHER THAN 95% OF ENM  NODES (750)
[1575018594] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1575018672] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1575019815] SERVICE ALERT: ieatnetsimv7035-56;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575019835] SERVICE ALERT: ieatnetsimv7035-66;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575019967] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575019968] SERVICE ALERT: ieatnetsimv7035-56;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575019987] SERVICE ALERT: ieatnetsimv7035-66;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575019990] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.74%
[1575020015] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MINI-LINK-6352 NODES !
[1575020057] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1575020090] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575020140] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 21.11%
[1575020220] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1575020229] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13148) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575020244] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (20) IS LOWER THAN 80% OF ENM  NODES (50)
[1575020375] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13306) IS HIGHER THAN 95% OF ENM  NODES (14007)
[1575020380] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (20) IS LOWER THAN 80% OF ENM  NODES (50)
[1575020512] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13083) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575020531] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (33) IS LOWER THAN 80% OF ENM  NODES (50)
[1575020661] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13002) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575020667] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575020806] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12814) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575020818] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 58.28%
[1575020852] SERVICE ALERT: ieatnetsimv7035-103;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.81%
[1575021003] SERVICE ALERT: ieatnetsimv7035-103;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.63%
[1575021052] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (91) IS FROM 80% TO 95% OF ENM  NODES (98)
[1575021055] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1575021122] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (110) IS LOWER THAN 80% OF ENM  NODES (150)
[1575021132] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.24%
[1575021175] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575021178] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (92) IS FROM 80% TO 95% OF ENM  NODES (98)
[1575021269] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1575021276] SERVICE ALERT: ieatnetsimv7035-31;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.36%
[1575021283] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 73.29%
[1575021311] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (96) IS HIGHER THAN 95% OF ENM  NODES (98)
[1575021360] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575021391] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575021426] SERVICE ALERT: ieatnetsimv7035-31;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.88%
[1575021528] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575021546] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575021654] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575022340] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1575022621] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13377) IS HIGHER THAN 95% OF ENM  NODES (14007)
[1575022690] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575022707] SERVICE ALERT: ieatnetsimv7035-98;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.74%
[1575022841] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 25.43%
[1575022858] SERVICE ALERT: ieatnetsimv7035-98;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 8.93%
[1575022984] SERVICE ALERT: ieatnetsimv9041-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1575023135] SERVICE ALERT: ieatnetsimv9041-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.61%
[1575023163] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (682) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1575023286] SERVICE ALERT: ieatnetsimv9041-13;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.05%
[1575023400] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1138) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1575023530] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1197) IS HIGHER THAN 95% OF ENM  NODES (1200)
[1575024271] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (31.2% change >= 20.0% threshold)
[1575024271] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (29.3% change >= 20.0% threshold)
[1575024280] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (15) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575024452] SERVICE ALERT: ieatlms5735-1;check pm collection statistical BSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575025034] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1575025185] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 23.33%
[1575025191] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (72) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575025212] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1575025212] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575025261] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1575025321] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (682) IS FROM 80% TO 95% OF ENM  NODES (750)
[1575025333] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF PM FILES (10000) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575025336] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.45%
[1575025445] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (667) IS FROM 80% TO 95% OF ENM  NODES (750)
[1575025487] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575025578] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (664) IS FROM 80% TO 95% OF ENM  NODES (750)
[1575025591] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (608) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575025638] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575025822] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (22) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575026172] SERVICE ALERT: ieatnetsimv9041-11;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575026323] SERVICE ALERT: ieatnetsimv9041-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.46%
[1575026609] SERVICE ALERT: ieatnetsimv7035-98;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.43%
[1575026764] SERVICE ALERT: ieatnetsimv7035-98;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 10.72%
[1575026769] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (563) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575026917] SERVICE ALERT: ieatnetsimv9041-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 1.95%
[1575026955] SERVICE ALERT: ieatlms5735-1;check synchronized FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (33) IS LOWER THAN 80% OF ENM  NODES (100)
[1575027031] SERVICE ALERT: ieatnetsimv9041-04;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.37%
[1575027071] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.2%
[1575027079] SERVICE ALERT: ieatlms5735-1;check synchronized FRONTHAUL-6080;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (100)!
[1575027145] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SBG-IS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575027155] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575027157] SERVICE ALERT: ieatlms5735-1;check pm collection statistical EPG;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575027175] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.88%
[1575027181] SERVICE ALERT: ieatnetsimv9041-04;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 72.23%
[1575027222] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.97%
[1575027313] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1575027325] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.79%
[1575027372] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.88%
[1575027466] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.54%
[1575027476] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.93%
[1575027517] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: )
[1575027617] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.79%
[1575029268] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.68%
[1575029649] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (33.0% change >= 20.0% threshold)
[1575029649] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1575029649] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1575029928] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (33.0% change >= 20.0% threshold)
[1575029928] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1575029928] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1575030074] SERVICE ALERT: ieatnetsimv9041-11;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575030225] SERVICE ALERT: ieatnetsimv9041-11;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575030377] SERVICE ALERT: ieatnetsimv9041-11;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575030395] SERVICE ALERT: ieatnetsimv9041-06;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575030413] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.91%
[1575030440] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1575030440] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1575030440] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1575030546] SERVICE ALERT: ieatnetsimv9041-06;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.72%
[1575030564] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575030710] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1575030710] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1575030710] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1575030796] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575030849] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1575030849] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1575030849] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1575031000] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1575031000] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1575031000] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1575031097] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1575031097] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1575031097] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1575031308] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1575031308] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1575031308] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1575031685] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1575031781] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.83%
[1575031823] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1575031825] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1575031879] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1575031932] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.6%
[1575031974] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.97%
[1575031979] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.99%
[1575031994] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.96%
[1575031994] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.98%
[1575032073] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1575032124] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575032129] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575032144] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.06%
[1575032145] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 39.59%
[1575033363] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.85%
[1575033514] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79.05%
[1575033664] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.53%
[1575033815] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.96%
[1575033869] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1575033947] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1575033966] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.48%
[1575034277] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575034323] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1575034323] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1575034463] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575034533] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-BSP;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-BSP NODES !
[1575034630] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575034649] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (26.6% change >= 20.0% threshold)
[1575034649] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575034753] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575034848] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.02%
[1575034877] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575034935] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (41) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575035061] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (39) IS LOWER THAN 80% OF ENM  NODES (50)
[1575035185] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (39) IS LOWER THAN 80% OF ENM  NODES (50)
[1575035234] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10426) IS LOWER THAN 80% OF ENM  NODES (14007)
[1575035368] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10227) IS LOWER THAN 80% OF ENM  NODES (14007)
[1575035500] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10452) IS LOWER THAN 80% OF ENM  NODES (14007)
[1575035595] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575035722] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575035755] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 59.93%
[1575035760] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;OK;HARD;3;OK- CPU IS USED 59.13%
[1575035846] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575035983] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 27.11%
[1575036075] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (66) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575036192] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1575036192] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (26.6% change >= 20.0% threshold)
[1575036204] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (69) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575036328] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (76) IS HIGHER THAN 95% OF ENM  NODES (80)
[1575036515] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575036534] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 62.46%
[1575036681] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1575036990] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized EPG;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1575036990] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1575037322] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575037322] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13673) IS HIGHER THAN 95% OF ENM  NODES (14007)
[1575037650] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1575038440] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1575038440] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized EPG;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1575038440] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575038440] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (31.4% change >= 20.0% threshold)
[1575038479] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.68%
[1575038629] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.9%
[1575038780] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.7%
[1575039055] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1575039176] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1575039282] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.8%
[1575039296] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (25.3% change >= 20.0% threshold)
[1575039296] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1575039385] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.78%
[1575039390] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.89%
[1575039433] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.98%
[1575039536] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.73%
[1575039541] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.77%
[1575039584] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575039614] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1575039688] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.86%
[1575039693] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.97%
[1575039765] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575039915] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.99%
[1575040165] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1575040308] SERVICE ALERT: ieatnetsimv9034-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1575040315] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575040459] SERVICE ALERT: ieatnetsimv9034-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575040466] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575040610] SERVICE ALERT: ieatnetsimv9034-71;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.98%
[1575042385] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1575042505] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1575042625] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1575042897] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1575044560] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (11732) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575044712] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (11539) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575044848] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (11361) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575045332] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (65) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575045467] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (69) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575045593] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (70) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575045687] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575045810] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575046661] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13394) IS HIGHER THAN 95% OF ENM  NODES (14007)
[1575047399] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1575047399] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (79) IS HIGHER THAN 95% OF ENM  NODES (80)
[1575048474] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13194) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575048610] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13399) IS HIGHER THAN 95% OF ENM  NODES (14007)
[1575049535] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575049658] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575052060] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575052190] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575052314] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575052419] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (39) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1575052542] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (39) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1575052670] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (39) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1575057811] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;HARD;3;OK
[1575058594] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575058719] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575058843] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575059575] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10323) IS LOWER THAN 80% OF ENM  NODES (14007)
[1575059585] SERVICE ALERT: ieatlms3901;check pm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (150) IS LOWER THAN 80% OF ENM NODES (300)
[1575059628] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575059707] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10237) IS LOWER THAN 80% OF ENM  NODES (14007)
[1575059725] SERVICE ALERT: ieatlms3901;check pm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1575059764] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;SOFT;2;OK
[1575059847] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10164) IS LOWER THAN 80% OF ENM  NODES (14007)
[1575060004] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (53) IS LOWER THAN 80% OF ENM  NODES (80)
[1575060127] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (56) IS LOWER THAN 80% OF ENM  NODES (80)
[1575060253] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (53) IS LOWER THAN 80% OF ENM  NODES (80)
[1575060647] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (30) IS LOWER THAN 80% OF ENM  NODES (50)
[1575060898] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1575061018] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1575061139] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1575062057] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (67) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575063119] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (23) IS LOWER THAN 80% OF ENM  NODES (30)
[1575063243] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (24) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575063369] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575063474] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12390) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575063493] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575063619] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575064185] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575064252] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575064311] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575064519] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (27) IS LOWER THAN 80% OF ENM  NODES (98)
[1575064642] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (27) IS LOWER THAN 80% OF ENM  NODES (98)
[1575064769] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (27) IS LOWER THAN 80% OF ENM  NODES (98)
[1575064920] SERVICE FLAPPING ALERT: ieatnetsimv9034-60;check netsim cpu;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1575064920] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 30.13%
[1575065288] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10455) IS LOWER THAN 80% OF ENM  NODES (14007)
[1575065660] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (53) IS LOWER THAN 80% OF ENM  NODES (80)
[1575066057] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (31) IS LOWER THAN 80% OF ENM  NODES (50)
[1575066242] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (20) IS LOWER THAN 80% OF ENM  NODES (30)
[1575066368] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (20) IS LOWER THAN 80% OF ENM  NODES (30)
[1575066493] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (20) IS LOWER THAN 80% OF ENM  NODES (30)
[1575066579] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (91) IS FROM 80% TO 95% OF ENM  NODES (98)
[1575067101] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (11801) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575067223] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1575067223] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (24) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575067464] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575067864] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575068297] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1575068387] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (97) IS HIGHER THAN 95% OF ENM  NODES (98)
[1575068924] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13329) IS HIGHER THAN 95% OF ENM  NODES (14007)
[1575069951] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.01%
[1575070102] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 10.51%
[1575070737] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13211) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575070780] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575070869] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13200) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575071004] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (13141) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575071475] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575071599] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575071723] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)