[1573604168] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (614) IS FROM 80% TO 95% OF ENM  NODES (750)
[1573605942] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1573609523] SERVICE FLAPPING ALERT: ieatlms3901;check node fs hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573611825] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (666) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573611959] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (666) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573612087] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (666) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573614790] SERVICE ALERT: ieatlms3901;check pm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (20) IS LOWER THAN 80% OF ENM NODES (60)
[1573614915] SERVICE ALERT: ieatlms3901;check pm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1573615695] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (737) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573617653] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573617838] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1573618024] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1573618783] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573618911] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573619039] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573620267] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1573620394] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1573620522] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1573622339] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1573622644] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (24)
[1573624147] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1573624275] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1573624403] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1573626211] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1573626211] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1573629075] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1573629876] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2232) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573630037] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (2124) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573630165] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (2134) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573630611] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.83%
[1573630762] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 74.91%
[1573631619] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1573631747] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1573631875] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1573631973] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2554) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1573633640] SERVICE ALERT: ieatlms3901;check synchronized RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (3045) IS LOWER THAN 80% OF ENM  NODES (3849)
[1573633772] SERVICE ALERT: ieatlms3901;check synchronized RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (2181) IS LOWER THAN 80% OF ENM  NODES (3849)
[1573633782] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1580) IS LOWER THAN 80% OF ENM  NODES (2560)
[1573633875] SERVICE ALERT: ieatlms3901;check synchronized MGW;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (127) IS FROM 80% TO 95% OF ENM  NODES (150)
[1573633901] SERVICE ALERT: ieatlms3901;check synchronized RBS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (3239) IS FROM 80% TO 95% OF ENM  NODES (3849)
[1573633926] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1956) IS LOWER THAN 80% OF ENM  NODES (2560)
[1573634000] SERVICE ALERT: ieatlms3901;check synchronized MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (99) IS LOWER THAN 80% OF ENM  NODES (150)
[1573634057] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1855) IS LOWER THAN 80% OF ENM  NODES (2560)
[1573634125] SERVICE ALERT: ieatlms3901;check synchronized MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (111) IS LOWER THAN 80% OF ENM  NODES (150)
[1573635768] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1573636002] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1573636466] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1573636995] SERVICE ALERT: ieatlms3901;check synchronized RBS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (3849)!
[1573637037] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized ERBS;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1573637037] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (2560)!
[1573637045] SERVICE ALERT: ieatlms3901;check synchronized MGW;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1573637391] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized ERBS;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1573637391] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1573637877] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized ERBS;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1573637877] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1573638034] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573638185] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 17.55%
[1573638305] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (650) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573638439] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (650) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573638566] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (650) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573638764] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1573639317] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized ERBS;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1573639317] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (33.2% change >= 20.0% threshold)
[1573639346] SERVICE ALERT: ieatwlvm9034;check workload profiles;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1573639470] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized ERBS;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1573639470] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (33.2% change >= 20.0% threshold)
[1573639533] SERVICE ALERT: ieatwlvm9034;check workload profiles;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1573639601] HOST ALERT: ieatrcxb5978;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.210.244.252)
[1573639611] HOST ALERT: ieatrcxb6310;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.210.244.253)
[1573639613] HOST ALERT: ieatrcxb6422;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.210.245.0)
[1573639613] HOST ALERT: ieatrcxb6412;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.210.244.255)
[1573639665] SERVICE ALERT: ieatwlvm9034;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1573639666] HOST ALERT: ieatrcxb5978;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.210.244.252)
[1573639676] HOST ALERT: ieatrcxb6310;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.210.244.253)
[1573639678] HOST ALERT: ieatrcxb6422;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.210.245.0)
[1573639678] HOST ALERT: ieatrcxb6412;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.210.244.255)
[1573639730] HOST ALERT: ieatrcxb5978;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.210.244.252)
[1573639739] HOST ALERT: ieatrcxb6412;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.210.244.255)
[1573639740] HOST ALERT: ieatrcxb6310;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.210.244.253)
[1573639768] HOST ALERT: ieatrcxb6422;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1573639795] HOST ALERT: ieatrcxb5978;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.210.244.252)
[1573639800] SERVICE ALERT: ieatlms7435;check pm collection statistical SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573639804] HOST ALERT: ieatrcxb6412;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.210.244.255)
[1573639805] HOST ALERT: ieatrcxb6310;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.210.244.253)
[1573639832] HOST ALERT: ieatrcxb6422;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.210.245.0)
[1573639859] HOST ALERT: ieatrcxb5978;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.210.244.252)
[1573639868] SERVICE ALERT: ieatwlvm9034;check workload profiles started;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1573639868] HOST ALERT: ieatrcxb6412;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.210.244.255)
[1573639869] HOST ALERT: ieatrcxb6310;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.210.244.253)
[1573639897] HOST ALERT: ieatrcxb6422;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.210.245.0)
[1573639902] SERVICE ALERT: ieatnetsimv9034-45;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1573639923] SERVICE ALERT: ieatlms7435;check pm collection statistical SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573639924] HOST ALERT: ieatrcxb5978;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.210.244.252)
[1573639933] HOST ALERT: ieatrcxb6412;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.210.244.255)
[1573639934] HOST ALERT: ieatrcxb6310;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.210.244.253)
[1573639961] HOST ALERT: ieatrcxb6422;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.210.245.0)
[1573639988] HOST ALERT: ieatrcxb5978;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.210.244.252)
[1573639995] SERVICE ALERT: ieatwlvm9034;check workload profiles started;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1573639997] HOST ALERT: ieatrcxb6412;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.210.244.255)
[1573639998] HOST ALERT: ieatrcxb6310;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.210.244.253)
[1573640015] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (12545) IS LOWER THAN 80% OF ENM NODES (20060)
[1573640023] SERVICE ALERT: ieatnetsimv9034-45;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1573640026] HOST ALERT: ieatrcxb6422;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.210.245.0)
[1573640047] SERVICE ALERT: ieatlms7435;check pm collection statistical SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573640053] HOST ALERT: ieatrcxb5978;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.210.244.252)
[1573640062] HOST ALERT: ieatrcxb6412;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.210.244.255)
[1573640063] HOST ALERT: ieatrcxb6310;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.210.244.253)
[1573640090] HOST ALERT: ieatrcxb6422;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.210.245.0)
[1573640117] HOST ALERT: ieatrcxb5978;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.210.244.252)
[1573640124] SERVICE ALERT: ieatwlvm9034;check workload profiles started;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1573640126] HOST ALERT: ieatrcxb6412;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.210.244.255)
[1573640127] HOST ALERT: ieatrcxb6310;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.210.244.253)
[1573640143] SERVICE ALERT: ieatnetsimv9034-45;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1573640150] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (12545) IS LOWER THAN 80% OF ENM NODES (20060)
[1573640155] HOST ALERT: ieatrcxb6422;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.210.245.0)
[1573640182] HOST ALERT: ieatrcxb5978;DOWN;HARD;10;CRITICAL - Host Unreachable (10.210.244.252)
[1573640191] HOST ALERT: ieatrcxb6412;DOWN;HARD;10;CRITICAL - Host Unreachable (10.210.244.255)
[1573640192] HOST ALERT: ieatrcxb6310;DOWN;HARD;10;CRITICAL - Host Unreachable (10.210.244.253)
[1573640219] HOST ALERT: ieatrcxb6422;DOWN;HARD;10;CRITICAL - Host Unreachable (10.210.245.0)
[1573640267] SERVICE ALERT: ieatlms7435;check snapshots;WARNING;SOFT;1;CRITICAL- ENM SNAPSHOTS ARE PRESENT
[1573640284] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (12545) IS LOWER THAN 80% OF ENM NODES (20060)
[1573640407] SERVICE ALERT: ieatlms7435;check snapshots;WARNING;SOFT;2;CRITICAL- ENM SNAPSHOTS ARE PRESENT
[1573640468] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.51%
[1573640545] SERVICE ALERT: ieatlms7435;check snapshots;WARNING;HARD;3;CRITICAL- ENM SNAPSHOTS ARE PRESENT
[1573640562] SERVICE ALERT: ieatnetsimv5116-97;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.47%
[1573640564] SERVICE ALERT: ieatnetsimv5116-81;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.95%
[1573640572] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1573640619] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.76%
[1573640696] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (20)
[1573640699] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1573640712] SERVICE ALERT: ieatnetsimv5116-97;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.3%
[1573640714] SERVICE ALERT: ieatnetsimv5116-81;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 40.33%
[1573640738] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1573640770] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;OK;HARD;3;OK- CPU IS USED 19.7%
[1573640801] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1573640818] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (20)
[1573640827] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1573640895] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1573640922] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1573640940] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (20)
[1573641039] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1573641042] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1573641150] SERVICE ALERT: ieatlms7435;check rpm diff;WARNING;SOFT;1;WARNING - RPM ON SERVER ARE DIFFERENT FROM ISO (1.82.128)
[1573641175] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (18360) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (20034)
[1573641208] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1) IS LOWER THAN 80% OF ENM NODES (750)
[1573641273] SERVICE ALERT: ieatlms7435;check rpm diff;WARNING;SOFT;2;WARNING - RPM ON SERVER ARE DIFFERENT FROM ISO (1.82.128)
[1573641326] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (18451) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (20034)
[1573641334] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1573641396] SERVICE ALERT: ieatlms7435;check rpm diff;WARNING;HARD;3;WARNING - RPM ON SERVER ARE DIFFERENT FROM ISO (1.82.128)
[1573641475] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (18451) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (20034)
[1573643623] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized ERBS;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1573643623] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (37.7% change >= 20.0% threshold)
[1573643733] SERVICE ALERT: ieatnetsimv9041-07;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573643883] SERVICE ALERT: ieatnetsimv9041-07;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.38%
[1573645087] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20034)
[1573647036] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (862) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1573647164] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (862) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1573647291] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (862) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1573647451] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573647452] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573647575] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573647576] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573647699] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573647700] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573647899] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573648023] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573648147] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573649629] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573649782] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573649835] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1573649932] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.61%
[1573650898] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;OK;HARD;3;OK- NUMBER OF PM FILES (1200) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (1200)
[1573650976] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573651173] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573651308] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573651309] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (79) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573651370] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1573651750] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573651931] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.12%
[1573652082] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 50.79%
[1573653080] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573654657] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13317) IS FROM 80% TO 95% OF ENM  NODES (14026)
[1573654807] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13260) IS FROM 80% TO 95% OF ENM  NODES (14026)
[1573654952] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (13134) IS FROM 80% TO 95% OF ENM  NODES (14026)
[1573655220] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1573655345] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1573655470] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1573656777] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13787) IS HIGHER THAN 95% OF ENM  NODES (14026)
[1573657275] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (79) IS HIGHER THAN 95% OF ENM  NODES (80)
[1573659548] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.3%
[1573659699] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.24%
[1573660173] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1573660173] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (737) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573660446] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573663987] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573664045] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573664169] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573664226] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573664350] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573664414] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573665772] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573665968] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573666852] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1573666888] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1573666973] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1573667009] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1573667094] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1573667106] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573667130] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1573667142] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1573667227] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573667263] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1573667284] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573667348] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573667381] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (660) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573667383] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1573667405] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573667481] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573667509] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (672) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573667526] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573667546] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1573667602] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1573667604] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;SOFT;1;CRITICAL- san_alert_healthcheck FAILED
[1573667606] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1573667660] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (679) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573667667] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1573667723] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1573667725] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;SOFT;2;CRITICAL- san_alert_healthcheck FAILED
[1573667727] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1573667788] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1573667845] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;HARD;3;CRITICAL- san_alert_healthcheck FAILED
[1573667848] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1573667888] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573668009] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573668130] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1573669211] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1573669386] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573669588] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573669709] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573670006] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573670752] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1573670829] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1573671250] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized ERBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573671269] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (741) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573671477] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1573671518] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1573676842] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1573677026] SERVICE ALERT: ieatlms3901;check pm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (24) IS LOWER THAN 80% OF ENM NODES (48)
[1573677170] SERVICE ALERT: ieatlms3901;check pm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1573677433] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.78%
[1573677583] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 4.46%
[1573680612] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1573680675] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573681512] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93%
[1573681533] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.21%
[1573681662] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.57%
[1573681684] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 18.56%
[1573685684] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (672) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573685725] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573685843] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (672) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573685921] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573685970] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (672) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573686117] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1573686906] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (862) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1573687036] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (862) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1573687163] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (862) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1573687993] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573688735] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1573688856] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1573688876] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573688977] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1573688997] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1573689118] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1573689146] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1573689267] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1573689388] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1573689577] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (733) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)