[1573517494] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1573520019] SERVICE FLAPPING ALERT: ieatlms3901;check vcs llt heartbeat hc;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1573520019] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1573520168] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573520298] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573520349] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573520530] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1573521880] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573522061] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573522159] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573522242] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573522342] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573522522] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1573535122] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.69%
[1573535269] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.42%
[1573535273] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.92%
[1573535420] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 20.08%
[1573535424] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.47%
[1573536816] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.64%
[1573536966] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.6%
[1573537117] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 58.35%
[1573538546] SERVICE ALERT: ieatnetsimv017-18;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573538592] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573538667] SERVICE ALERT: ieatnetsimv017-18;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1573538726] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.37%
[1573538773] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573538799] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1573538799] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573538878] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.97%
[1573538958] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573539028] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;OK;HARD;3;OK- CPU IS USED 52.56%
[1573539054] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;OK;HARD;3;OK- CPU IS USED 69.52%
[1573539517] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1573540043] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1573540053] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573540097] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1573540728] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1573540948] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1573541004] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573541133] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573541147] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1573541168] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1750)
[1573541262] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573541299] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1750)
[1573541427] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1750)
[1573541429] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (48) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1573541556] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (48) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1573541683] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (48) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1573541919] SERVICE ALERT: ieatlms3901;check pm collection statistical BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (96)
[1573542047] SERVICE ALERT: ieatlms3901;check pm collection statistical BSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (96)
[1573542173] SERVICE ALERT: ieatlms3901;check pm collection statistical BSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (96)
[1573542588] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1573542695] SERVICE FLAPPING ALERT: ieatlms3901;check node fs;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1573542695] SERVICE FLAPPING ALERT: ieatlms3901;check vcs llt heartbeat hc;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1573543095] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573543177] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1573543228] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573543308] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1573543373] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573543443] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1573544396] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1573544524] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1573544555] SERVICE ALERT: ieatnetsimv7035-105;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.86%
[1573544651] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1573544705] SERVICE ALERT: ieatnetsimv7035-105;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.62%
[1573545250] SERVICE FLAPPING ALERT: ieatlms3901;check vcs llt heartbeat hc;STARTED; Service appears to have started flapping (26.8% change >= 20.0% threshold)
[1573546269] SERVICE FLAPPING ALERT: ieatlms3901;check vcs llt heartbeat hc;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1573546458] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1573546581] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1573546704] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1573546881] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1573546882] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1573547006] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1573547007] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1573547059] SERVICE ALERT: ieatlms5735-1;check fm hb failures EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (50) IS HIGHER THAN 20% OF ENM NODES (50)
[1573547130] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1573547132] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1573547184] SERVICE ALERT: ieatlms5735-1;check fm hb failures EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (50) IS HIGHER THAN 20% OF ENM NODES (50)
[1573547307] SERVICE ALERT: ieatlms5735-1;check fm hb failures EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (50) IS HIGHER THAN 20% OF ENM NODES (50)
[1573547473] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (11) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1573547481] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (135) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1573547600] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (11) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1573547606] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (135) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1573547726] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (11) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1573547732] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (135) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1573548011] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (1750) IS HIGHER THAN 95% OF ENM NODES (1750)
[1573548076] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1573548086] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1983) IS BETWEEN 5% AND 20% OF ENM NODES (14008)
[1573548200] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1573548214] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1983) IS BETWEEN 5% AND 20% OF ENM NODES (14008)
[1573548259] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1573548324] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1573548345] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1983) IS BETWEEN 5% AND 20% OF ENM NODES (14008)
[1573550480] HOST ALERT: ieatnetsimv5116-17;DOWN;SOFT;1;PING NOK
[1573550480] HOST ALERT: ieatnetsimv5116-36;DOWN;SOFT;1;PING NOK
[1573550542] HOST ALERT: ieatnetsimv5116-36;UP;SOFT;1;PING OK
[1573550542] HOST ALERT: ieatnetsimv5116-17;UP;SOFT;1;PING OK
[1573550972] HOST ALERT: ieatnetsimv7035-116;DOWN;SOFT;1;PING NOK
[1573551034] HOST ALERT: ieatnetsimv7035-116;UP;SOFT;1;PING OK
[1573554504] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573554514] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573554685] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573554695] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573554866] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573554876] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573556410] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573556672] HOST ALERT: ieatnetsimv7035-14;DOWN;SOFT;1;PING NOK
[1573556692] SERVICE FLAPPING ALERT: ieatlms3901;check node fs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1573556692] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1573556734] HOST ALERT: ieatnetsimv7035-14;UP;SOFT;1;PING OK
[1573556873] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573557023] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 29.04%
[1573557527] SERVICE FLAPPING ALERT: ieatlms3901;check node fs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1573557549] SERVICE ALERT: ieatlms5735-1;check synchronized ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (1000)
[1573557683] SERVICE ALERT: ieatlms5735-1;check synchronized ESC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (1000)
[1573557807] SERVICE ALERT: ieatlms5735-1;check synchronized ESC;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (1000)
[1573557997] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (12)
[1573558051] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573558120] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (12)
[1573558239] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573558243] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (12)
[1573558310] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.43%
[1573558329] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.33%
[1573558445] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (13)
[1573558450] SERVICE ALERT: ieatlms5735-1;check synchronized SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (20)
[1573558460] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 58.5%
[1573558480] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 26.94%
[1573558501] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.98%
[1573558571] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (13)
[1573558578] SERVICE ALERT: ieatlms5735-1;check synchronized SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (20)
[1573558652] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.41%
[1573558696] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (13)
[1573558705] SERVICE ALERT: ieatlms5735-1;check synchronized SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (20)
[1573558803] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.2%
[1573558904] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (49) IS LOWER THAN 80% OF ENM  NODES (750)
[1573558914] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (24)
[1573559028] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (125) IS LOWER THAN 80% OF ENM  NODES (750)
[1573559039] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (24)
[1573559156] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (178) IS LOWER THAN 80% OF ENM  NODES (750)
[1573559163] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (24)
[1573559427] SERVICE FLAPPING ALERT: ieatlms3901;check node fs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1573560295] SERVICE ALERT: ieatnetsimv9041-06;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573560446] SERVICE ALERT: ieatnetsimv9041-06;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573560597] SERVICE ALERT: ieatnetsimv9041-06;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573560891] SERVICE FLAPPING ALERT: ieatlms3901;check node fs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1573561222] SERVICE ALERT: ieatnetsimv9041-08;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573561374] SERVICE ALERT: ieatnetsimv9041-08;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573561525] SERVICE ALERT: ieatnetsimv9041-08;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573561787] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573563437] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573563588] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.73%
[1573566360] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (602) IS FROM 80% TO 95% OF ENM  NODES (750)
[1573569872] HOST ALERT: ieatnetsimv7035-64;DOWN;SOFT;1;PING NOK
[1573569934] HOST ALERT: ieatnetsimv7035-64;UP;SOFT;1;PING OK
[1573574472] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573574654] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573574753] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573574835] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573574934] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573575116] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1573576696] SERVICE FLAPPING ALERT: ieatlms3901;check node fs hc;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1573576696] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573578139] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.48%
[1573578289] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.37%
[1573578468] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (1579) IS BETWEEN 80% AND 95% OF ENM NODES (1750)
[1573578529] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (12517) IS BETWEEN 80% AND 95% OF ENM NODES (14008)
[1573578594] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (1579) IS BETWEEN 80% AND 95% OF ENM NODES (1750)
[1573578648] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (27) IS BETWEEN 80% AND 95% OF ENM NODES (30)
[1573578661] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (12517) IS BETWEEN 80% AND 95% OF ENM NODES (14008)
[1573578671] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes CISCO-ASR900;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (136) IS BETWEEN 80% AND 95% OF ENM NODES (150)
[1573578718] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (1579) IS BETWEEN 80% AND 95% OF ENM NODES (1750)
[1573578771] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (27) IS BETWEEN 80% AND 95% OF ENM NODES (30)
[1573578779] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (46) IS BETWEEN 80% AND 95% OF ENM NODES (50)
[1573578781] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (26) IS BETWEEN 80% AND 95% OF ENM NODES (30)
[1573578794] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (12517) IS BETWEEN 80% AND 95% OF ENM NODES (14008)
[1573578796] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (1076) IS BETWEEN 80% AND 95% OF ENM NODES (1200)
[1573578798] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (89) IS BETWEEN 80% AND 95% OF ENM NODES (100)
[1573578801] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes CISCO-ASR900;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (136) IS BETWEEN 80% AND 95% OF ENM NODES (150)
[1573578827] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (678) IS BETWEEN 80% AND 95% OF ENM NODES (750)
[1573578889] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ESC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (903) IS BETWEEN 80% AND 95% OF ENM NODES (1000)
[1573578895] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (27) IS BETWEEN 80% AND 95% OF ENM NODES (30)
[1573578902] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (46) IS BETWEEN 80% AND 95% OF ENM NODES (50)
[1573578905] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (26) IS BETWEEN 80% AND 95% OF ENM NODES (30)
[1573578921] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (1076) IS BETWEEN 80% AND 95% OF ENM NODES (1200)
[1573578922] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (89) IS BETWEEN 80% AND 95% OF ENM NODES (100)
[1573578925] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes CISCO-ASR900;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (136) IS BETWEEN 80% AND 95% OF ENM NODES (150)
[1573578952] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1573579015] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1573579032] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1573579034] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes DSC;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1573579051] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1573579053] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1573579534] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1573579536] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1573579624] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (1750) IS HIGHER THAN 95% OF ENM NODES (1750)
[1573579707] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (13940) IS HIGHER THAN 95% OF ENM NODES (14008)
[1573579800] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1573579830] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes CISCO-ASR900;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1573580765] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (557) IS LOWER THAN 80% OF ENM  NODES (750)
[1573583138] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1573583141] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1573583262] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1573583265] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1573583394] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1573583400] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1573584101] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1573584221] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1573584342] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1573591377] SERVICE FLAPPING ALERT: ieatlms3901;check node fs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573594760] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573594942] SERVICE ALERT: ieatlms3901;check node fs hc;OK;SOFT;2;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573602290] SERVICE ALERT: ieatlms3901;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (12)
[1573602414] SERVICE ALERT: ieatlms3901;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !