[1571701347] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1571703588] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2590) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571704120] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3849) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571704737] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2710) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571716177] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571717571] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.86%
[1571717722] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.89%
[1571717873] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.49%
[1571717945] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3166) IS LOWER THAN 80% OF ENM CPP NODES (4254)
[1571718085] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3166) IS LOWER THAN 80% OF ENM CPP NODES (4254)
[1571718218] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4055) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1571719852] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571720103] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571720223] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5589) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1571720392] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5588) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1571720565] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2710) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571720569] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5589) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1571721504] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;OK;HARD;3;OK- CPU IS USED 54.01%
[1571722400] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (6589) IS HIGHER THAN 95% OF ENM CPP NODES (6589)
[1571724761] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571727374] SERVICE ALERT: ieatlms3901;check pm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (26212) IS HIGHER THAN 95% OF ENM NODES (27030)
[1571727882] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2352) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571728021] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2424) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1571729003] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571729188] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571729373] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571731239] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571734904] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571735097] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571735282] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571735690] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (150) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571736214] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571736257] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.52%
[1571736262] SERVICE ALERT: ieatnetsimv5116-57;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571736267] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571736408] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.72%
[1571736413] SERVICE ALERT: ieatnetsimv5116-57;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.39%
[1571736417] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 10.61%
[1571736424] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (180) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571738431] SERVICE ALERT: ieatnetsimv017-22;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571738489] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571738579] SERVICE ALERT: ieatnetsimv017-22;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1571738677] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571738729] SERVICE ALERT: ieatnetsimv017-22;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1571738864] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571738947] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1571738947] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571740653] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571740731] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571742509] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.99%
[1571742659] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.73%
[1571742968] SERVICE ALERT: ieatnetsimv7035-110;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.63%
[1571743118] SERVICE ALERT: ieatnetsimv7035-110;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 43%
[1571747756] SERVICE ALERT: ieatnetsimv017-22;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1571748014] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571748203] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571748390] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571750804] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (32.6% change >= 20.0% threshold)
[1571750804] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1571750877] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571750880] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571751028] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571751033] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571751180] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571751185] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571751229] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571751383] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571751537] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571752108] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571752168] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (31.8% change >= 20.0% threshold)
[1571752261] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571752293] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (31.8% change >= 20.0% threshold)
[1571752414] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571752705] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571752706] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571752708] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571752709] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571752710] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571752744] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571752865] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571752865] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571752870] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571752873] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571752878] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571752907] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571753018] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571753018] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571753022] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571753026] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571753031] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571753064] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571753308] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753310] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753312] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753412] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571753462] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571753463] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571753464] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571753597] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571753611] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753615] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571753616] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571753617] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571753763] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571753857] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1571753857] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571753875] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753885] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753889] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753893] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753894] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753898] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753900] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753902] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753906] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753908] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753909] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753909] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753910] SERVICE ALERT: ieatnetsimv7035-88;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753911] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753914] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753915] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571753915] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571753995] SERVICE ALERT: ieatnetsimv7035-49;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1571754026] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754037] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754041] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754045] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754047] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754050] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754052] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754054] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754057] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754059] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 11.7%
[1571754060] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754061] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754062] SERVICE ALERT: ieatnetsimv7035-88;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754064] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754066] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754069] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571754178] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754189] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754193] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754197] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754202] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754204] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754205] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754206] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754208] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754212] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754212] SERVICE ALERT: ieatnetsimv7035-88;check netsim cpu;OK;HARD;3;OK- CPU IS USED 23.6%
[1571754212] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754218] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754218] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754222] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571754807] SERVICE ALERT: ieatnetsimv7035-49;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571756652] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.45%
[1571756657] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;OK;HARD;3;OK- CPU IS USED 66%
[1571756694] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;HARD;3;OK- CPU IS USED 24.86%
[1571757247] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.32%
[1571757247] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.05%
[1571757545] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.78%
[1571757602] SERVICE ALERT: ieatnetsimv7035-52;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1571757808] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 11.87%
[1571757820] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.96%
[1571757824] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 27.12%
[1571757827] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;OK;HARD;3;OK- CPU IS USED 17.57%
[1571757834] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.6%
[1571757836] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.36%
[1571757837] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.54%
[1571757839] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.89%
[1571757842] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;OK;HARD;3;OK- CPU IS USED 17.04%
[1571757843] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;OK;HARD;3;OK- CPU IS USED 11.1%
[1571757848] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.63%
[1571757897] SERVICE ALERT: ieatnetsimv7035-52;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571760248] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.97%
[1571760248] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 52.64%
[1571762547] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1571762667] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1571762788] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1571762855] SERVICE ALERT: ieatnetsimv7035-66;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571762984] SERVICE ALERT: ieatnetsimv7035-66;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571765610] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.56%
[1571765616] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.39%
[1571766844] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.31%
[1571767187] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13162) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14033)
[1571767328] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13153) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14033)
[1571767462] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.33%
[1571767469] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13153) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14033)
[1571768045] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.83%
[1571768633] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.41%
[1571769753] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571770274] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3999) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571770548] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (180) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571771124] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.87%
[1571771283] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 59.18%
[1571774795] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571775034] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.14%
[1571775103] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13373) IS FROM 80% TO 95% OF ENM COMECIM NODES (14151)
[1571775184] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78%
[1571775279] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13373) IS FROM 80% TO 95% OF ENM COMECIM NODES (14151)
[1571775335] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.87%
[1571775430] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13373) IS FROM 80% TO 95% OF ENM COMECIM NODES (14151)
[1571775486] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 59.47%
[1571784525] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)