[1570748449] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1570748454] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1570748462] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1570748569] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1570748575] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1570748582] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1570748598] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1570748690] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1570748695] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1570748719] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1570748735] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570748840] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1570748855] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1570748925] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1570748976] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1570749045] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1570749166] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1570749412] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1570749533] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1570749654] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1570751705] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13387) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570751850] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13355) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570751995] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13331) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570752100] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.6%
[1570752251] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.88%
[1570755323] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570755477] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570755608] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1570759222] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13539) IS HIGHER THAN 95% OF ENM COMECIM NODES (14097)
[1570766979] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (180) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570767373] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2710) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570768088] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570768975] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570771705] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570771710] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570772158] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1570772301] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1570772450] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570773572] SERVICE ALERT: ieatnetsimv7035-96;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.29%
[1570773723] SERVICE ALERT: ieatnetsimv7035-96;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.77%
[1570773874] SERVICE ALERT: ieatnetsimv7035-96;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.81%
[1570773885] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1570773982] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1570774129] SERVICE FLAPPING ALERT: ieatnetsimv7035-52;check netsim cpu;STARTED; Service appears to have started flapping (27.2% change >= 20.0% threshold)
[1570774129] SERVICE FLAPPING ALERT: ieatnetsimv7035-61;check netsim cpu;STARTED; Service appears to have started flapping (25.7% change >= 20.0% threshold)
[1570774129] SERVICE FLAPPING ALERT: ieatnetsimv7035-62;check netsim cpu;STARTED; Service appears to have started flapping (27.2% change >= 20.0% threshold)
[1570774543] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1570774595] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1570774673] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1570774798] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1570774899] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (1038) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1570774910] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (1038) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1570774912] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1570775052] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1038) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1570776317] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.18%
[1570776374] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570776468] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.11%
[1570776524] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 19.86%
[1570776928] HOST ALERT: ieatnetsimv7035-93;DOWN;SOFT;1;PING NOK
[1570776990] HOST ALERT: ieatnetsimv7035-93;UP;SOFT;1;PING OK
[1570777470] SERVICE ALERT: ieatnetsimv7035-70;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.28%
[1570777621] SERVICE ALERT: ieatnetsimv7035-70;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 22.98%
[1570782030] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1570782281] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1570784123] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570784311] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570784498] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570785636] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (608) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1570785764] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (608) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1570785888] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570785889] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (608) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1570786018] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570786143] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570787105] SERVICE ALERT: ieatwlvm5116;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1570787716] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1570788165] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570789991] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570790060] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570790142] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570790211] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570790295] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570790304] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.82%
[1570790305] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.42%
[1570790315] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570790362] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570790365] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.7%
[1570790422] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570790454] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 56.28%
[1570790455] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 20.71%
[1570790465] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 33.14%
[1570790516] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570790572] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 40.5%
[1570790667] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;OK;HARD;3;OK- CPU IS USED 25.92%
[1570790723] SERVICE FLAPPING ALERT: ieatnetsimv7035-61;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1570792960] SERVICE ALERT: ieatlms3901;check pm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (26267) IS HIGHER THAN 95% OF ENM NODES (27052)
[1570793481] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (10151) IS LOWER THAN 80% OF ENM COMECIM NODES (14097)
[1570793635] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12050) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570793803] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12967) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570793926] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;OK;HARD;3;OK- CPU IS USED 18.74%
[1570793977] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23752) IS BETWEEN 80% AND 95% OF ENM NODES (27052)
[1570793993] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.46%
[1570794232] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23752) IS BETWEEN 80% AND 95% OF ENM NODES (27052)
[1570794454] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23752) IS BETWEEN 80% AND 95% OF ENM NODES (27052)
[1570795643] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13864) IS HIGHER THAN 95% OF ENM COMECIM NODES (14097)
[1570798218] SERVICE FLAPPING ALERT: ieatnetsimv7035-62;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1570799276] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13248) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570799372] SERVICE ALERT: ieatnetsimv7035-70;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1570799422] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13323) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570799523] SERVICE ALERT: ieatnetsimv7035-70;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 40.75%
[1570799569] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13928) IS HIGHER THAN 95% OF ENM COMECIM NODES (14097)
[1570801065] SERVICE ALERT: ieatnetsimv7035-74;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570801195] SERVICE ALERT: ieatnetsimv7035-74;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570801325] SERVICE ALERT: ieatnetsimv7035-74;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570802802] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.6%
[1570802952] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 35.27%
[1570805205] SERVICE FLAPPING ALERT: ieatnetsimv7035-52;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570810323] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570810518] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1570810711] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1570812532] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1570812653] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1570812773] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1570813106] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570813111] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570813227] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570813231] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570813337] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570813347] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570813352] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1570813463] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570813486] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1570813592] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1570813752] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;SOFT;2;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1570814384] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570814635] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1570815213] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570816497] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1570816678] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1570817530] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1570817709] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1570817776] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (9701) IS LOWER THAN 80% OF ENM COMECIM NODES (14097)
[1570817887] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1570817917] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;2;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (9838) IS LOWER THAN 80% OF ENM COMECIM NODES (14097)
[1570818159] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (9966) IS LOWER THAN 80% OF ENM COMECIM NODES (14097)
[1570818812] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570819053] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570819178] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570819302] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570819925] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570819993] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized comecim;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1570819993] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12810) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570820050] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570820174] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570822154] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570822234] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3708) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1570822384] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3867) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1570822562] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4059) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1570822831] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570822907] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570823012] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570823497] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3879) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570823629] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (8236) IS LOWER THAN 80% OF ENM COMECIM NODES (14097)
[1570823787] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570823893] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570824277] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570824402] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570824462] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570824527] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570824647] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570824655] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570825456] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13574) IS HIGHER THAN 95% OF ENM COMECIM NODES (14097)
[1570827289] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13165) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570827431] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12481) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570827579] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12735) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570829193] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2069) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1570829326] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1874) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1570829401] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13734) IS HIGHER THAN 95% OF ENM COMECIM NODES (14097)
[1570829458] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1874) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1570831229] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12998) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570831278] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2485) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1570831373] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13179) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570831519] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12991) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570831626] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6439) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570831912] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1570831912] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570832400] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1570833353] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13715) IS HIGHER THAN 95% OF ENM COMECIM NODES (14097)