[1564786965] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4925) IS HIGHER THAN 95% OF ENM SNMP NODES (5017)
[1564787373] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1564787373] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14027) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564788789] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564788855] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.86%
[1564788973] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564789008] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.82%
[1564789080] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1564789080] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2535) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564789159] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564790195] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564790266] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564790376] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564790451] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564790557] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564790635] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564790957] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564791063] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564791140] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564791252] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564791441] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1564791560] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1564791661] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564791741] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1564791922] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1564792528] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7151) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564792758] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564792911] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2535) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564793095] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564793124] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564793284] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564793307] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564793449] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1564793470] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1564793662] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.92%
[1564793813] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.36%
[1564794369] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1564794741] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1564795172] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564795987] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4277) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1564796135] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4296) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1564796277] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4322) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1564796390] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.41%
[1564796817] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1564796882] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1564796991] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1564797191] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1564798100] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1564798100] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4866) IS HIGHER THAN 95% OF ENM SNMP NODES (5017)
[1564798676] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564798858] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564798982] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1564799039] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1564799163] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1564799345] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1564799455] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.35%
[1564799605] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.11%
[1564803172] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564804092] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1564804917] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564805055] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564807195] SERVICE ALERT: ieatnetsimv5116-85;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1564807220] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.94%
[1564807371] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 72.36%
[1564807521] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.06%
[1564807672] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.63%
[1564807713] SERVICE ALERT: ieatnetsimv5116-85;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1564807808] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.15%
[1564807959] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.06%
[1564808110] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 11.71%
[1564808141] SERVICE ALERT: ieatnetsimv5116-85;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1564808720] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564808906] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564809091] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1564811511] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1564811872] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1564812233] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1564817520] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6557) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1564817671] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6522) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1564817814] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6327) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1564818617] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (24222) IS BETWEEN 80% AND 95% OF ENM NODES (25520)
[1564818906] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH CM SUPERVISION (24191) IS BETWEEN 80% AND 95% OF ENM NODES (25520)
[1564819137] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (24007) IS BETWEEN 80% AND 95% OF ENM NODES (25520)
[1564819637] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (5245) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1564823329] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4231) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1564823471] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4197) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1564823691] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4159) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1564824251] SERVICE ALERT: ieatnetsimv5116-50;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1564824372] SERVICE ALERT: ieatnetsimv5116-50;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1564824492] SERVICE ALERT: ieatnetsimv5116-50;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1564824652] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (19461) IS LOWER THAN 80% OF ENM NODES (25520)
[1564825144] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11056) IS FROM 80% TO 95% OF ENM COMECIM NODES (12144)
[1564825304] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11046) IS FROM 80% TO 95% OF ENM COMECIM NODES (12144)
[1564825460] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11015) IS FROM 80% TO 95% OF ENM COMECIM NODES (12144)
[1564825517] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (3263) IS LOWER THAN 80% OF ENM SNMP NODES (5017)
[1564827305] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (9471) IS LOWER THAN 80% OF ENM COMECIM NODES (12144)
[1564828207] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1564829701] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1564830472] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1564831468] SERVICE ALERT: ieatnetsimv5116-58;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1564831589] SERVICE ALERT: ieatnetsimv5116-58;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1564831709] SERVICE ALERT: ieatnetsimv5116-58;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1564834517] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564835299] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564836219] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564836960] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (5693) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564837208] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1564837329] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1564837450] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1564837463] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (1459) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564837673] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1564837794] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1564837906] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1564837915] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1564837942] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564837962] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1564837966] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1564838027] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1564838064] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564838083] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1564838087] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1564838148] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1564838185] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564838204] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1564838209] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1564838428] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564838549] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1564838670] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1564839461] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5870) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1564840083] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1564840204] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1564840325] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1564840564] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1564840685] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1564840806] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1564842339] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564842922] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1564843144] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1564843273] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1564843402] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1564843560] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (10550) IS FROM 80% TO 95% OF ENM COMECIM NODES (12144)
[1564843654] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1564843668] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1564843691] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1564844465] SERVICE ALERT: ieatlms3901;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1564845353] SERVICE FLAPPING ALERT: ieatlms3901;check nas hc;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1564845353] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1564845415] SERVICE ALERT: ieatlms3901;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1564845524] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564845712] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564845904] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564846002] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564846185] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564846368] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564846551] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1564846735] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1564848162] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1564848598] SERVICE FLAPPING ALERT: ieatlms3901;check vcs cluster hc;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1564848598] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564848751] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1564849066] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (24991) IS HIGHER THAN 95% OF ENM NODES (25520)
[1564849142] SERVICE ALERT: ieatlms3901;check unsync nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1564850156] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.93%
[1564850223] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23812) IS BETWEEN 80% AND 95% OF ENM NODES (25520)
[1564850306] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 60.14%
[1564850461] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564850644] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1564850827] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564850974] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1564853173] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1564853173] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1564855052] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564855241] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564855429] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564855430] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1564855560] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1564855693] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1564856195] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564856295] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564856478] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1564856662] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1564858287] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1564859737] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1564859850] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1564864969] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.11%
[1564865118] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.98%
[1564865120] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.85%
[1564865268] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.62%
[1564865271] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.65%
[1564865726] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564865962] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (7032) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21429)
[1564867589] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1564867772] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1564867955] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1564868558] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564868715] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564868868] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564868901] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;HARD;3;OK- CPU IS USED 75.09%
[1564869024] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564869180] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564869336] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564869642] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (18847) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21429)
[1564869818] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1564869818] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1564869995] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1564870088] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (5336) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1564870326] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1564870706] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1564870916] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1564871182] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1564871278] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1564871930] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5870) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1564872535] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1564872806] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1564873014] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK