[1562281247] SERVICE FLAPPING ALERT: ieatnetsimv7035-27;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1562282115] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.14%
[1562282266] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.63%
[1562282419] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.73%
[1562282508] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.79%
[1562282659] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.47%
[1562282811] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.18%
[1562282961] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.81%
[1562283112] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.14%
[1562283208] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.87%
[1562283359] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.18%
[1562283510] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 92.63%
[1562285015] SERVICE FLAPPING ALERT: ieatlms3901;check vcs sg;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562286049] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1562286050] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.04%
[1562287141] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;HARD;3;OK- CPU IS USED 81.02%
[1562287911] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562288092] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562288273] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562289680] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.37%
[1562289839] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562289991] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.18%
[1562291683] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.33%
[1562291834] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.82%
[1562291985] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.03%
[1562292550] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;OK;HARD;3;OK
[1562293622] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.98%
[1562295263] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1562295615] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.89%
[1562295623] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1562295954] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1562296150] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562296271] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562296391] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562297545] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.44%
[1562299250] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.98%
[1562299401] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.07%
[1562299523] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562299552] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.87%
[1562299655] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562299704] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.94%
[1562299788] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562299854] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.74%
[1562299992] SERVICE FLAPPING ALERT: ieatnetsimv5116-16;check netsim disk;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1562299992] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;OK;HARD;3;OK
[1562300182] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;OK;HARD;3;OK
[1562300233] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.05%
[1562300383] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.69%
[1562301601] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562303485] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 79.36%
[1562304453] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.68%
[1562304604] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.3%
[1562304686] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562304759] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1562304759] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562304814] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562304888] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562305018] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562305148] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562305210] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562305395] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562305580] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562306157] SERVICE ALERT: ieatnetsimv5116-13;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562306190] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.31%
[1562306278] SERVICE ALERT: ieatnetsimv5116-13;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562306341] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 21.91%
[1562306398] SERVICE ALERT: ieatnetsimv5116-13;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562306407] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562306412] SERVICE ALERT: ieatnetsimv5116-14;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562306414] SERVICE ALERT: ieatnetsimv5116-08;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562306415] SERVICE ALERT: ieatnetsimv5116-12;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562306527] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562306530] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562306533] SERVICE ALERT: ieatnetsimv5116-14;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562306534] SERVICE ALERT: ieatnetsimv5116-08;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562306535] SERVICE ALERT: ieatnetsimv5116-12;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562306648] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562306650] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562306651] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562306651] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562306654] SERVICE ALERT: ieatnetsimv5116-14;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562306655] SERVICE ALERT: ieatnetsimv5116-08;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562306656] SERVICE ALERT: ieatnetsimv5116-12;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562306771] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562306772] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562306772] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562306891] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562306892] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562307193] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1562307313] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1562307382] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562307434] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1562307445] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562307503] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562307623] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562307734] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.88%
[1562307885] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.58%
[1562308035] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 70.55%
[1562309310] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562309500] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562309685] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562309757] HOST ALERT: ieatnetsimv5116-111;DOWN;SOFT;1;PING NOK
[1562309820] HOST ALERT: ieatnetsimv5116-111;UP;SOFT;1;PING OK
[1562310001] SERVICE ALERT: ieatnetsimv5116-20;check netsim disk;CRITICAL;SOFT;1;CRITICAL- FILESYSTEM IS FULL!
[1562310121] SERVICE ALERT: ieatnetsimv5116-20;check netsim disk;CRITICAL;SOFT;2;CRITICAL- FILESYSTEM IS FULL!
[1562310242] SERVICE ALERT: ieatnetsimv5116-20;check netsim disk;CRITICAL;HARD;3;CRITICAL- FILESYSTEM IS FULL!
[1562311550] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1562311550] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562312694] SERVICE ALERT: ieatnetsimv7035-17;check netsim cpu;OK;HARD;3;OK- CPU IS USED 31.07%
[1562312817] SERVICE ALERT: ieatnetsimv7035-17;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562312951] SERVICE ALERT: ieatnetsimv7035-17;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562313085] SERVICE ALERT: ieatnetsimv7035-17;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562313432] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1562314900] SERVICE ALERT: ieatnetsimv7035-17;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562315216] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1562315334] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1562315364] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11244) IS FROM 80% TO 95% OF ENM COMECIM NODES (12144)
[1562315401] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1562315603] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1562316435] SERVICE ALERT: ieatnetsimv7035-22;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562316564] SERVICE ALERT: ieatnetsimv7035-22;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562316648] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1562316694] SERVICE ALERT: ieatnetsimv7035-22;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562316959] SERVICE ALERT: ieatnetsimv5116-64;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1562316992] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.83%
[1562317080] SERVICE ALERT: ieatnetsimv5116-64;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1562317140] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4750) IS FROM 80% TO 95% OF ENM SNMP NODES (5017)
[1562317143] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.17%
[1562317199] SERVICE ALERT: ieatnetsimv5116-13;check netsim disk;OK;HARD;3;OK
[1562317204] SERVICE ALERT: ieatnetsimv5116-64;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1562317281] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4776) IS HIGHER THAN 95% OF ENM SNMP NODES (5017)
[1562317295] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.94%
[1562317443] SERVICE ALERT: ieatnetsimv5116-20;check netsim disk;OK;HARD;3;OK
[1562317448] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;OK;HARD;3;OK
[1562317454] SERVICE ALERT: ieatnetsimv5116-14;check netsim disk;OK;HARD;3;OK
[1562317457] SERVICE ALERT: ieatnetsimv5116-12;check netsim disk;OK;HARD;3;OK
[1562317572] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;OK;HARD;3;OK
[1562317655] SERVICE ALERT: ieatnetsimv5116-08;check netsim disk;OK;HARD;3;OK
[1562317692] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;OK;HARD;3;OK
[1562317693] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;OK;HARD;3;OK
[1562317697] SERVICE ALERT: ieatnetsimv5116-19;check netsim disk;OK;HARD;3;OK
[1562317791] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1562318032] SERVICE ALERT: ieatnetsimv7035-93;check netsim disk;OK;HARD;3;OK
[1562318234] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;OK;HARD;3;OK
[1562318423] SERVICE FLAPPING ALERT: ieatnetsimv5116-18;check netsim disk;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1562318423] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;OK;HARD;3;OK
[1562318505] SERVICE ALERT: ieatnetsimv7035-22;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562318892] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562318944] SERVICE ALERT: ieatnetsimv5116-17;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319020] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319022] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319072] SERVICE ALERT: ieatnetsimv5116-17;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319085] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319109] SERVICE ALERT: ieatnetsimv5116-15;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319147] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562319150] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319189] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.78%
[1562319191] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319199] SERVICE ALERT: ieatnetsimv5116-17;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562319213] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319228] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319238] SERVICE ALERT: ieatnetsimv5116-15;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319268] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1562319277] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562319313] SERVICE ALERT: ieatnetsimv5116-12;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319313] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319320] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319326] SERVICE ALERT: ieatnetsimv5116-19;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319341] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562319342] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319356] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319369] SERVICE ALERT: ieatnetsimv5116-15;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562319443] SERVICE ALERT: ieatnetsimv5116-12;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319443] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319448] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562319457] SERVICE ALERT: ieatnetsimv5116-19;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319470] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319484] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562319574] SERVICE ALERT: ieatnetsimv5116-12;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562319575] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562319589] SERVICE ALERT: ieatnetsimv5116-19;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562319596] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562319678] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562319806] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562319933] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562320258] SERVICE ALERT: ieatnetsimv7035-20;check netsim cpu;OK;HARD;3;OK- CPU IS USED 15.8%
[1562320595] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562320747] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.42%
[1562320810] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (11706) IS HIGHER THAN 95% OF ENM COMECIM NODES (12144)
[1562320844] SERVICE ALERT: ieatnetsimv7035-27;check netsim cpu;OK;HARD;3;OK- CPU IS USED 51.3%
[1562321086] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562321406] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562321415] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562321534] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562321546] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562321662] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562321679] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562321964] SERVICE ALERT: ieatnetsimv7035-25;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562322101] SERVICE ALERT: ieatnetsimv7035-25;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562322756] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562322808] SERVICE ALERT: ieatnetsimv5116-17;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562322952] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562322981] SERVICE ALERT: ieatnetsimv5116-15;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562323053] SERVICE ALERT: ieatnetsimv7035-30;check netsim cpu;OK;HARD;3;OK- CPU IS USED 52.13%
[1562323057] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562323093] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562323186] SERVICE ALERT: ieatnetsimv5116-12;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562323470] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562323541] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562324502] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562324525] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;OK;HARD;3;OK- CPU IS USED 20.07%
[1562324653] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.43%
[1562324989] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562325007] SERVICE ALERT: ieatnetsimv5116-19;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562325294] SERVICE FLAPPING ALERT: ieatnetsimv5116-18;check netsim nodes;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1562325294] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1562326561] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6407) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (1544)
[1562332006] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1562332157] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.99%
[1562333513] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562333697] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562333921] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562334106] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562334424] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562334582] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1562334741] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1562337574] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.92%
[1562337727] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.66%
[1562337879] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.67%
[1562338559] HOST ALERT: ieatnetsimv5116-64;DOWN;SOFT;1;PING NOK
[1562338591] HOST ALERT: ieatnetsimv5116-64;UP;SOFT;1;PING OK
[1562341509] SERVICE FLAPPING ALERT: ieatnetsimv7035-29;check netsim cpu;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1562341509] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.03%
[1562345141] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.01%
[1562345294] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.29%
[1562346894] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562347063] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562351734] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562353292] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.78%
[1562353442] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.49%
[1562356247] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562356448] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1562357708] SERVICE FLAPPING ALERT: ieatnetsimv5116-18;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1562361576] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.45%
[1562361727] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 63.1%
[1562363841] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1562364010] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES