[1575936078] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1575936236] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (34) IS HIGHER THAN 20% OF ENM NODES (50)
[1575937602] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1575937754] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RNC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575942724] SERVICE ALERT: ieatnetsimv9041-18;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575942724] HOST ALERT: ieatnetsimv9041-18;DOWN;SOFT;1;PING NOK
[1575942787] HOST ALERT: ieatnetsimv9041-18;UP;SOFT;1;PING OK
[1575942875] SERVICE ALERT: ieatnetsimv9041-18;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575947674] SERVICE FLAPPING ALERT: ieatnetsimv5116-10;check netsim disk;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1575947878] SERVICE FLAPPING ALERT: ieatnetsimv5116-08;check netsim disk;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1575948823] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575949096] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1575951720] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2556) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1575952708] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (10) IS HIGHER THAN 20% OF ENM NODES (30)
[1575952832] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (7) IS HIGHER THAN 20% OF ENM NODES (30)
[1575952957] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (7) IS HIGHER THAN 20% OF ENM NODES (30)
[1575954153] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575954324] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1575954506] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1575954633] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1575954760] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1575956076] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.33%
[1575956227] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.77%
[1575956555] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures RBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575957688] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (16) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575957757] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8) IS HIGHER THAN 20% OF ENM NODES (30)
[1575958930] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (5660) IS HIGHER THAN 20% OF ENM NODES (14035)
[1575959823] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575961293] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (34) IS HIGHER THAN 20% OF ENM NODES (100)
[1575961362] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1575962554] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2620) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1575962954] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575963144] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575963330] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575964056] SERVICE ALERT: ieatwlvm9034;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1575964795] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (750)
[1575964811] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575964897] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575964930] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1575964978] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (16) IS HIGHER THAN 20% OF ENM NODES (30)
[1575964983] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (41) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575965146] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575965160] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (30)
[1575965195] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575965270] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575965298] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (11) IS HIGHER THAN 20% OF ENM NODES (30)
[1575965880] SERVICE ALERT: ieatnetsimv7048-08;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1575966000] SERVICE ALERT: ieatnetsimv7048-08;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1575966077] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (8300) IS HIGHER THAN 95% OF ENM NODES (8300)
[1575966120] SERVICE ALERT: ieatnetsimv7048-08;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1575966176] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3978) IS HIGHER THAN 20% OF ENM NODES (14035)
[1575966908] SERVICE ALERT: ieatlms5735-1;check synchronized FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (75) IS LOWER THAN 80% OF ENM  NODES (100)
[1575967034] SERVICE ALERT: ieatlms5735-1;check synchronized FRONTHAUL-6080;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (100)!
[1575967077] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575967158] SERVICE ALERT: ieatlms5735-1;check synchronized FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (75) IS LOWER THAN 80% OF ENM  NODES (100)
[1575967213] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575967287] SERVICE ALERT: ieatlms5735-1;check synchronized FRONTHAUL-6080;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (100)!
[1575967344] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (41) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575967462] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1575968504] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1575968907] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1575969159] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (49) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575970661] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575970846] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575971031] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575971066] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1575971190] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1575971376] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1575972518] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (17) IS HIGHER THAN 20% OF ENM NODES (30)
[1575972650] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (14) IS HIGHER THAN 20% OF ENM NODES (30)
[1575972780] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (10) IS HIGHER THAN 20% OF ENM NODES (30)
[1575973560] SERVICE ALERT: ieatnetsimv9041-12;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1575974980] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1575977648] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (11684) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1575977782] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13980) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1575977923] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (11444) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1575978058] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13980) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1575978181] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (34) IS LOWER THAN 80% OF ENM  NODES (50)
[1575978308] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (39) IS LOWER THAN 80% OF ENM  NODES (50)
[1575978438] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (39) IS LOWER THAN 80% OF ENM  NODES (50)
[1575978588] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1575978752] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1575978876] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1575979016] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1575979139] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1575979139] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1575979538] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575979628] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575979662] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575979755] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575979796] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575979880] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575979962] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575980001] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12596) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1575980097] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575980097] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575980143] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12515) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1575980243] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized Router6274;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1575980243] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575980286] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13983) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1575980566] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1575980716] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 22.13%
[1575981196] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575981342] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575981685] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1575981962] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575982104] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13274) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1575982148] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575982269] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13347) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1575982333] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575982748] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1575982824] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN CISCO-ASR900 NODES !
[1575983106] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MINI-LINK-6352 NODES !
[1575983110] SERVICE ALERT: ieatnetsimv9034-45;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.8%
[1575983261] SERVICE ALERT: ieatnetsimv9034-45;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575983412] SERVICE ALERT: ieatnetsimv9034-45;check netsim cpu;OK;HARD;3;OK- CPU IS USED 53.52%
[1575983414] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1575983564] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.98%
[1575983592] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (12) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575983639] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575983715] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575983789] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575983870] SERVICE ALERT: ieatnetsimv9034-105;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.9%
[1575983903] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575983940] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575984021] SERVICE ALERT: ieatnetsimv9034-105;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.98%
[1575984053] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575984172] SERVICE ALERT: ieatnetsimv9034-105;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575984204] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575984209] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13251) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1575984305] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575984359] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13983) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1575984402] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1575984455] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575984522] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1575984606] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575984643] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1575984760] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575984765] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575984841] SERVICE ALERT: ieatnetsimv9034-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1575984911] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575984915] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575984992] SERVICE ALERT: ieatnetsimv9034-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.34%
[1575985063] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575985067] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575985868] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1557) IS LOWER THAN 80% OF ENM  NODES (2560)
[1575985954] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1575985994] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2555) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1575986104] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.92%
[1575986123] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1557) IS LOWER THAN 80% OF ENM  NODES (2560)
[1575986255] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2556) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1575986255] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575986268] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1575986285] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.95%
[1575986382] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1557) IS LOWER THAN 80% OF ENM  NODES (2560)
[1575986395] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1575986441] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (150)
[1575986445] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 57.41%
[1575986529] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1557) IS LOWER THAN 80% OF ENM  NODES (2560)
[1575986534] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575986551] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1575986569] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (150)
[1575986583] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1575986671] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1557) IS LOWER THAN 80% OF ENM  NODES (2560)
[1575986677] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1575986695] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (150)
[1575986704] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1575986723] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (90) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1575986813] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1575986824] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1575986848] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (90) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1575986977] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MINI-LINK-6352 NODES !
[1575987199] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (25) IS HIGHER THAN 20% OF ENM NODES (80)
[1575987346] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;OK;HARD;3;OK- CPU IS USED 45.21%
[1575987737] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1575987803] SERVICE ALERT: ieatnetsimv9034-105;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.03%
[1575987834] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.86%
[1575988236] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;OK;HARD;3;OK- CPU IS USED 57.33%
[1575988477] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2556) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1575988528] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (217) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1575988694] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.82%
[1575988698] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.19%
[1575988965] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575989598] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575989885] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.91%
[1575990311] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN CISCO-ASR900 NODES !
[1575990425] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1575991170] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;OK;HARD;3;OK- CPU IS USED 15.19%
[1575992148] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (996) IS HIGHER THAN 20% OF ENM NODES (1200)
[1575992324] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575992475] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 3.21%
[1575992568] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1575993263] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575993448] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575993633] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575993742] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (41) IS HIGHER THAN 20% OF ENM NODES (100)
[1575993922] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (150)
[1575994050] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (150)
[1575994202] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN CISCO-ASR900 NODES !
[1575994215] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575994366] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575994408] SERVICE ALERT: ieatnetsimv9034-83;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575994516] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;OK;HARD;3;OK- CPU IS USED 70.61%
[1575994558] SERVICE ALERT: ieatnetsimv9034-83;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 11.89%
[1575994801] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575994952] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575995034] SERVICE ALERT: ieatnetsimv9034-105;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.93%
[1575995103] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575995185] SERVICE ALERT: ieatnetsimv9034-105;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.92%
[1575995335] SERVICE ALERT: ieatnetsimv9034-105;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.94%
[1575995467] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575995618] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575995628] SERVICE ALERT: ieatnetsimv9034-05;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575995751] SERVICE ALERT: ieatnetsimv9034-111;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.93%
[1575995753] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (66) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1575995769] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575995778] SERVICE ALERT: ieatnetsimv9034-05;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.24%
[1575995901] SERVICE ALERT: ieatnetsimv9034-111;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 3.96%
[1575995929] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575996080] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575996171] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575996226] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575996230] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575996377] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575996528] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575997298] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575997327] HOST ALERT: ieatnetsimv9034-42;DOWN;SOFT;1;PING NOK
[1575997350] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1575997389] HOST ALERT: ieatnetsimv9034-42;UP;SOFT;1;PING OK
[1575997396] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575997547] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575997698] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575998734] SERVICE FLAPPING ALERT: ieatnetsimv9034-115;check netsim cpu;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1575998734] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.91%
[1575998966] SERVICE FLAPPING ALERT: ieatnetsimv9034-105;check netsim cpu;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1575998966] SERVICE ALERT: ieatnetsimv9034-105;check netsim cpu;OK;HARD;3;OK- CPU IS USED 16.23%
[1575999357] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1575999399] SERVICE FLAPPING ALERT: ieatnetsimv9034-100;check netsim cpu;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1575999399] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;OK;HARD;3;OK- CPU IS USED 27.71%
[1575999861] SERVICE FLAPPING ALERT: ieatnetsimv9034-121;check netsim cpu;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1575999861] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.61%
[1576000158] SERVICE FLAPPING ALERT: ieatnetsimv9034-114;check netsim cpu;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576000158] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.07%
[1576000972] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (34) IS HIGHER THAN 20% OF ENM NODES (100)
[1576001108] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (74) IS HIGHER THAN 20% OF ENM NODES (100)
[1576001233] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (36) IS HIGHER THAN 20% OF ENM NODES (100)
[1576001328] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.27%
[1576001398] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (90) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1576001410] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (150)
[1576001528] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MINI-LINK-6352 NODES !
[1576001538] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (150)
[1576001662] SERVICE ALERT: ieatlms5735-1;check fm hb failures CISCO-ASR900;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN CISCO-ASR900 NODES !
[1576001848] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576001972] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (41) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576002098] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (38) IS LOWER THAN 80% OF ENM  NODES (50)
[1576002963] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (559) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576003092] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (927) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576003217] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (947) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576003327] HOST ALERT: ieatnetsimv9034-48;DOWN;SOFT;1;PING NOK
[1576003375] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1576003389] HOST ALERT: ieatnetsimv9034-48;UP;SOFT;1;PING OK
[1576003903] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1576004642] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576004846] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (12) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576004959] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576005048] SERVICE ALERT: ieatnetsimv9034-117;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576005110] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576005185] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1576005199] SERVICE ALERT: ieatnetsimv9034-117;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576005246] SERVICE ALERT: ieatnetsimv9034-101;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1576005255] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (90) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1576005261] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576005350] SERVICE ALERT: ieatnetsimv9034-117;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576005377] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576005378] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MINI-LINK-6352 NODES !
[1576005397] SERVICE ALERT: ieatnetsimv9034-101;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.97%
[1576005504] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (90) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1576005527] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576005547] SERVICE ALERT: ieatnetsimv9034-101;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.97%
[1576005631] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (90) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1576005678] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576005711] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576005757] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (90) IS BETWEEN 5% AND 20% OF ENM NODES (750)
[1576005841] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576005865] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576005965] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1576005965] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576006015] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576006116] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576006166] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.99%
[1576006266] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576006630] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576006668] SERVICE ALERT: ieatnetsimv9034-118;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576006781] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576006818] SERVICE ALERT: ieatnetsimv9034-118;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576006822] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1576006932] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576006969] SERVICE ALERT: ieatnetsimv9034-118;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576007389] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576007540] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576007691] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576007747] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1576007897] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.14%
[1576008050] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 92.54%
[1576008425] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576008508] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (42) IS HIGHER THAN 20% OF ENM NODES (100)
[1576008546] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576008666] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576008789] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (16) IS HIGHER THAN 20% OF ENM NODES (30)
[1576008918] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (15) IS HIGHER THAN 20% OF ENM NODES (30)
[1576008980] SERVICE ALERT: ieatnetsimv9034-117;check netsim cpu;OK;HARD;3;OK- CPU IS USED 18.31%
[1576009042] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (17) IS HIGHER THAN 20% OF ENM NODES (30)
[1576009178] SERVICE ALERT: ieatnetsimv9034-101;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.11%
[1576009308] SERVICE FLAPPING ALERT: ieatnetsimv9034-106;check netsim cpu;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1576009308] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;OK;HARD;3;OK- CPU IS USED 31.99%
[1576009378] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576009499] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576009619] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576009797] SERVICE FLAPPING ALERT: ieatnetsimv9034-59;check netsim cpu;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1576009797] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.6%
[1576010426] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (122) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576010557] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1120) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576010562] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.91%
[1576010599] SERVICE ALERT: ieatnetsimv9034-118;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.88%
[1576010682] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (107) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576011321] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.62%
[1576011682] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;OK;HARD;3;OK- CPU IS USED 32.62%
[1576012492] SERVICE FLAPPING ALERT: ieatnetsimv9034-72;check netsim cpu;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1576012492] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.81%
[1576012647] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1576012962] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MINI-LINK-6352 NODES !
[1576013220] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;OK;HARD;3;OK
[1576013497] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;OK;HARD;3;OK- CPU IS USED 19.26%
[1576014286] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (972) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576015866] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1576015866] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576016004] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (11) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576016251] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576016376] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576016502] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576016597] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1826) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576016664] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1128) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1576016770] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576016789] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1200)!
[1576016898] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576017024] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (41) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576018771] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1) IS LOWER THAN 80% OF ENM  NODES (4)
[1576018830] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (49) IS HIGHER THAN 95% OF ENM  NODES (50)
[1576018895] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (3) IS HIGHER THAN 95% OF ENM  NODES (4)
[1576018978] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576019024] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (4)
[1576019081] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (2560)
[1576019122] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (1514)
[1576019151] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (2) IS LOWER THAN 80% OF ENM  NODES (4)
[1576019164] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576019212] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2555) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1576019246] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (126) IS LOWER THAN 80% OF ENM  NODES (1514)
[1576019278] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1) IS LOWER THAN 80% OF ENM  NODES (4)
[1576019349] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576019372] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1506) IS HIGHER THAN 95% OF ENM  NODES (1514)
[1576019609] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (80)
[1576020210] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (6131) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576020493] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13224) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576020633] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13983) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576021099] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (4)!
[1576021215] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576021492] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (15) IS LOWER THAN 5% OF ENM NODES (1200)