[1565650843] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1565651881] SERVICE ALERT: ieatnetsimv5116-54;check netsim disk;OK;HARD;3;OK
[1565651942] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565652072] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565652203] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565652924] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.85%
[1565653074] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.59%
[1565653225] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.19%
[1565653375] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.22%
[1565653526] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.08%
[1565653677] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.27%
[1565653827] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 93.46%
[1565659082] SERVICE ALERT: ieatnetsimv5116-54;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565659202] SERVICE ALERT: ieatnetsimv5116-54;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565659323] SERVICE ALERT: ieatnetsimv5116-54;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565661058] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim cpu;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1565661058] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.47%
[1565664689] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94%
[1565664839] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.83%
[1565664993] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.33%
[1565668623] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.43%
[1565670399] SERVICE ALERT: ieatnetsimv7035-105;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.99%
[1565670550] SERVICE ALERT: ieatnetsimv7035-105;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.95%
[1565671080] SERVICE ALERT: ieatnetsimv7035-107;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.59%
[1565671231] SERVICE ALERT: ieatnetsimv7035-107;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 8.19%
[1565672775] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565673221] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565673237] SERVICE ALERT: ieatnetsimv5116-84;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565673359] SERVICE ALERT: ieatnetsimv5116-84;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565673480] SERVICE ALERT: ieatnetsimv5116-84;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565675853] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.53%
[1565676004] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.91%
[1565676155] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 92.39%
[1565677143] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1565677743] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1565677826] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565679609] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1565679800] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1565679985] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1565680438] SERVICE ALERT: ieatnetsimv5116-24;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565680438] HOST ALERT: ieatnetsimv5116-24;DOWN;SOFT;1;PING NOK
[1565680500] HOST ALERT: ieatnetsimv5116-24;UP;SOFT;1;PING OK
[1565680558] SERVICE ALERT: ieatnetsimv5116-24;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565680679] SERVICE ALERT: ieatnetsimv5116-24;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565680751] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1565680879] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1565681007] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1565682818] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565684037] HOST ALERT: ieatnetsimv5116-84;DOWN;SOFT;1;PING NOK
[1565684100] HOST ALERT: ieatnetsimv5116-84;UP;SOFT;1;PING OK
[1565684279] SERVICE ALERT: ieatnetsimv5116-24;check netsim disk;OK;HARD;3;OK
[1565684524] SERVICE FLAPPING ALERT: ieatnetsimv5116-54;check netsim disk;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1565684524] SERVICE ALERT: ieatnetsimv5116-54;check netsim disk;OK;HARD;3;OK
[1565689715] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565689818] SERVICE ALERT: ieatnetsimv7035-99;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565689850] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565689956] SERVICE ALERT: ieatnetsimv7035-99;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565690087] SERVICE ALERT: ieatnetsimv7035-99;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565690676] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13526) IS HIGHER THAN 95% OF ENM COMECIM NODES (14149)
[1565691606] SERVICE ALERT: ieatnetsimv7035-92;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565691757] SERVICE ALERT: ieatnetsimv7035-92;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565693691] SERVICE ALERT: ieatnetsimv7035-95;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565693826] SERVICE ALERT: ieatnetsimv7035-95;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565693955] SERVICE ALERT: ieatnetsimv7035-95;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565694136] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565695057] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.43%
[1565695131] SERVICE ALERT: ieatnetsimv7035-64;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565695208] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.44%
[1565695262] SERVICE ALERT: ieatnetsimv7035-64;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565695359] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 69.26%
[1565695690] SERVICE ALERT: ieatnetsimv7035-65;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565695900] SERVICE ALERT: ieatnetsimv7035-105;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.06%
[1565696051] SERVICE ALERT: ieatnetsimv7035-105;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 30.62%
[1565696925] SERVICE ALERT: ieatnetsimv7035-62;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565697053] SERVICE ALERT: ieatnetsimv7035-62;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565697186] SERVICE ALERT: ieatnetsimv7035-62;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565698438] SERVICE ALERT: ieatnetsimv5116-22;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565698559] SERVICE ALERT: ieatnetsimv5116-22;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565698679] SERVICE ALERT: ieatnetsimv5116-22;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565698681] SERVICE ALERT: ieatnetsimv5116-84;check netsim disk;OK;HARD;3;OK
[1565698989] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.18%
[1565699140] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 38.79%
[1565699483] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565699667] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565699674] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1565699853] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565701378] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565701568] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565701782] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14031) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565702369] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.51%
[1565702520] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.84%
[1565704415] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1565705265] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1565707130] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1565707316] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1565707501] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1565707501] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1565708023] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1565708151] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1565708279] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1565708279] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1565708585] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.82%
[1565709053] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565709236] SERVICE ALERT: ieatnetsimv5116-35;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565709238] HOST ALERT: ieatnetsimv5116-22;DOWN;SOFT;1;PING NOK
[1565709243] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565709270] HOST ALERT: ieatnetsimv5116-22;UP;SOFT;1;PING OK
[1565709357] SERVICE ALERT: ieatnetsimv5116-35;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565709412] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565709477] SERVICE ALERT: ieatnetsimv5116-35;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565710087] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1565710683] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2515) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565710896] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1565711167] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1565711393] SERVICE ALERT: ieatnetsimv7035-85;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565711548] SERVICE ALERT: ieatnetsimv7035-85;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565711708] SERVICE ALERT: ieatnetsimv7035-85;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565711829] SERVICE ALERT: ieatnetsimv7035-26;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565711829] SERVICE ALERT: ieatnetsimv7035-85;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565712074] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1565712218] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.47%
[1565712369] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.01%
[1565712757] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1565712938] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1565713118] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1565713645] SERVICE ALERT: ieatnetsimv7035-26;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565713806] SERVICE ALERT: ieatnetsimv7035-26;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565713989] SERVICE ALERT: ieatnetsimv7035-26;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565714289] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565714416] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2515) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565714440] HOST ALERT: ieatnetsimv5116-50;DOWN;SOFT;1;PING NOK
[1565714502] HOST ALERT: ieatnetsimv5116-50;UP;SOFT;1;PING OK
[1565715681] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1565715808] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1565715934] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1565716634] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14031) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565717407] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1565717605] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1565717801] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1565720036] HOST ALERT: ieatnetsimv5116-35;DOWN;SOFT;1;PING NOK
[1565720069] HOST ALERT: ieatnetsimv5116-35;UP;SOFT;1;PING OK
[1565721477] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1565727312] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565727495] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565727568] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1565727677] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565727750] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1565727931] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1565729651] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.2%
[1565729802] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 61.41%
[1565731325] SERVICE FLAPPING ALERT: ieatnetsimv5116-54;check netsim disk;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1565733139] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1565733392] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED