[1563324239] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563324427] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1563324612] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1563324614] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1563324749] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1563324879] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1563330484] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563330792] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2370) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563330931] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2413) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1563331201] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563332990] SERVICE FLAPPING ALERT: ieatlms3901;check node fs hc;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1563335722] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1563336449] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563337369] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563338078] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563343814] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1563343990] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6754) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1563344133] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6755) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1563344276] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6752) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1563344811] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.97%
[1563344961] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 55.44%
[1563345679] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563345864] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1563346050] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1563346095] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (66) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1563349716] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1563356269] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1563356434] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1563356443] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.3%
[1563356593] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 9.38%
[1563360347] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1563360501] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1563360598] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563360654] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1563361317] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563362022] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1563362262] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1563363182] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563363400] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1563363535] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1563364285] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;OK;HARD;3;OK- CPU IS USED 23.3%
[1563365981] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563366167] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1563370848] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (1365) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563370974] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (1365) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563371099] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1365) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563371752] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563371938] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1563374705] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1453) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563378347] SERVICE ALERT: ieatnetsimv7035-02;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1563378476] SERVICE ALERT: ieatnetsimv7035-02;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1563378605] SERVICE ALERT: ieatnetsimv7035-02;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1563378999] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563379137] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1563379179] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563379318] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1563379360] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563379498] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1563380797] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13103) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563380940] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13103) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563381082] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14035) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563384704] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12710) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563384821] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1563384845] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13039) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563384959] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1563384987] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13047) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563387614] SERVICE ALERT: ieatnetsimv7035-02;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1563387771] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1563387965] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1563388158] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1563388323] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563388509] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1563388608] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14035) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563389126] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (1352) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563389150] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563389253] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (1357) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563389331] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563389380] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1398) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563389398] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (114) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563389512] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1563390032] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1563390591] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1563390772] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1563390953] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1563391375] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563391556] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563391737] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563392987] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RBS;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1563392987] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563393005] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563403030] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12695) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563403160] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563403301] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12695) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)