[1574900344] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;OK;HARD;3;OK
[1574900396] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1574900396] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1574900396] SERVICE FLAPPING ALERT: ieatlms5735-1;check nas hc;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1574900396] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1574900396] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1574900396] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1574900433] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574900455] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574900550] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574900553] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1200)
[1574900558] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574900579] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574900596] HOST ALERT: ieatnetsimv5116-85;DOWN;SOFT;1;PING NOK
[1574900596] HOST ALERT: ieatnetsimv5116-80;DOWN;SOFT;1;PING NOK
[1574900658] HOST ALERT: ieatnetsimv5116-80;UP;SOFT;1;PING OK
[1574900658] HOST ALERT: ieatnetsimv5116-85;UP;SOFT;1;PING OK
[1574900676] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574900677] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1574900683] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574900708] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574900709] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574900761] SERVICE ALERT: ieatlms7555;check pm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (7500) IS HIGHER THAN 95% OF ENM NODES (7500)
[1574900802] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574900833] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574900889] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20*1=20)
[1574900896] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1750*4=7000)
[1574900899] SERVICE ALERT: ieatlms5735-1;check pm collection statistical EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574900959] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574901012] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574901013] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20*1=20)
[1574901023] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1750*4=7000)
[1574901025] SERVICE ALERT: ieatlms5735-1;check pm collection statistical EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574901136] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574901137] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20*1=20)
[1574901149] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1750*4=7000)
[1574901150] SERVICE ALERT: ieatlms5735-1;check pm collection statistical EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574901187] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574901261] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574901313] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574901333] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12*1=12)
[1574901437] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574901456] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (7500*4=30000)
[1574901457] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12*1=12)
[1574901582] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12*1=12)
[1574901584] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (7500*4=30000)
[1574901609] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574901619] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574901622] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574901636] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574901725] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (7500*4=30000)
[1574901734] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574901743] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1574901744] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574901748] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574901761] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574901783] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13*1=13)
[1574901787] SERVICE ALERT: ieatlms5735-1;check pm collection statistical BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (98*1=98)
[1574901871] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574901872] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574901874] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574901886] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574901908] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13*1=13)
[1574901911] SERVICE ALERT: ieatlms5735-1;check pm collection statistical BSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (98*1=98)
[1574902027] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935*1=13935)
[1574902032] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13*1=13)
[1574902042] SERVICE ALERT: ieatlms5735-1;check pm collection statistical BSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (98*1=98)
[1574902052] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574902062] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574902158] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935*1=13935)
[1574902176] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574902186] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574902288] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935*1=13935)
[1574902302] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574902311] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574902931] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514*1=1514)
[1574902971] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574903058] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514*1=1514)
[1574903096] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574903183] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514*1=1514)
[1574903211] HOST ALERT: ieatnetsimv7035-38;DOWN;SOFT;1;PING NOK
[1574903211] HOST ALERT: ieatnetsimv7035-15;DOWN;SOFT;1;PING NOK
[1574903221] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1574903221] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574903273] HOST ALERT: ieatnetsimv7035-38;UP;SOFT;1;PING OK
[1574903273] HOST ALERT: ieatnetsimv7035-15;UP;SOFT;1;PING OK
[1574903944] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574904065] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574904185] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574904196] HOST ALERT: ieatnetsimv5116-95;DOWN;SOFT;1;PING NOK
[1574904258] HOST ALERT: ieatnetsimv5116-95;UP;SOFT;1;PING OK
[1574904561] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574905628] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (6029) IS FROM 80% TO 95% OF ENM  NODES (7500)
[1574905910] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (6097) IS BETWEEN 80% AND 95% OF ENM NODES (7500)
[1574905948] SERVICE ALERT: ieatlms7555;check cm supervision nodes MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (6101) IS BETWEEN 80% AND 95% OF ENM NODES (7500)
[1574905966] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.79%
[1574906117] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 60.63%
[1574908219] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574908398] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574908576] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574912740] SERVICE ALERT: ieatnetsimv017-17;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.05%
[1574912891] SERVICE ALERT: ieatnetsimv017-17;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.09%
[1574914917] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (7132) IS HIGHER THAN 95% OF ENM NODES (7500)
[1574914953] HOST ALERT: ieatnetsimv7035-05;DOWN;SOFT;1;PING NOK
[1574914954] SERVICE ALERT: ieatlms7555;check cm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (7135) IS HIGHER THAN 95% OF ENM NODES (7500)
[1574915015] HOST ALERT: ieatnetsimv7035-05;UP;SOFT;1;PING OK
[1574915446] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574915567] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;SOFT;2;OK
[1574916434] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (7291) IS HIGHER THAN 95% OF ENM  NODES (7500)
[1574919480] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1574919480] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574919892] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;OK;HARD;3;OK- NUMBER OF PM FILES (1499) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574920036] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13206) IS FROM 80% TO 95% OF ENM  NODES (14013)
[1574920040] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574920161] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;HARD;3;OK
[1574920175] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (11096) IS LOWER THAN 80% OF ENM  NODES (14013)
[1574920177] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574920316] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (71) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574920317] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13734) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935*1=13935)
[1574920324] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (7705) IS LOWER THAN 80% OF ENM  NODES (14013)
[1574920700] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (15) IS LOWER THAN 80% OF ENM  NODES (50)
[1574920707] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (14) IS LOWER THAN 80% OF ENM  NODES (30)
[1574920827] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (27) IS LOWER THAN 80% OF ENM  NODES (50)
[1574920831] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (20) IS LOWER THAN 80% OF ENM  NODES (30)
[1574920951] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (38) IS LOWER THAN 80% OF ENM  NODES (50)
[1574920955] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (20) IS LOWER THAN 80% OF ENM  NODES (30)
[1574920996] HOST ALERT: ieatnetsimv5116-65;DOWN;SOFT;1;PING NOK
[1574921058] HOST ALERT: ieatnetsimv5116-65;UP;SOFT;1;PING OK
[1574921232] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574921549] SERVICE ALERT: ieatnetsimv7035-93;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.08%
[1574921699] SERVICE ALERT: ieatnetsimv7035-93;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.03%
[1574921928] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574921977] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574922052] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574922114] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;SOFT;2;OK
[1574922120] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574922153] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13638) IS HIGHER THAN 95% OF ENM  NODES (14013)
[1574922250] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574922288] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574922386] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;SOFT;2;OK
[1574922408] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574922566] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;OK;HARD;3;OK- NUMBER OF PM FILES (1199) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574922759] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574922761] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1574922822] SERVICE ALERT: ieatnetsimv7035-93;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574922865] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (22) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574922997] SERVICE ALERT: ieatnetsimv7035-93;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574923042] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574923162] SERVICE ALERT: ieatnetsimv7035-93;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574923484] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (3) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574923486] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574923637] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-IS;OK;HARD;3;OK- NUMBER OF PM FILES (72) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13*1=13)
[1574923907] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (100) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574923916] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;OK;HARD;3;OK- NUMBER OF PM FILES (50) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574924268] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.6%
[1574924419] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.36%
[1574924569] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 60.23%
[1574924854] HOST ALERT: ieatnetsimv7035-110;DOWN;SOFT;1;PING NOK
[1574924916] HOST ALERT: ieatnetsimv7035-110;UP;SOFT;1;PING OK
[1574925488] SERVICE ALERT: ieatnetsimv7048-01;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574925495] SERVICE ALERT: ieatnetsimv5116-87;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.92%
[1574925498] SERVICE ALERT: ieatnetsimv7048-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574925506] SERVICE ALERT: ieatnetsimv7048-03;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.81%
[1574925509] SERVICE ALERT: ieatnetsimv7048-07;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.64%
[1574925615] SERVICE ALERT: ieatnetsimv7048-01;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574925624] SERVICE ALERT: ieatnetsimv7048-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574925646] SERVICE ALERT: ieatnetsimv5116-87;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574925656] SERVICE ALERT: ieatnetsimv7048-03;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.92%
[1574925659] SERVICE ALERT: ieatnetsimv7048-07;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.66%
[1574925758] SERVICE ALERT: ieatnetsimv7048-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574925797] SERVICE ALERT: ieatnetsimv5116-87;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.55%
[1574925805] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (12545) IS LOWER THAN 80% OF ENM NODES (20060)
[1574925841] SERVICE ALERT: ieatlms7435;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1574925885] SERVICE ALERT: ieatnetsimv7048-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574925939] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (12545) IS LOWER THAN 80% OF ENM NODES (20060)
[1574926029] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1574926073] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (12545) IS LOWER THAN 80% OF ENM NODES (20060)
[1574926429] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (20)
[1574926487] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574926551] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (20)
[1574926568] SERVICE ALERT: ieatlms7435;check node fs;OK;HARD;3;OK
[1574926608] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574926674] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (20)
[1574926728] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574926856] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574927015] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574927159] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574928558] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574928687] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574928815] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574929816] SERVICE ALERT: ieatnetsimv7102-02;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.4%
[1574929966] SERVICE ALERT: ieatnetsimv7102-02;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.33%
[1574930470] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1574930470] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1574930470] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1574930622] SERVICE ALERT: ieatwlvm5116;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1574931156] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574931229] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574931344] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574931403] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1574931531] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574931533] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1574931663] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1574932234] SERVICE ALERT: ieatnetsimv017-04;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1574932385] SERVICE ALERT: ieatnetsimv017-04;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 1.03%
[1574933473] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1574933929] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1574934799] SERVICE ALERT: ieatnetsimv5116-07;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574934950] SERVICE ALERT: ieatnetsimv5116-07;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574934972] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935004] SERVICE ALERT: ieatnetsimv5116-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935071] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935090] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935101] SERVICE ALERT: ieatnetsimv5116-07;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574935116] SERVICE ALERT: ieatlms7435;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574935127] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574935156] SERVICE ALERT: ieatnetsimv5116-82;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574935198] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574935223] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574935251] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574935283] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574935309] SERVICE ALERT: ieatnetsimv5116-82;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574935344] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935375] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574935419] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574935455] SERVICE ALERT: ieatnetsimv5116-85;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935496] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574935520] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935608] SERVICE ALERT: ieatnetsimv5116-85;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574935649] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574935672] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574935761] SERVICE ALERT: ieatnetsimv5116-85;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574935767] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935769] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935794] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935824] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574935844] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935918] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574935921] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574935946] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574935970] SERVICE ALERT: ieatnetsimv5116-69;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574935995] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 11.26%
[1574936070] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574936072] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574936098] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574936121] SERVICE ALERT: ieatnetsimv5116-69;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574936217] HOST ALERT: ieatnetsimv9034-67;DOWN;SOFT;1;PING NOK
[1574936272] SERVICE ALERT: ieatnetsimv5116-69;check netsim cpu;OK;HARD;3;OK- CPU IS USED 22.09%
[1574936279] HOST ALERT: ieatnetsimv9034-67;UP;SOFT;1;PING OK
[1574936294] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574936445] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574936596] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574936729] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574936880] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574937035] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574937203] SERVICE ALERT: ieatnetsimv7035-67;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574937360] SERVICE ALERT: ieatnetsimv7035-67;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574937496] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574937504] SERVICE ALERT: ieatnetsimv7035-67;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574937516] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574937517] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574937647] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574937648] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574937667] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574937669] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574937799] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574937799] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574937818] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574937821] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574937893] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574937950] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574937968] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574938044] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574938119] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574938195] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574938270] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574938416] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574938467] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574938492] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574938513] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574938568] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574938618] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574938643] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574938664] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574938720] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574938731] SERVICE ALERT: ieatnetsimv5116-07;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.65%
[1574938769] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574938795] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574938913] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.51%
[1574938939] SERVICE ALERT: ieatnetsimv5116-82;check netsim cpu;OK;HARD;3;OK- CPU IS USED 38.37%
[1574939049] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;OK;HARD;3;OK- CPU IS USED 15.71%
[1574939143] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574939281] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;HARD;3;OK- CPU IS USED 18.45%
[1574939290] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574939342] SERVICE ALERT: ieatnetsimv7035-67;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574939391] SERVICE ALERT: ieatnetsimv5116-85;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.04%
[1574939443] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574939703] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.98%
[1574939729] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 53.82%
[1574940126] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574940171] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13247) IS FROM 80% TO 95% OF ENM  NODES (14013)
[1574940260] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574940322] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13371) IS HIGHER THAN 95% OF ENM  NODES (14013)
[1574940385] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (76) IS HIGHER THAN 95% OF ENM  NODES (80)
[1574940429] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574940615] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574940797] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574940806] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574940992] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574941188] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574942305] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2404) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1574942442] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2391) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1574942572] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2355) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1574942671] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574944055] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13174) IS FROM 80% TO 95% OF ENM  NODES (14013)
[1574944190] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13188) IS FROM 80% TO 95% OF ENM  NODES (14013)
[1574944322] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (13221) IS FROM 80% TO 95% OF ENM  NODES (14013)
[1574944863] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574945282] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (2558) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574946134] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1574946134] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13679) IS HIGHER THAN 95% OF ENM  NODES (14013)
[1574946280] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.57%
[1574946430] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 30.17%
[1574946901] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;OK;HARD;3;OK- CPU IS USED 27.17%
[1574947427] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.55%
[1574947866] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.69%
[1574948648] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.25%
[1574948781] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.94%
[1574949626] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 52.73%
[1574949853] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1574949883] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0*4=0)
[1574949920] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1574949933] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1574950215] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1574950245] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (7500*4=30000)
[1574950370] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (7500*4=30000)
[1574950496] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (7500*4=30000)
[1574951012] SERVICE ALERT: ieatlms3901;check unsync nodes;OK;HARD;3;OK
[1574951725] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (750) IS LOWER THAN 80% OF ENM NODES (1021)
[1574951739] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (750) IS LOWER THAN 80% OF ENM NODES (1021)
[1574951856] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (750) IS LOWER THAN 80% OF ENM NODES (1021)
[1574951867] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (750) IS LOWER THAN 80% OF ENM NODES (1021)
[1574951929] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574951938] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574951982] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (750) IS LOWER THAN 80% OF ENM NODES (1021)
[1574951993] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (750) IS LOWER THAN 80% OF ENM NODES (1021)
[1574952050] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574952117] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574952149] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (11) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1574952151] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (877) IS LOWER THAN 80% OF ENM  NODES (1200)
[1574952170] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1574952170] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574952230] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.2%
[1574952251] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 28.9%
[1574952279] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (994) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1574952295] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1574952295] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574952410] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1162) IS HIGHER THAN 95% OF ENM  NODES (1200)
[1574952625] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.43%
[1574952861] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574953026] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574953192] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574953200] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.37%
[1574953537] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574953723] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574953855] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.41%
[1574953908] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574954201] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (80) IS FROM 80% TO 95% OF ENM  NODES (98)
[1574954215] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1104) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1574954277] SERVICE ALERT: ieatnetsimv7048-05;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574954325] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (97) IS HIGHER THAN 95% OF ENM  NODES (98)
[1574954341] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (927) IS LOWER THAN 80% OF ENM  NODES (1200)
[1574954428] SERVICE ALERT: ieatnetsimv7048-05;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574954466] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1125) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1574954676] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1574954720] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (619) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574955396] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1574955577] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1574955744] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574955758] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (22) IS HIGHER THAN 20% OF ENM NODES (80)
[1574955758] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1574955771] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1574955864] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574955984] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574956134] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574956272] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1200)!
[1574956301] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 77.18%
[1574956316] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574956321] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (37) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574956452] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (37) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574956497] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574956580] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (37) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574956682] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574956864] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574956966] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574957005] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.67%
[1574957091] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574957217] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574957957] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1574958076] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1071) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1574958103] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1300) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574958204] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1153) IS HIGHER THAN 95% OF ENM  NODES (1200)
[1574958230] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1315) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574958298] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (100) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574958325] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (484) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574958357] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1333) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574958426] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (100) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574958482] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (549) IS LOWER THAN 80% OF ENM  NODES (750)
[1574958557] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (100) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574958579] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (829) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574958704] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (829) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574958829] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (829) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574958870] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (20) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574959002] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (20) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574959126] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574959917] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (76) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574960047] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (76) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574960164] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1638) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574960172] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (76) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574960483] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (647) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574961120] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN RadioNode NODES !
[1574961975] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1749) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574963510] HOST ALERT: ieatnetsimv7035-47;DOWN;SOFT;1;PING NOK
[1574963572] HOST ALERT: ieatnetsimv7035-47;UP;SOFT;1;PING OK
[1574964246] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574964428] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574964524] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (18) IS LOWER THAN 5% OF ENM NODES (1200)
[1574964609] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574965161] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1574965531] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1574965531] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (641) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574965785] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574965785] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574966003] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (100)
[1574966034] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6672;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574966034] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;OK;HARD;3;OK- NUMBER OF PM FILES (1199) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (NO. OF EXPECTED FILES: 1200)
[1574966471] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574966562] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (80)
[1574966897] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574967276] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1574967378] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574967378] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (100) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574967385] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6274;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1574967385] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;OK;HARD;3;OK- NUMBER OF PM FILES (50) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574967403] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1574967532] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1574967550] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.89%
[1574967782] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574967817] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1574968506] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1574969138] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (597) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574969340] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1574970133] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574970314] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574970496] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574972747] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (624) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574973240] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (858) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574973384] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (858) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574973509] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574973771] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574973891] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574974012] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574974587] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (61) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574974592] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (39) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574974719] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (39) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574974722] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (61) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574974844] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (24) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574974847] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (33) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574974887] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (507) IS LOWER THAN 80% OF ENM  NODES (750)
[1574975192] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (105) IS LOWER THAN 80% OF ENM  NODES (150)
[1574975316] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (148) IS HIGHER THAN 95% OF ENM  NODES (150)
[1574975756] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574975882] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574976008] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574976329] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (988) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1574976356] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (214) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574976380] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1499) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574976451] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (26) IS LOWER THAN 80% OF ENM  NODES (50)
[1574976456] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1041) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1574976506] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1522) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574976577] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (42) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574976582] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1101) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1574976596] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (77) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574976632] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1574976632] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1543) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574976701] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1574976701] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (624) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574976709] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574976724] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (77) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574976851] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (77) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574977757] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1574977757] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574978387] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized Router6672;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1574978387] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1200)!
[1574978438] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1748) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574978518] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574979958] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1574980348] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574980974] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574981020] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1574981196] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574981213] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1574982044] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1386) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574982107] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (545) IS LOWER THAN 80% OF ENM  NODES (750)
[1574982148] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12977) IS FROM 80% TO 95% OF ENM  NODES (14013)
[1574982170] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1406) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574982281] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13982) IS HIGHER THAN 95% OF ENM  NODES (14013)
[1574982300] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1428) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574983912] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (655) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574984108] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1743) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574984539] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574984735] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574984931] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED