[1573344203] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1573344896] HOST ALERT: ieatrcxb3778;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573344986] HOST ALERT: ieatrcxb3778;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573345076] HOST ALERT: ieatrcxb3778;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1573345140] HOST ALERT: ieatrcxb3778;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1573345673] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573345797] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573345923] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573347527] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1573347602] HOST ALERT: ieatrcxb3911;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573347692] HOST ALERT: ieatrcxb3911;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573347719] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1573347782] HOST ALERT: ieatrcxb3911;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1573347817] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN ERBS NODES !
[1573347846] HOST ALERT: ieatrcxb3911;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1573347905] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1573348751] SERVICE ALERT: ieatlms3901;check synchronized FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (83) IS FROM 80% TO 95% OF ENM  NODES (100)
[1573348872] SERVICE ALERT: ieatlms3901;check synchronized Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (900) IS LOWER THAN 80% OF ENM  NODES (1200)
[1573348880] SERVICE ALERT: ieatlms3901;check synchronized FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (83) IS FROM 80% TO 95% OF ENM  NODES (100)
[1573348999] SERVICE ALERT: ieatlms3901;check synchronized Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (900) IS LOWER THAN 80% OF ENM  NODES (1200)
[1573349005] SERVICE ALERT: ieatlms3901;check synchronized FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (83) IS FROM 80% TO 95% OF ENM  NODES (100)
[1573349125] SERVICE ALERT: ieatlms3901;check synchronized Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (958) IS LOWER THAN 80% OF ENM  NODES (1200)
[1573349768] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573349932] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573350811] SERVICE ALERT: ieatlms3901;check synchronized FRONTHAUL-6080;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (100)!
[1573350815] SERVICE FLAPPING ALERT: ieatlms3901;check hw resources hc;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1573350815] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1573350932] SERVICE ALERT: ieatlms3901;check synchronized Router6672;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1199) IS HIGHER THAN 95% OF ENM  NODES (1200)
[1573351131] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (580) IS LOWER THAN 80% OF ENM  NODES (750)
[1573351214] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2448) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573351256] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (580) IS LOWER THAN 80% OF ENM  NODES (750)
[1573351382] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (580) IS LOWER THAN 80% OF ENM  NODES (750)
[1573353129] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573353189] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1573353257] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573353387] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573353673] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12315) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573353816] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12315) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573353957] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12315) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573356992] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573357578] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13695) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573361200] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12317) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573361342] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12317) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573361484] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12317) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573364199] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573364324] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573364448] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573368023] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1573368054] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573368707] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13336) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573370739] SERVICE ALERT: ieatlms3901;check synchronized Router6672;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1024) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1573370868] SERVICE ALERT: ieatlms3901;check synchronized Router6672;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1067) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1573370995] SERVICE ALERT: ieatlms3901;check synchronized Router6672;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1090) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1573372161] SERVICE FLAPPING ALERT: ieatlms3901;check unsync nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573372804] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized Router6672;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1573372804] SERVICE ALERT: ieatlms3901;check synchronized Router6672;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1200)!
[1573379327] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized Router6672;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1573379327] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (31.7% change >= 20.0% threshold)
[1573379327] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (32.2% change >= 20.0% threshold)
[1573383073] HOST ALERT: ieatnetsimv5116-114;DOWN;SOFT;1;PING NOK
[1573383129] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12333) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573383135] HOST ALERT: ieatnetsimv5116-12;DOWN;SOFT;1;PING NOK
[1573383135] HOST ALERT: ieatnetsimv017-14;DOWN;SOFT;1;PING NOK
[1573383135] HOST ALERT: ieatnetsimv5116-49;DOWN;SOFT;1;PING NOK
[1573383135] HOST ALERT: ieatnetsimv5116-114;UP;SOFT;1;PING OK
[1573383197] HOST ALERT: ieatnetsimv5116-49;UP;SOFT;1;PING OK
[1573383197] HOST ALERT: ieatnetsimv017-14;UP;SOFT;1;PING OK
[1573383197] HOST ALERT: ieatnetsimv5116-12;UP;SOFT;1;PING OK
[1573383273] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12333) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573383415] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12333) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573385992] SERVICE ALERT: ieatlms3901;check cm supervision nodes RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (3849) IS LOWER THAN 80% OF ENM NODES (11547)
[1573386059] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573386122] SERVICE ALERT: ieatlms3901;check cm supervision nodes RBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (3849) IS HIGHER THAN 95% OF ENM NODES (3849)
[1573386184] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573386310] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573387035] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13392) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573387837] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1573388015] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1573388231] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1573388444] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.71%
[1573388594] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 58.41%
[1573389915] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (143) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573392345] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.71%
[1573392496] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.1%
[1573393034] HOST ALERT: ieatnetsimv5116-40;DOWN;SOFT;1;PING NOK
[1573393096] HOST ALERT: ieatnetsimv5116-40;UP;SOFT;1;PING OK
[1573397120] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573397245] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573397371] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (145) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573397858] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13096) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573398000] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13096) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573398142] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12328) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573400976] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573401103] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573401233] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573401764] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13695) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1573404843] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573406834] HOST ALERT: ieatnetsimv5116-27;DOWN;SOFT;1;PING NOK
[1573406896] HOST ALERT: ieatnetsimv5116-27;UP;SOFT;1;PING OK
[1573407011] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized Router6672;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573407434] HOST ALERT: ieatnetsimv5116-41;DOWN;SOFT;1;PING NOK
[1573407496] HOST ALERT: ieatnetsimv5116-41;UP;SOFT;1;PING OK
[1573408334] HOST ALERT: ieatnetsimv5116-15;DOWN;SOFT;1;PING NOK
[1573408396] HOST ALERT: ieatnetsimv5116-15;UP;SOFT;1;PING OK
[1573410134] HOST ALERT: ieatnetsimv5116-47;DOWN;SOFT;1;PING NOK
[1573410134] HOST ALERT: ieatnetsimv5116-09;DOWN;SOFT;1;PING NOK
[1573410197] HOST ALERT: ieatnetsimv5116-09;UP;SOFT;1;PING OK
[1573410197] HOST ALERT: ieatnetsimv5116-47;UP;SOFT;1;PING OK
[1573412604] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1573412724] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1573412848] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1573420860] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.71%
[1573421012] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.99%
[1573421162] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;OK;HARD;3;OK- CPU IS USED 48.25%