[1572045980] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1572046241] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;SOFT;2;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1572048850] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572048974] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572049100] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572052233] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13761) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14001)
[1572052470] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572052595] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572052720] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572055689] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23685) IS BETWEEN 80% AND 95% OF ENM NODES (26997)
[1572055842] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14001)
[1572055985] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (13761) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14001)
[1572055997] SERVICE ALERT: ieatlms3901;check pm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (26245) IS HIGHER THAN 95% OF ENM NODES (26997)
[1572056306] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572057005] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572057132] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572059714] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13966)
[1572059844] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13961)
[1572059977] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13953)
[1572061322] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (23146) IS BETWEEN 80% AND 95% OF ENM NODES (26996)
[1572061557] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (26995) IS HIGHER THAN 95% OF ENM NODES (26996)
[1572063512] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572063526] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572063639] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572063763] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572063888] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572064014] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572064140] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572064267] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572067872] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572068000] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572068059] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572068185] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572068311] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572068437] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572068564] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572074398] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13711) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572075325] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572075450] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572075576] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572075701] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572078008] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572078149] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (13711) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572078279] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572078420] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572078562] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572081532] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572081656] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572081781] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572082971] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572083027] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572083096] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572083151] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572083222] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572083277] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572083402] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572083527] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572083652] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572083776] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572085386] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572086828] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572088992] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572089118] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572089243] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572089374] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572089499] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572090135] SERVICE ALERT: ieatnetsimv7035-07;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572090286] SERVICE ALERT: ieatnetsimv7035-07;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572090434] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572090437] SERVICE ALERT: ieatnetsimv7035-07;check netsim cpu;OK;HARD;3;OK- CPU IS USED 72.14%
[1572090561] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572092984] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13711) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572093823] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093826] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093830] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093841] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093844] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093848] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093850] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093861] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093863] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093880] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093884] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093887] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093888] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093893] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093897] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093900] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093904] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093906] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093909] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093912] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093914] SERVICE ALERT: ieatnetsimv7035-88;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1572093932] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572093974] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572093978] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572093981] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.92%
[1572093991] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 35.87%
[1572093995] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 52.09%
[1572094000] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094001] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094013] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094014] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094032] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094037] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094039] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094040] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094044] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094049] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094051] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 53.99%
[1572094055] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094056] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 59.8%
[1572094059] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.62%
[1572094063] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 8.58%
[1572094065] SERVICE ALERT: ieatnetsimv7035-88;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.48%
[1572094085] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094125] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094129] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094131] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;OK;HARD;3;OK- CPU IS USED 27.62%
[1572094150] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.22%
[1572094151] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.93%
[1572094165] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094166] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094182] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.86%
[1572094188] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094190] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094190] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;OK;HARD;3;OK- CPU IS USED 16.63%
[1572094195] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.38%
[1572094196] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572094200] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;OK;HARD;3;OK- CPU IS USED 17.05%
[1572094200] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572094206] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572094206] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.85%
[1572094206] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572094239] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094287] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572094348] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094352] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094357] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094358] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572094414] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572094499] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094503] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094509] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094509] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572094582] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572094706] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572094831] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572096606] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572096735] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572096877] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13711) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572097322] SERVICE ALERT: ieatnetsimv7035-01;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572097324] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572097326] SERVICE ALERT: ieatnetsimv7035-03;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572097329] SERVICE ALERT: ieatnetsimv7035-04;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572097473] SERVICE ALERT: ieatnetsimv7035-01;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572097475] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572097477] SERVICE ALERT: ieatnetsimv7035-03;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572097480] SERVICE ALERT: ieatnetsimv7035-04;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572097624] SERVICE ALERT: ieatnetsimv7035-01;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572097626] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572097628] SERVICE ALERT: ieatnetsimv7035-03;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572097630] SERVICE ALERT: ieatnetsimv7035-04;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572097869] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;OK;HARD;3;OK- CPU IS USED 44.29%
[1572098438] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572098563] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572098688] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572101254] SERVICE ALERT: ieatnetsimv7035-01;check netsim cpu;OK;HARD;3;OK- CPU IS USED 1.9%
[1572101261] SERVICE ALERT: ieatnetsimv7035-04;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.66%
[1572104858] SERVICE ALERT: ieatnetsimv7035-03;check netsim cpu;OK;HARD;3;OK- CPU IS USED 1.66%
[1572104956] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 30.59%
[1572104960] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.59%
[1572105341] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572105467] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572105595] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572107687] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572107816] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572107946] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13951)
[1572108457] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;OK;HARD;3;OK- CPU IS USED 1.46%
[1572108596] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.15%
[1572108618] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;OK;HARD;3;OK- CPU IS USED 15.6%
[1572108621] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 27.3%
[1572108939] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.98%
[1572108940] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.66%
[1572109208] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572109336] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572109463] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572109589] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572109715] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572112195] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;OK;HARD;3;OK- CPU IS USED 15.55%
[1572113321] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572113448] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572113575] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572116129] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.51%
[1572116133] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.66%