[1563663922] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.61%
[1563664077] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.28%
[1563664228] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.47%
[1563665551] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563667858] SERVICE FLAPPING ALERT: ieatnetsimv7035-15;check netsim cpu;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1563667858] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;HARD;3;OK- CPU IS USED 75.43%
[1563669951] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563670865] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563670871] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1563671045] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1563671046] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563671226] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1563671227] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563671408] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1563671791] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1563672984] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13313) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563673132] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563673272] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13313) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563676730] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RBS;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1563676893] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14035) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563681661] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563682088] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1563682269] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1563691368] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2343) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563691499] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2343) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563691636] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2343) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563691787] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1563691924] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1563692466] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (116) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563693332] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1563699671] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (120) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563703276] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (121) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563703400] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563703528] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563707133] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (120) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563707489] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.38%
[1563707639] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.28%
[1563707790] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.1%
[1563709315] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12763) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563709456] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12763) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563709598] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13534) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563710737] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (52) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563710861] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (52) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563710986] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (52) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563711020] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1002) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563711421] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;HARD;3;OK- CPU IS USED 74.81%
[1563713208] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563713339] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563713470] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563717107] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563717589] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2679) IS LOWER THAN 80% OF ENM CPP NODES (4254)
[1563717723] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4193) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1563718264] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563718464] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1563718592] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1563718652] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.17%
[1563718687] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (462) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1563718720] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1563718803] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.64%
[1563718955] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.37%
[1563719106] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.43%
[1563720499] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2033) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563722939] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563723655] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563724230] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1563724230] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2033) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1563725129] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (1544) IS LOWER THAN 80% OF ENM CPP NODES (4254)
[1563725302] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4193) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1563725470] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2158) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563725512] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (1544) IS LOWER THAN 80% OF ENM CPP NODES (4254)
[1563725596] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2158) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563725646] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4193) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1563725737] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563726133] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (462) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1563726265] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2033) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563726402] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2033) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563727161] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563727259] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563727339] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563727384] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563727536] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563727703] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (1514) IS LOWER THAN 80% OF ENM CPP NODES (4254)
[1563727908] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4193) IS HIGHER THAN 95% OF ENM CPP NODES (2740)
[1563728259] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1563728991] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563729179] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1563729381] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1563729381] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2399) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563730867] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1563732987] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563733113] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563733246] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563733412] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563733526] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2557) IS LOWER THAN 80% OF ENM CPP NODES (4254)
[1563733703] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4193) IS HIGHER THAN 95% OF ENM CPP NODES (4104)
[1563734472] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563734640] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563734807] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563736196] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (110) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563738377] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563739341] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2679) IS LOWER THAN 80% OF ENM CPP NODES (4254)
[1563739349] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (112) IS LOWER THAN 80% OF ENM COMECIM NODES (14069)
[1563739480] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4193) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1563739537] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (14114) IS HIGHER THAN 95% OF ENM COMECIM NODES (14149)
[1563739654] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.44%
[1563739801] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563739804] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 8.8%
[1563741566] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563741731] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;SOFT;2;OK- NUMBER OF PM FILES (1438) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563741981] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563742147] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563742274] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563742440] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563742565] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563742689] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563742865] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1563742865] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563746877] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.63%
[1563747006] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (120) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563747028] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 40.02%
[1563747956] SERVICE ALERT: ieatnetsimv5116-97;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.57%
[1563747987] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1563748050] SERVICE ALERT: ieatnetsimv7035-24;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.92%
[1563748111] SERVICE ALERT: ieatnetsimv5116-97;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 25.45%
[1563748204] SERVICE ALERT: ieatnetsimv7035-24;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1563748357] SERVICE ALERT: ieatnetsimv7035-24;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1563748733] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2679) IS LOWER THAN 80% OF ENM CPP NODES (4104)
[1563748906] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4193) IS HIGHER THAN 95% OF ENM CPP NODES (4104)
[1563749877] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES