[1576540920] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576541050] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576541062] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (50)
[1576541112] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576541173] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576541236] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576541305] HOST ALERT: ieatnetsimv5116-22;DOWN;SOFT;1;PING NOK
[1576541367] HOST ALERT: ieatnetsimv5116-22;UP;SOFT;1;PING OK
[1576541381] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576541555] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1555) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576541696] HOST ALERT: ieatnetsimv7035-17;DOWN;SOFT;1;PING NOK
[1576541709] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1555) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576541758] HOST ALERT: ieatnetsimv7035-17;UP;SOFT;1;PING OK
[1576541845] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (276) IS LOWER THAN 5% OF ENM NODES (14035)
[1576542061] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (31) IS HIGHER THAN 20% OF ENM NODES (100)
[1576542165] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1576542184] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576542226] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (714) IS FROM 80% TO 95% OF ENM  NODES (767)
[1576542310] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (31) IS HIGHER THAN 20% OF ENM NODES (100)
[1576542352] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (715) IS FROM 80% TO 95% OF ENM  NODES (767)
[1576542359] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1576542435] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576542478] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (719) IS FROM 80% TO 95% OF ENM  NODES (767)
[1576542560] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (34) IS HIGHER THAN 20% OF ENM NODES (100)
[1576542686] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (34) IS HIGHER THAN 20% OF ENM NODES (100)
[1576542815] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576542948] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576543074] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576543185] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MTAS;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1576543185] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1576543202] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576543438] HOST ALERT: ieatnetsimv9034-80;DOWN;SOFT;1;PING NOK
[1576543438] HOST ALERT: ieatnetsimv9034-15;DOWN;SOFT;1;PING NOK
[1576543500] HOST ALERT: ieatnetsimv9034-80;UP;SOFT;1;PING OK
[1576543500] HOST ALERT: ieatnetsimv9034-15;UP;SOFT;1;PING OK
[1576544094] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1576544229] HOST ALERT: ieatnetsimv017-28;DOWN;SOFT;1;PING NOK
[1576544291] HOST ALERT: ieatnetsimv017-28;UP;SOFT;1;PING OK
[1576544367] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (91) IS FROM 80% TO 95% OF ENM  NODES (98)
[1576544479] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1576544497] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (94) IS HIGHER THAN 95% OF ENM  NODES (98)
[1576544633] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576544660] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1576544781] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1576544781] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576544785] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576544842] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1576544914] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576545038] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576545200] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12977) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576545286] SERVICE ALERT: ieatwlvm9034;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1576545333] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13676) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576545897] SERVICE ALERT: ieatwlvm9034;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1576546806] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1576546806] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1576546814] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1576546843] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized EPG;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576546843] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (49) IS HIGHER THAN 95% OF ENM  NODES (50)
[1576546942] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1576547070] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1576547342] HOST ALERT: ieatnetsimv9034-119;DOWN;SOFT;1;PING NOK
[1576547404] HOST ALERT: ieatnetsimv9034-119;UP;SOFT;1;PING OK
[1576547453] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1576547548] SERVICE FLAPPING ALERT: ieatlms7435;check system service hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576547980] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (40) IS BETWEEN 5% AND 20% OF ENM NODES (767)
[1576548107] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (40) IS BETWEEN 5% AND 20% OF ENM NODES (767)
[1576548231] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (40) IS BETWEEN 5% AND 20% OF ENM NODES (767)
[1576548239] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576548363] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576548465] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (543) IS LOWER THAN 80% OF ENM  NODES (750)
[1576548486] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576548611] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576548735] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576548831] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1446) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576548858] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576548956] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1472) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576548975] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576549081] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1748) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1576549095] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576549216] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576550269] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (650) IS FROM 80% TO 95% OF ENM  NODES (750)
[1576550651] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (37) IS LOWER THAN 80% OF ENM  NODES (50)
[1576550779] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576550868] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13134) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576550905] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1576551006] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (14015) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576551835] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (34) IS LOWER THAN 5% OF ENM NODES (767)
[1576552816] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576555143] HOST ALERT: ieatnetsimv9034-120;DOWN;SOFT;1;PING NOK
[1576555205] HOST ALERT: ieatnetsimv9034-120;UP;SOFT;1;PING OK
[1576555440] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (40) IS BETWEEN 5% AND 20% OF ENM NODES (767)
[1576555564] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (42) IS BETWEEN 5% AND 20% OF ENM NODES (767)
[1576555673] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (530) IS LOWER THAN 80% OF ENM  NODES (750)
[1576555689] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (41) IS BETWEEN 5% AND 20% OF ENM NODES (767)
[1576556146] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1576556322] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1576556335] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1576556511] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1576556525] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1576556701] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1576557238] HOST ALERT: ieatnetsimv9034-18;DOWN;SOFT;1;PING NOK
[1576557238] HOST ALERT: ieatnetsimv9034-16;DOWN;SOFT;1;PING NOK
[1576557300] HOST ALERT: ieatnetsimv9034-16;UP;SOFT;1;PING OK
[1576557300] HOST ALERT: ieatnetsimv9034-18;UP;SOFT;1;PING OK
[1576557477] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576557477] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (650) IS FROM 80% TO 95% OF ENM  NODES (750)
[1576558144] HOST ALERT: ieatnetsimv9034-92;DOWN;SOFT;1;PING NOK
[1576558206] HOST ALERT: ieatnetsimv9034-92;UP;SOFT;1;PING OK
[1576558744] HOST ALERT: ieatnetsimv9034-118;DOWN;SOFT;1;PING NOK
[1576558806] HOST ALERT: ieatnetsimv9034-118;UP;SOFT;1;PING OK
[1576559678] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1576559889] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1636) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576559941] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1559) IS LOWER THAN 80% OF ENM  NODES (2562)
[1576560019] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1667) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1576560071] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1559) IS LOWER THAN 80% OF ENM  NODES (2562)
[1576560196] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2559) IS HIGHER THAN 95% OF ENM  NODES (2562)
[1576560844] HOST ALERT: ieatnetsimv9034-91;DOWN;SOFT;1;PING NOK
[1576560907] HOST ALERT: ieatnetsimv9034-91;UP;SOFT;1;PING OK
[1576561912] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1576562090] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1576562242] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (80)
[1576562271] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1576562335] HOST ALERT: ieatnetsimv9034-127;DOWN;SOFT;1;PING NOK
[1576562335] HOST ALERT: ieatnetsimv9034-100;DOWN;SOFT;1;PING NOK
[1576562335] HOST ALERT: ieatnetsimv9034-73;DOWN;SOFT;1;PING NOK
[1576562335] HOST ALERT: ieatnetsimv9034-66;DOWN;SOFT;1;PING NOK
[1576562377] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1576562397] HOST ALERT: ieatnetsimv9034-100;UP;SOFT;1;PING OK
[1576562397] HOST ALERT: ieatnetsimv9034-73;UP;SOFT;1;PING OK
[1576562397] HOST ALERT: ieatnetsimv9034-127;UP;SOFT;1;PING OK
[1576562397] HOST ALERT: ieatnetsimv9034-66;UP;SOFT;1;PING OK
[1576562634] HOST ALERT: ieatnetsimv9034-105;DOWN;SOFT;1;PING NOK
[1576562638] HOST ALERT: ieatnetsimv9034-52;DOWN;SOFT;1;PING NOK
[1576562696] HOST ALERT: ieatnetsimv9034-105;UP;SOFT;1;PING OK
[1576562700] HOST ALERT: ieatnetsimv9034-52;UP;SOFT;1;PING OK
[1576563592] HOST ALERT: ieatnetsimv7035-23;DOWN;SOFT;1;PING NOK
[1576563596] HOST ALERT: ieatnetsimv7035-92;DOWN;SOFT;1;PING NOK
[1576563626] HOST ALERT: ieatnetsimv7035-92;UP;SOFT;1;PING OK
[1576563654] HOST ALERT: ieatnetsimv7035-23;UP;SOFT;1;PING OK
[1576563794] SERVICE ALERT: ieatlms7435;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1576563972] SERVICE ALERT: ieatlms7435;check node fs;OK;HARD;3;OK
[1576564682] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (563) IS LOWER THAN 80% OF ENM  NODES (750)
[1576564703] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1576564922] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (135) IS FROM 80% TO 95% OF ENM  NODES (150)
[1576565047] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (135) IS FROM 80% TO 95% OF ENM  NODES (150)
[1576565178] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (136) IS FROM 80% TO 95% OF ENM  NODES (150)
[1576565549] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1576565754] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1576565959] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1576566486] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (650) IS FROM 80% TO 95% OF ENM  NODES (750)
[1576566984] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1576567579] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2562)
[1576567707] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2562)
[1576567844] SERVICE ALERT: ieatlms7435;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1576568236] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576568424] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576568611] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576569472] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576570094] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1576570094] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (35) IS LOWER THAN 5% OF ENM NODES (767)
[1576570816] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576570937] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576571057] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576571620] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576571891] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (562) IS LOWER THAN 80% OF ENM  NODES (750)
[1576572435] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576572555] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576572675] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576572895] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1576573024] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1576573696] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (650) IS FROM 80% TO 95% OF ENM  NODES (750)
[1576574078] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576574554] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1598) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576574658] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576574679] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1623) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576574761] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1135) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1576574808] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1745) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1576574888] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1154) IS HIGHER THAN 95% OF ENM  NODES (1200)
[1576576276] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;OK;HARD;3;OK
[1576576781] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1576576904] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1576577031] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1576577187] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576577365] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1576577389] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MTAS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576577547] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1576577729] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1576579210] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576579595] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1576580094] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576580213] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1413) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576580215] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576580272] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576580335] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576580356] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1430) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576580400] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1576580400] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1576580482] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1740) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1576580902] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (48) IS BETWEEN 5% AND 20% OF ENM NODES (767)
[1576580902] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (496) IS LOWER THAN 80% OF ENM  NODES (750)
[1576581027] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (47) IS BETWEEN 5% AND 20% OF ENM NODES (767)
[1576581048] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized EPG;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576581150] SERVICE ALERT: ieatlms3901;check fm hb failures MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (46) IS BETWEEN 5% AND 20% OF ENM NODES (767)
[1576581510] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576581644] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1576582706] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (650) IS FROM 80% TO 95% OF ENM  NODES (750)
[1576586412] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576586536] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1576587775] SERVICE ALERT: ieatnetsimv7035-119;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.2%
[1576587926] SERVICE ALERT: ieatnetsimv7035-119;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 40.31%
[1576588130] SERVICE ALERT: ieatwlvm7102;check workload profiles;OK;HARD;3;OK- ALL WORKLOAD PROFILES ARE RUNNING
[1576589558] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.09%
[1576589709] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.84%
[1576589859] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.35%
[1576591330] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1576591534] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1576591590] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576591738] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1576592658] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576592668] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1576592668] SERVICE ALERT: ieatnetsimv7048-14;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592668] SERVICE ALERT: ieatnetsimv9034-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592668] HOST ALERT: ieatnetsimv9034-09;DOWN;SOFT;1;PING NOK
[1576592668] HOST ALERT: ieatnetsimv7048-14;DOWN;SOFT;1;PING NOK
[1576592668] SERVICE ALERT: ieatlms7435;check unsync nodes;OK;HARD;3;OK
[1576592669] HOST ALERT: ieatnetsimv7048-18;DOWN;SOFT;1;PING NOK
[1576592669] HOST ALERT: ieatnetsimv7048-16;DOWN;SOFT;1;PING NOK
[1576592669] HOST ALERT: ieatnetsimv7048-17;DOWN;SOFT;1;PING NOK
[1576592670] HOST ALERT: ieatnetsimv7048-19;DOWN;SOFT;1;PING NOK
[1576592670] HOST ALERT: ieatnetsimv7048-22;DOWN;SOFT;1;PING NOK
[1576592673] SERVICE ALERT: ieatnetsimv9034-118;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592673] HOST ALERT: ieatnetsimv9034-118;DOWN;SOFT;1;PING NOK
[1576592680] HOST ALERT: ieatnetsimv9034-01;DOWN;SOFT;1;PING NOK
[1576592680] SERVICE ALERT: ieatnetsimv9034-91;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592680] HOST ALERT: ieatnetsimv9034-91;DOWN;SOFT;1;PING NOK
[1576592681] HOST ALERT: ieatnetsimv9034-03;DOWN;SOFT;1;PING NOK
[1576592681] HOST ALERT: ieatnetsimv9034-47;DOWN;SOFT;1;PING NOK
[1576592681] HOST ALERT: ieatnetsimv9034-04;DOWN;SOFT;1;PING NOK
[1576592716] HOST ALERT: ieatnetsimv9034-06;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592717] HOST ALERT: ieatnetsimv9034-08;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592718] HOST ALERT: ieatnetsimv9034-103;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592718] HOST ALERT: ieatnetsimv9034-10;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592724] HOST ALERT: ieatnetsimv9034-105;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592725] HOST ALERT: ieatnetsimv7048-12;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1576592725] HOST ALERT: ieatnetsimv9034-66;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592726] HOST ALERT: ieatnetsimv9034-73;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1576592726] HOST ALERT: ieatnetsimv9034-127;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1576592726] HOST ALERT: ieatnetsimv9034-100;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1576592727] HOST ALERT: ieatwlvm9034;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592728] HOST ALERT: ieatnetsimv9034-15;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1576592728] HOST ALERT: ieatnetsimv9034-52;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592728] HOST ALERT: ieatnetsimv9034-80;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592728] HOST ALERT: ieatnetsimv9034-16;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1576592729] HOST ALERT: ieatnetsimv9034-18;DOWN;SOFT;1;(Host check timed out after 30.88 seconds)
[1576592733] HOST ALERT: ieatnetsimv9034-119;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1576592733] HOST ALERT: ieatnetsimv9034-120;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592735] HOST ALERT: ieatnetsimv9034-93;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592735] HOST ALERT: ieatnetsimv9034-92;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1576592736] HOST ALERT: ieatnetsimv9034-55;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1576592747] HOST ALERT: ieatnetsimv9034-106;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1576592747] HOST ALERT: ieatnetsimv9034-108;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1576592758] HOST ALERT: ieatnetsimv9034-09;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592759] HOST ALERT: ieatnetsimv7048-18;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592759] HOST ALERT: ieatnetsimv7048-17;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592759] HOST ALERT: ieatnetsimv7048-16;DOWN;SOFT;2;(Host check timed out after 30.46 seconds)
[1576592759] HOST ALERT: ieatnetsimv7048-14;DOWN;SOFT;2;(Host check timed out after 30.54 seconds)
[1576592760] HOST ALERT: ieatnetsimv7048-19;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592760] HOST ALERT: ieatnetsimv7048-22;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592763] HOST ALERT: ieatnetsimv9034-118;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592770] HOST ALERT: ieatnetsimv9034-01;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592770] HOST ALERT: ieatnetsimv9034-91;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592771] HOST ALERT: ieatnetsimv9034-03;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592771] HOST ALERT: ieatnetsimv9034-04;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592772] HOST ALERT: ieatnetsimv9034-47;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1576592778] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576592806] HOST ALERT: ieatnetsimv9034-06;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592807] HOST ALERT: ieatnetsimv9034-08;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592808] HOST ALERT: ieatnetsimv9034-103;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592808] HOST ALERT: ieatnetsimv9034-10;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592815] HOST ALERT: ieatnetsimv7048-12;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592815] HOST ALERT: ieatnetsimv9034-66;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592815] HOST ALERT: ieatnetsimv9034-105;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1576592817] HOST ALERT: ieatwlvm9034;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592817] HOST ALERT: ieatnetsimv9034-73;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1576592817] HOST ALERT: ieatnetsimv9034-100;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1576592817] HOST ALERT: ieatnetsimv9034-127;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1576592817] SERVICE ALERT: ieatnetsimv9034-90;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592819] HOST ALERT: ieatnetsimv9034-15;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1576592819] HOST ALERT: ieatnetsimv9034-52;DOWN;SOFT;2;(Host check timed out after 31.00 seconds)
[1576592819] HOST ALERT: ieatnetsimv9034-80;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1576592819] HOST ALERT: ieatnetsimv9034-16;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1576592820] HOST ALERT: ieatnetsimv9034-18;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1576592822] SERVICE ALERT: ieatnetsimv9034-109;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592823] HOST ALERT: ieatnetsimv9034-119;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592823] SERVICE ALERT: ieatnetsimv9034-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592823] HOST ALERT: ieatnetsimv9034-31;DOWN;SOFT;1;PING NOK
[1576592823] HOST ALERT: ieatnetsimv9034-118;DOWN;SOFT;3;PING NOK
[1576592823] HOST ALERT: ieatnetsimv7048-17;DOWN;SOFT;3;PING NOK
[1576592823] HOST ALERT: ieatnetsimv7048-19;DOWN;SOFT;3;PING NOK
[1576592823] HOST ALERT: ieatnetsimv7048-22;DOWN;SOFT;3;PING NOK
[1576592823] SERVICE ALERT: ieatnetsimv7048-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592823] HOST ALERT: ieatnetsimv7048-10;DOWN;SOFT;1;PING NOK
[1576592823] HOST ALERT: ieatnetsimv9034-120;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592824] SERVICE ALERT: ieatwlvm9034;check workload profiles;OK;HARD;3;OK- ALL WORKLOAD PROFILES ARE RUNNING
[1576592824] HOST ALERT: ieatwlvm9034;DOWN;SOFT;3;PING NOK
[1576592824] SERVICE ALERT: ieatnetsimv9034-118;check netsim nodes;CRITICAL;HARD;1;CRITICAL- THERE ARE STOPPED NODES
[1576592824] HOST ALERT: ieatnetsimv9034-90;DOWN;SOFT;1;PING NOK
[1576592825] SERVICE ALERT: ieatnetsimv9034-114;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592825] HOST ALERT: ieatnetsimv9034-114;DOWN;SOFT;1;PING NOK
[1576592825] HOST ALERT: ieatnetsimv9034-109;DOWN;SOFT;1;PING NOK
[1576592825] SERVICE ALERT: ieatnetsimv7048-12;check netsim nodes;CRITICAL;HARD;1;CRITICAL- THERE ARE STOPPED NODES
[1576592825] HOST ALERT: ieatnetsimv9034-93;DOWN;SOFT;2;(Host check timed out after 30.08 seconds)
[1576592826] HOST ALERT: ieatnetsimv9034-09;DOWN;SOFT;3;PING NOK
[1576592826] HOST ALERT: ieatnetsimv7048-14;DOWN;SOFT;3;PING NOK
[1576592826] SERVICE ALERT: ieatnetsimv9034-46;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592826] HOST ALERT: ieatnetsimv9034-46;DOWN;SOFT;1;PING NOK
[1576592826] HOST ALERT: ieatnetsimv7048-18;DOWN;SOFT;3;PING NOK
[1576592826] HOST ALERT: ieatnetsimv7048-16;DOWN;SOFT;3;PING NOK
[1576592826] HOST ALERT: ieatnetsimv9034-55;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592826] SERVICE ALERT: ieatnetsimv9034-112;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592826] HOST ALERT: ieatnetsimv9034-112;DOWN;SOFT;1;PING NOK
[1576592826] HOST ALERT: ieatnetsimv9034-92;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1576592826] HOST ALERT: ieatnetsimv9034-55;DOWN;SOFT;3;PING NOK
[1576592827] HOST ALERT: ieatnetsimv9034-16;DOWN;SOFT;3;PING NOK
[1576592828] SERVICE ALERT: ieatnetsimv9034-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592828] HOST ALERT: ieatnetsimv9034-11;DOWN;SOFT;1;PING NOK
[1576592828] SERVICE ALERT: ieatnetsimv9034-121;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592828] HOST ALERT: ieatnetsimv9034-121;DOWN;SOFT;1;PING NOK
[1576592828] SERVICE ALERT: ieatnetsimv9034-36;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592828] HOST ALERT: ieatnetsimv9034-36;DOWN;SOFT;1;PING NOK
[1576592830] HOST ALERT: ieatnetsimv9034-01;DOWN;SOFT;3;PING NOK
[1576592830] SERVICE ALERT: ieatnetsimv9034-116;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592830] HOST ALERT: ieatnetsimv9034-116;DOWN;SOFT;1;PING NOK
[1576592830] HOST ALERT: ieatnetsimv9034-91;DOWN;SOFT;3;PING NOK
[1576592831] HOST ALERT: ieatnetsimv9034-119;DOWN;SOFT;3;PING NOK
[1576592831] HOST ALERT: ieatnetsimv9034-04;DOWN;SOFT;3;PING NOK
[1576592831] HOST ALERT: ieatnetsimv9034-03;DOWN;SOFT;3;PING NOK
[1576592831] SERVICE ALERT: ieatnetsimv9034-91;check netsim nodes;CRITICAL;HARD;1;CRITICAL- THERE ARE STOPPED NODES
[1576592832] HOST ALERT: ieatnetsimv9034-47;DOWN;SOFT;3;PING NOK
[1576592833] SERVICE ALERT: ieatnetsimv9034-95;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592833] HOST ALERT: ieatnetsimv9034-95;DOWN;SOFT;1;PING NOK
[1576592833] SERVICE ALERT: ieatnetsimv9034-44;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592833] HOST ALERT: ieatnetsimv9034-44;DOWN;SOFT;1;PING NOK
[1576592837] HOST ALERT: ieatnetsimv9034-106;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592837] HOST ALERT: ieatnetsimv9034-108;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1576592838] HOST ALERT: ieatnetsimv9034-09;DOWN;SOFT;4;PING NOK
[1576592840] SERVICE ALERT: ieatnetsimv9034-79;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592840] HOST ALERT: ieatnetsimv9034-79;DOWN;SOFT;1;PING NOK
[1576592841] HOST ALERT: ieatnetsimv9034-15;DOWN;SOFT;3;PING NOK
[1576592844] SERVICE ALERT: ieatnetsimv9034-96;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592844] HOST ALERT: ieatnetsimv9034-96;DOWN;SOFT;1;PING NOK
[1576592844] SERVICE ALERT: ieatnetsimv9034-66;check netsim nodes;CRITICAL;HARD;1;CRITICAL- THERE ARE STOPPED NODES
[1576592845] SERVICE ALERT: ieatnetsimv9034-82;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592845] HOST ALERT: ieatnetsimv9034-82;DOWN;SOFT;1;PING NOK
[1576592845] SERVICE ALERT: ieatnetsimv7048-22;check netsim nodes;CRITICAL;HARD;1;CRITICAL- THERE ARE STOPPED NODES
[1576592847] SERVICE ALERT: ieatnetsimv7048-03;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592847] HOST ALERT: ieatnetsimv7048-03;DOWN;SOFT;1;PING NOK
[1576592851] SERVICE ALERT: ieatnetsimv9034-39;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592851] HOST ALERT: ieatnetsimv9034-39;DOWN;SOFT;1;PING NOK
[1576592852] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1576592852] SERVICE ALERT: ieatnetsimv9034-09;check netsim nodes;CRITICAL;HARD;1;CRITICAL- THERE ARE STOPPED NODES
[1576592852] SERVICE ALERT: ieatnetsimv7048-14;check netsim nodes;CRITICAL;HARD;1;CRITICAL- THERE ARE STOPPED NODES
[1576592852] HOST ALERT: ieatnetsimv7048-05;DOWN;SOFT;1;PING NOK
[1576592852] HOST ALERT: ieatnetsimv7048-06;DOWN;SOFT;1;PING NOK
[1576592853] HOST ALERT: ieatnetsimv7048-09;DOWN;SOFT;1;PING NOK
[1576592853] HOST ALERT: ieatnetsimv7048-07;DOWN;SOFT;1;PING NOK
[1576592853] HOST ALERT: ieatnetsimv7048-04;DOWN;SOFT;1;PING NOK
[1576592854] HOST ALERT: ieatnetsimv7048-11;DOWN;SOFT;1;PING NOK
[1576592854] HOST ALERT: ieatnetsimv7048-02;DOWN;SOFT;1;PING NOK
[1576592854] SERVICE ALERT: ieatnetsimv7048-04;check netsim nodes;CRITICAL;HARD;1;CRITICAL- THERE ARE STOPPED NODES
[1576592854] HOST ALERT: ieatnetsimv9034-31;DOWN;SOFT;2;PING NOK
[1576592855] SERVICE ALERT: ieatnetsimv7048-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592855] HOST ALERT: ieatnetsimv7048-08;DOWN;SOFT;1;PING NOK
[1576592856] HOST ALERT: ieatnetsimv7048-21;DOWN;SOFT;1;PING NOK
[1576592857] HOST ALERT: ieatnetsimv7048-20;DOWN;SOFT;1;PING NOK
[1576592858] SERVICE ALERT: ieatnetsimv9034-19;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592858] HOST ALERT: ieatnetsimv9034-19;DOWN;SOFT;1;PING NOK
[1576592860] SERVICE ALERT: ieatnetsimv9034-48;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592868] SERVICE ALERT: ieatlms7435;check fm supervision nodes SBG-IS;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1576592869] HOST ALERT: ieatnetsimv7048-19;UP;SOFT;1;PING OK
[1576592869] HOST ALERT: ieatnetsimv9034-06;UP;SOFT;1;PING OK
[1576592870] HOST ALERT: ieatnetsimv9034-08;UP;SOFT;1;PING OK
[1576592870] HOST ALERT: ieatnetsimv9034-10;UP;SOFT;1;PING OK
[1576592870] HOST ALERT: ieatnetsimv9034-103;UP;SOFT;1;PING OK
[1576592871] SERVICE ALERT: ieatnetsimv9034-111;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592875] HOST ALERT: ieatnetsimv9034-03;UP;SOFT;1;PING OK
[1576592877] HOST ALERT: ieatnetsimv7048-12;UP;SOFT;1;PING OK
[1576592877] HOST ALERT: ieatnetsimv9034-66;UP;SOFT;1;PING OK
[1576592877] SERVICE ALERT: ieatnetsimv9034-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576592877] HOST ALERT: ieatnetsimv9034-105;UP;SOFT;1;PING OK
[1576592879] HOST ALERT: ieatnetsimv9034-127;UP;SOFT;1;PING OK
[1576592879] HOST ALERT: ieatnetsimv9034-100;UP;SOFT;1;PING OK
[1576592879] HOST ALERT: ieatnetsimv9034-73;UP;SOFT;1;PING OK
[1576592881] HOST ALERT: ieatnetsimv9034-52;UP;SOFT;1;PING OK
[1576592881] HOST ALERT: ieatnetsimv9034-80;UP;SOFT;1;PING OK
[1576592882] HOST ALERT: ieatnetsimv9034-18;UP;SOFT;1;PING OK
[1576592886] HOST ALERT: ieatnetsimv9034-118;UP;SOFT;1;PING OK
[1576592886] HOST ALERT: ieatnetsimv7048-17;UP;SOFT;1;PING OK
[1576592886] HOST ALERT: ieatnetsimv7048-22;UP;SOFT;1;PING OK
[1576592886] HOST ALERT: ieatnetsimv7048-10;UP;SOFT;1;PING OK
[1576592886] HOST ALERT: ieatnetsimv9034-120;UP;SOFT;1;PING OK
[1576592886] HOST ALERT: ieatwlvm9034;UP;SOFT;1;PING OK
[1576592887] HOST ALERT: ieatnetsimv9034-90;UP;SOFT;1;PING OK
[1576592887] HOST ALERT: ieatnetsimv9034-114;UP;SOFT;1;PING OK
[1576592887] SERVICE ALERT: ieatnetsimv9034-108;check netsim nodes;CRITICAL;HARD;1;CRITICAL- THERE ARE STOPPED NODES
[1576592887] HOST ALERT: ieatnetsimv9034-109;UP;SOFT;1;PING OK
[1576592888] HOST ALERT: ieatnetsimv9034-93;UP;SOFT;1;PING OK
[1576592888] HOST ALERT: ieatnetsimv7048-14;UP;SOFT;1;PING OK
[1576592888] HOST ALERT: ieatnetsimv9034-46;UP;SOFT;1;PING OK
[1576592889] HOST ALERT: ieatnetsimv7048-16;UP;SOFT;1;PING OK
[1576592889] HOST ALERT: ieatnetsimv7048-18;UP;SOFT;1;PING OK
[1576592889] HOST ALERT: ieatnetsimv9034-112;UP;SOFT;1;PING OK
[1576592889] HOST ALERT: ieatnetsimv9034-55;UP;SOFT;1;PING OK
[1576592889] HOST ALERT: ieatnetsimv9034-92;UP;SOFT;1;PING OK
[1576592889] HOST ALERT: ieatnetsimv9034-16;UP;SOFT;1;PING OK
[1576592890] HOST ALERT: ieatnetsimv9034-11;UP;SOFT;1;PING OK
[1576592890] HOST ALERT: ieatnetsimv9034-121;UP;SOFT;1;PING OK
[1576592891] HOST ALERT: ieatnetsimv9034-36;UP;SOFT;1;PING OK
[1576592892] HOST ALERT: ieatnetsimv9034-01;UP;SOFT;1;PING OK
[1576592892] HOST ALERT: ieatnetsimv9034-116;UP;SOFT;1;PING OK
[1576592892] HOST ALERT: ieatnetsimv9034-91;UP;SOFT;1;PING OK
[1576592893] HOST ALERT: ieatnetsimv9034-119;UP;SOFT;1;PING OK
[1576592893] HOST ALERT: ieatnetsimv9034-04;UP;SOFT;1;PING OK
[1576592894] HOST ALERT: ieatnetsimv9034-47;UP;SOFT;1;PING OK
[1576592895] HOST ALERT: ieatnetsimv9034-95;UP;SOFT;1;PING OK
[1576592895] HOST ALERT: ieatnetsimv9034-44;UP;SOFT;1;PING OK
[1576592899] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576592899] HOST ALERT: ieatnetsimv9034-106;UP;SOFT;1;PING OK
[1576592899] HOST ALERT: ieatnetsimv9034-108;UP;SOFT;1;PING OK
[1576592900] HOST ALERT: ieatnetsimv9034-09;UP;SOFT;1;PING OK
[1576592903] HOST ALERT: ieatnetsimv9034-79;UP;SOFT;1;PING OK
[1576592904] HOST ALERT: ieatnetsimv9034-15;UP;SOFT;1;PING OK
[1576592906] HOST ALERT: ieatnetsimv9034-96;UP;SOFT;1;PING OK
[1576592907] HOST ALERT: ieatnetsimv9034-82;UP;SOFT;1;PING OK
[1576592909] HOST ALERT: ieatnetsimv7048-03;UP;SOFT;1;PING OK
[1576592913] HOST ALERT: ieatnetsimv9034-39;UP;SOFT;1;PING OK
[1576592915] HOST ALERT: ieatnetsimv7048-05;UP;SOFT;1;PING OK
[1576592915] HOST ALERT: ieatnetsimv7048-06;UP;SOFT;1;PING OK
[1576592916] HOST ALERT: ieatnetsimv7048-09;UP;SOFT;1;PING OK
[1576592916] HOST ALERT: ieatnetsimv7048-04;UP;SOFT;1;PING OK
[1576592916] HOST ALERT: ieatnetsimv7048-07;UP;SOFT;1;PING OK
[1576592916] HOST ALERT: ieatnetsimv7048-02;UP;SOFT;1;PING OK
[1576592916] HOST ALERT: ieatnetsimv7048-11;UP;SOFT;1;PING OK
[1576592917] HOST ALERT: ieatnetsimv9034-31;UP;SOFT;1;PING OK
[1576592917] HOST ALERT: ieatnetsimv7048-08;UP;SOFT;1;PING OK
[1576592918] HOST ALERT: ieatnetsimv7048-21;UP;SOFT;1;PING OK
[1576592919] HOST ALERT: ieatnetsimv7048-20;UP;SOFT;1;PING OK
[1576592920] HOST ALERT: ieatnetsimv9034-19;UP;SOFT;1;PING OK
[1576592949] SERVICE ALERT: ieatnetsimv9034-90;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592954] SERVICE ALERT: ieatnetsimv9034-109;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592955] SERVICE ALERT: ieatnetsimv7048-10;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592957] SERVICE ALERT: ieatnetsimv9034-36;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592957] SERVICE ALERT: ieatnetsimv9034-46;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592959] SERVICE ALERT: ieatnetsimv9034-112;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592960] SERVICE ALERT: ieatnetsimv9034-121;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592961] SERVICE ALERT: ieatnetsimv9034-114;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592962] SERVICE ALERT: ieatnetsimv9034-116;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592965] SERVICE ALERT: ieatnetsimv9034-44;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592965] SERVICE ALERT: ieatnetsimv9034-95;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592972] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1576592972] SERVICE ALERT: ieatnetsimv9034-79;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592974] SERVICE ALERT: ieatnetsimv9034-31;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592976] SERVICE ALERT: ieatnetsimv9034-96;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592978] SERVICE ALERT: ieatnetsimv9034-82;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592978] SERVICE ALERT: ieatnetsimv7048-03;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592979] SERVICE ALERT: ieatnetsimv9034-39;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592987] SERVICE ALERT: ieatnetsimv7048-08;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592987] SERVICE ALERT: ieatnetsimv9034-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576592992] SERVICE ALERT: ieatnetsimv9034-48;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576593007] SERVICE ALERT: ieatnetsimv9034-111;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576593013] SERVICE ALERT: ieatnetsimv9034-71;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576593014] SERVICE ALERT: ieatnetsimv9034-19;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576593308] SERVICE ALERT: ieatwlvm9034;check workload profiles;CRITICAL;SOFT;1;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1576593439] SERVICE ALERT: ieatwlvm9034;check workload profiles;CRITICAL;SOFT;2;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1576593571] SERVICE ALERT: ieatwlvm9034;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1576593623] SERVICE FLAPPING ALERT: ieatlms7435;check system service hc;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1576593623] SERVICE ALERT: ieatlms7435;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1576593770] SERVICE ALERT: ieatlms7435;check fm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1576593892] SERVICE ALERT: ieatlms7435;check fm supervision nodes SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1576593945] SERVICE ALERT: ieatlms7435;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1576593947] SERVICE ALERT: ieatlms7435;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1576594015] SERVICE ALERT: ieatlms7435;check fm supervision nodes SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1576594066] SERVICE ALERT: ieatlms7435;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1576594067] SERVICE ALERT: ieatlms7435;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1576594186] SERVICE ALERT: ieatlms7435;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1576594188] SERVICE ALERT: ieatlms7435;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1576594276] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576594396] SERVICE ALERT: ieatlms7435;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1576594397] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576594499] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1576594517] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576594539] SERVICE ALERT: ieatnetsimv7048-04;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1576594571] SERVICE ALERT: ieatnetsimv9034-108;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1576594583] SERVICE ALERT: ieatlms7435;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1576594594] SERVICE ALERT: ieatnetsimv7048-22;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1576594599] SERVICE ALERT: ieatnetsimv7048-14;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1576594605] SERVICE ALERT: ieatnetsimv9034-118;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1576594605] SERVICE ALERT: ieatnetsimv7048-12;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1576594617] SERVICE ALERT: ieatnetsimv9034-91;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1576594622] SERVICE ALERT: ieatnetsimv9034-09;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1576594643] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1576594657] SERVICE ALERT: ieatnetsimv9034-66;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1576594704] SERVICE ALERT: ieatlms7435;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1576594832] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1576595179] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576595253] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576595263] SERVICE ALERT: ieatlms7435;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1576595300] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576595311] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (578) IS LOWER THAN 80% OF ENM  NODES (750)
[1576595384] SERVICE ALERT: ieatlms7435;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1576595420] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576595424] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1576595504] SERVICE ALERT: ieatlms7435;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1576595545] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1576595665] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1576595831] SERVICE ALERT: ieatlms7435;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1576595849] SERVICE ALERT: ieatlms7435;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1576595878] SERVICE ALERT: ieatlms7435;check san alert healthcheck;CRITICAL;SOFT;1;CRITICAL- san_alert_healthcheck FAILED
[1576595951] SERVICE ALERT: ieatlms7435;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1576595970] SERVICE ALERT: ieatlms7435;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1576595999] SERVICE ALERT: ieatlms7435;check san alert healthcheck;CRITICAL;SOFT;2;CRITICAL- san_alert_healthcheck FAILED
[1576596072] SERVICE ALERT: ieatlms7435;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1576596090] SERVICE ALERT: ieatlms7435;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1576596119] SERVICE ALERT: ieatlms7435;check san alert healthcheck;CRITICAL;HARD;3;CRITICAL- san_alert_healthcheck FAILED
[1576596499] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576596792] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1576597295] SERVICE ALERT: ieatlms7435;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576597375] SERVICE ALERT: ieatlms7435;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1576597550] SERVICE ALERT: ieatlms7435;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1576597859] SERVICE ALERT: ieatlms7435;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1576597962] SERVICE ALERT: ieatlms7435;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1576598117] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;OK;HARD;3;OK
[1576598333] SERVICE ALERT: ieatwlvm7102;check workload profiles;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1576598441] SERVICE ALERT: ieatlms7435;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1576598455] SERVICE ALERT: ieatwlvm7102;check workload profiles;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1576598578] SERVICE ALERT: ieatwlvm7102;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1576598612] SERVICE ALERT: ieatlms7435;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1576598793] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576598852] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576598926] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (50)
[1576599010] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576599052] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576599136] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576599178] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (50)
[1576599302] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1576599387] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576599434] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576599436] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (719) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576599454] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1122) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576599565] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576599567] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1576599594] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (8) IS LOWER THAN 5% OF ENM NODES (14035)
[1576599739] SERVICE ALERT: ieatlms7435;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1576599770] SERVICE ALERT: ieatlms7435;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1576599790] SERVICE ALERT: ieatlms7435;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1576600424] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (10) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576600563] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (22) IS HIGHER THAN 20% OF ENM NODES (100)
[1576600648] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.69%
[1576600687] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (15) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576600798] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.38%
[1576600979] SERVICE ALERT: ieatwlvm7102;check workload profiles;OK;HARD;3;OK- ALL WORKLOAD PROFILES ARE RUNNING
[1576602739] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN DSC NODES !
[1576603175] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1576604291] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576612517] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576612638] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576612758] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576613094] HOST ALERT: ieatnetsimv9041-10;DOWN;SOFT;1;PING NOK
[1576613156] HOST ALERT: ieatnetsimv9041-10;UP;SOFT;1;PING OK
[1576613857] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576613996] HOST ALERT: ieatnetsimv7035-93;DOWN;SOFT;1;PING NOK
[1576614058] HOST ALERT: ieatnetsimv7035-93;UP;SOFT;1;PING OK
[1576614500] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576614620] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576614740] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576617008] SERVICE ALERT: ieatnetsimv5116-05;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576617129] SERVICE ALERT: ieatnetsimv5116-05;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576617250] SERVICE ALERT: ieatnetsimv5116-05;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576618715] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1576621940] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576622993] HOST ALERT: ieatnetsimv7035-86;DOWN;SOFT;1;PING NOK
[1576623055] HOST ALERT: ieatnetsimv7035-86;UP;SOFT;1;PING OK
[1576623558] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;OK;HARD;3;OK
[1576624928] HOST ALERT: ieatnetsimv017-26;DOWN;SOFT;1;PING NOK
[1576624991] HOST ALERT: ieatnetsimv017-26;UP;SOFT;1;PING OK