[1566774091] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1566774091] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566774342] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.67%
[1566774500] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79%
[1566775786] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2210) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566776029] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.02%
[1566776180] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 39.01%
[1566778894] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566779044] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566779196] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566782827] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.28%
[1566783711] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.63%
[1566783862] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.51%
[1566784012] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.39%
[1566784055] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.55%
[1566784163] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.49%
[1566784206] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.62%
[1566784357] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.44%
[1566784397] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.56%
[1566784547] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 12.9%
[1566786289] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3083) IS LOWER THAN 80% OF ENM CPP NODES (4255)
[1566790760] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566791163] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566791576] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566791657] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566791778] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566791898] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566794331] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.48%
[1566794482] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 22.2%
[1566799299] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1566801375] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.57%
[1566801525] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.95%
[1566801676] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 92.6%
[1566803006] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.16%
[1566803156] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 27.68%
[1566803307] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.26%
[1566803455] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566803458] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 35.39%
[1566803605] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 71.03%
[1566804757] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566804829] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1566804853] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566804942] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566805010] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1566805034] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566805128] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566805160] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1566805190] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1566805215] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566805288] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1566805416] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1566805549] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.41%
[1566805683] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.89%
[1566805700] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 25.03%
[1566805833] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 21.18%
[1566805850] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.18%
[1566806001] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 59.21%
[1566808517] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566808681] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566808851] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1566808863] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566808876] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1566809177] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.08%
[1566809327] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.28%
[1566810593] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566810703] SERVICE ALERT: ieatnetsimv7035-13;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566810824] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1566812507] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 80.91%
[1566813314] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3461) IS FROM 80% TO 95% OF ENM CPP NODES (4230)
[1566813370] SERVICE ALERT: ieatnetsimv5116-34;check netsim disk;OK;HARD;3;OK
[1566813413] SERVICE ALERT: ieatnetsimv5116-51;check netsim disk;OK;HARD;3;OK
[1566813498] SERVICE ALERT: ieatnetsimv5116-10;check netsim disk;OK;HARD;3;OK
[1566813578] SERVICE ALERT: ieatnetsimv5116-20;check netsim disk;OK;HARD;3;OK
[1566813626] SERVICE ALERT: ieatnetsimv5116-39;check netsim disk;OK;HARD;3;OK
[1566813633] SERVICE ALERT: ieatnetsimv5116-42;check netsim disk;OK;HARD;3;OK
[1566813670] SERVICE ALERT: ieatnetsimv5116-57;check netsim disk;OK;HARD;3;OK
[1566813682] SERVICE ALERT: ieatnetsimv5116-61;check netsim disk;OK;HARD;3;OK
[1566813687] SERVICE ALERT: ieatnetsimv5116-63;check netsim disk;OK;HARD;3;OK
[1566813804] SERVICE ALERT: ieatnetsimv5116-14;check netsim disk;OK;HARD;3;OK
[1566813879] SERVICE ALERT: ieatnetsimv5116-44;check netsim disk;OK;HARD;3;OK
[1566814090] SERVICE ALERT: ieatnetsimv5116-31;check netsim disk;OK;HARD;3;OK
[1566814259] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566814444] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566818229] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566818415] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566818601] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566818876] SERVICE FLAPPING ALERT: ieatnetsimv5116-68;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566822955] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6497) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1566823097] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6003) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1566823251] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6436) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1566823883] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566824067] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566824076] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566824246] SERVICE ALERT: ieatnetsimv5116-64;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566824270] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566824287] SERVICE ALERT: ieatnetsimv5116-80;check netsim disk;OK;HARD;3;OK
[1566824293] SERVICE ALERT: ieatnetsimv5116-82;check netsim disk;OK;HARD;3;OK
[1566824367] SERVICE ALERT: ieatnetsimv5116-64;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566824488] SERVICE ALERT: ieatnetsimv5116-64;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566824488] SERVICE ALERT: ieatnetsimv5116-62;check netsim disk;OK;HARD;3;OK
[1566824519] SERVICE ALERT: ieatnetsimv5116-76;check netsim disk;OK;HARD;3;OK
[1566824548] SERVICE ALERT: ieatnetsimv7035-107;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566824615] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1566824681] SERVICE ALERT: ieatnetsimv7035-107;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1566824796] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1566824977] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1566827033] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.38%
[1566827183] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 57.33%
[1566827587] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.62%
[1566827738] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.68%
[1566827893] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566828070] SERVICE ALERT: ieatnetsimv5116-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.56%
[1566828077] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1566828080] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;OK;HARD;3;OK
[1566828221] SERVICE ALERT: ieatnetsimv5116-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 67.06%
[1566828260] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566829532] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13942) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13945)
[1566830018] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2193) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2194)
[1566830123] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566830306] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1566830489] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566832353] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566832537] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566834522] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566834707] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1566834891] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566835088] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.46%
[1566835241] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.53%
[1566835392] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.76%
[1566836690] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.49%
[1566836848] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.73%
[1566837006] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566837738] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.19%
[1566837888] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.07%
[1566838554] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566838738] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1566838921] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566839023] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.51%
[1566840642] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;HARD;3;OK- CPU IS USED 75.97%
[1566840784] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566840967] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566844138] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566846704] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1566852438] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.83%
[1566852589] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.62%
[1566853160] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.83%
[1566853311] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.57%
[1566853461] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.94%
[1566854733] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1566854858] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566855010] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566855163] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566857054] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.08%
[1566857204] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.24%
[1566857756] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.54%
[1566857915] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566858067] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566858682] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.77%
[1566859523] HOST ALERT: ieatnetsimv5116-12;DOWN;SOFT;1;PING NOK
[1566859586] HOST ALERT: ieatnetsimv5116-12;UP;SOFT;1;PING OK