[1564961343] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1564961641] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1564962004] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1564963728] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.52%
[1564963879] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 38.59%
[1564963916] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1564964044] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1564964172] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1564964752] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564964938] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564965124] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564971431] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.18%
[1564971582] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.28%
[1564971592] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1564971733] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 61.54%
[1564971954] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1564972311] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1564973741] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.92%
[1564973891] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 42.52%
[1564974042] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.32%
[1564974193] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 47.89%
[1564975363] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.46%
[1564975514] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 15.73%
[1564986465] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.98%
[1564986616] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 62.6%
[1564986788] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1564986953] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1564987381] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1564987407] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1564987580] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1564988589] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1564988638] SERVICE ALERT: ieatnetsimv5116-110;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1564988775] SERVICE ALERT: ieatnetsimv5116-110;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1564989535] SERVICE ALERT: ieatnetsimv5116-62;check netsim disk;OK;HARD;3;OK
[1564989866] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1564990055] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1564990093] SERVICE ALERT: ieatnetsimv5116-50;check netsim disk;OK;HARD;3;OK
[1564990097] SERVICE ALERT: ieatnetsimv5116-85;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1564990110] SERVICE ALERT: ieatnetsimv5116-58;check netsim disk;OK;HARD;3;OK
[1564990112] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;OK;HARD;3;OK
[1564990681] SERVICE ALERT: ieatnetsimv5116-43;check netsim disk;OK;HARD;3;OK
[1564990685] SERVICE ALERT: ieatnetsimv5116-45;check netsim disk;OK;HARD;3;OK
[1564990931] SERVICE ALERT: ieatnetsimv5116-48;check netsim disk;OK;HARD;3;OK
[1564991016] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1564992255] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564992440] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1564992442] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (20545) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1564992989] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1564993117] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1564993245] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1564993423] SERVICE ALERT: ieatnetsimv5116-36;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1564993544] SERVICE ALERT: ieatnetsimv5116-36;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1564993561] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1564993664] SERVICE ALERT: ieatnetsimv5116-36;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1564993712] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 59.84%
[1564993864] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564994016] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564994167] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564994426] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1564994611] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1564994797] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1564995391] SERVICE ALERT: ieatnetsimv7035-20;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.92%
[1564995541] SERVICE ALERT: ieatnetsimv7035-20;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 25.81%
[1564995742] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.11%
[1564995786] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564995808] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.16%
[1564995923] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1564995943] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1564996155] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (25041) IS HIGHER THAN 95% OF ENM NODES (25570)
[1564996178] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (6891) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1564996637] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13379) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1564996659] SERVICE ALERT: ieatnetsimv7035-33;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1564996782] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13847) IS HIGHER THAN 95% OF ENM COMECIM NODES (14149)
[1564996789] SERVICE ALERT: ieatnetsimv7035-33;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1564996854] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1564996854] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1564996857] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1564996911] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1564997165] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23812) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1564997173] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (21164) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1564997265] SERVICE ALERT: ieatnetsimv5116-36;check netsim disk;OK;HARD;3;OK
[1564997365] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.84 seconds)
[1564997683] SERVICE ALERT: ieatnetsimv5116-85;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.96%
[1564997797] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;OK;HARD;3;OK- CPU IS USED 18.71%
[1564998216] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1564998254] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1564998406] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1564998557] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1564998729] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12849) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1564998866] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1564998872] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13123) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1564998875] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564999014] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13855) IS HIGHER THAN 95% OF ENM COMECIM NODES (14149)
[1564999136] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1564999657] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (25041) IS HIGHER THAN 95% OF ENM NODES (25570)
[1565000056] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1565000262] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1565000375] SERVICE FLAPPING ALERT: ieatlms3901;check pm supervision nodes;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1565000375] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23812) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1565000976] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565001267] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1565004878] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1565005006] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1565005131] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1565006073] SERVICE FLAPPING ALERT: ieatlms3901;check cm supervision nodes;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1565006073] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565006336] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.4%
[1565006487] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 36.45%
[1565007601] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (21786) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1565007777] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565008540] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (25041) IS HIGHER THAN 95% OF ENM NODES (25570)
[1565009328] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1565009460] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565009514] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1565010175] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565010202] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1565010238] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.83%
[1565010380] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1565010389] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 50.98%
[1565010540] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.66%
[1565010690] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 67.72%
[1565010953] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (25408) IS HIGHER THAN 95% OF ENM NODES (25570)
[1565011569] SERVICE ALERT: ieatwlvm5116;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1565011580] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (22061) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1565011603] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23812) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1565012100] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565012229] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565012245] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565012374] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565012381] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1565012533] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1565012777] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1565013281] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1565014071] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565014199] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565014205] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565014226] SERVICE ALERT: ieatnetsimv5116-19;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565014238] SERVICE ALERT: ieatnetsimv5116-17;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565014254] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6592) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1565014327] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565014334] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565014358] SERVICE ALERT: ieatnetsimv5116-19;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565014367] SERVICE ALERT: ieatnetsimv5116-17;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565014452] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6588) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1565014466] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565014492] SERVICE ALERT: ieatnetsimv5116-19;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565014494] SERVICE ALERT: ieatnetsimv5116-17;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565014642] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (22392) IS BETWEEN 80% AND 95% OF ENM NODES (25570)
[1565014705] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6272) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1565014836] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565014969] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565015003] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565015098] SERVICE ALERT: ieatnetsimv5116-18;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565015131] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565015259] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565015767] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565015895] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565015904] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565016001] SERVICE ALERT: ieatnetsimv5116-12;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565016022] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565016032] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565016106] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (6053) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1565016128] SERVICE ALERT: ieatnetsimv5116-12;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565016159] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565016163] SERVICE FLAPPING ALERT: ieatnetsimv5116-99;check netsim cpu;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1565016163] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;OK;HARD;3;OK- CPU IS USED 16.53%
[1565016256] SERVICE ALERT: ieatnetsimv5116-12;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565016268] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1565016287] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565016417] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565016449] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1565016547] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565016631] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1565017835] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565017974] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565018054] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1565018492] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1565018692] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (25408) IS HIGHER THAN 95% OF ENM NODES (25570)
[1565019227] SERVICE ALERT: ieatnetsimv5116-16;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1565019377] SERVICE ALERT: ieatnetsimv5116-16;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 8.03%
[1565019741] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565019870] SERVICE ALERT: ieatnetsimv5116-12;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565019881] SERVICE ALERT: ieatnetsimv5116-20;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565020165] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565020335] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (6802) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1565020506] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1565024032] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6771) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1565024190] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (6825) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1565024891] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1565024891] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6814) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1565028882] SERVICE FLAPPING ALERT: ieatlms3901;check pm supervision nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565029262] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1565029291] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565030129] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6814) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1565030971] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1565031049] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1565031168] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1565031365] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1565031969] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1565032024] SERVICE FLAPPING ALERT: ieatlms3901;check cm supervision nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565032757] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (6785) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1565036680] SERVICE ALERT: ieatnetsimv5116-64;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565036714] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (1121) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1565036800] SERVICE ALERT: ieatnetsimv5116-64;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565036921] SERVICE ALERT: ieatnetsimv5116-64;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565039641] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.88%
[1565039792] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 62.73%
[1565040897] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (5872) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)