[1561678952] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1561679421] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1561680120] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;OK;HARD;3;OK
[1561681014] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.76%
[1561690604] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.13%
[1561690755] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.48%
[1561690906] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.94%
[1561691748] HOST ALERT: ieatnetsimv5116-50;DOWN;SOFT;1;PING NOK
[1561691748] HOST ALERT: ieatnetsimv5116-27;DOWN;SOFT;1;PING NOK
[1561691748] HOST ALERT: ieatnetsimv5116-17;DOWN;SOFT;1;PING NOK
[1561691810] HOST ALERT: ieatnetsimv5116-17;UP;SOFT;1;PING OK
[1561691810] HOST ALERT: ieatnetsimv5116-27;UP;SOFT;1;PING OK
[1561691810] HOST ALERT: ieatnetsimv5116-50;UP;SOFT;1;PING OK
[1561696544] HOST ALERT: ieatnetsimv5116-56;DOWN;SOFT;1;PING NOK
[1561696544] HOST ALERT: ieatnetsimv5116-36;DOWN;SOFT;1;PING NOK
[1561696606] HOST ALERT: ieatnetsimv5116-56;UP;SOFT;1;PING OK
[1561696606] HOST ALERT: ieatnetsimv5116-36;UP;SOFT;1;PING OK
[1561702155] SERVICE ALERT: ieatnetsimv7035-79;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1561702276] SERVICE ALERT: ieatnetsimv7035-79;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1561702396] SERVICE ALERT: ieatnetsimv7035-79;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1561702930] SERVICE ALERT: ieatnetsimv7035-12;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1561704095] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1561704281] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1561704773] SERVICE ALERT: ieatnetsimv7035-12;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1561704922] SERVICE ALERT: ieatnetsimv7035-12;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1561705005] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.54%
[1561705070] SERVICE ALERT: ieatnetsimv7035-12;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1561705158] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.94%
[1561705290] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.49%
[1561705309] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.22%
[1561705441] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 73.28%
[1561705993] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4128) IS FROM 80% TO 95% OF ENM SNMP NODES (4893)
[1561705997] SERVICE ALERT: ieatnetsimv7035-79;check netsim disk;OK;HARD;3;OK
[1561706132] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (2368) IS LOWER THAN 80% OF ENM SNMP NODES (4893)
[1561706272] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (2230) IS LOWER THAN 80% OF ENM SNMP NODES (4893)
[1561706714] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (20468) IS BETWEEN 80% AND 95% OF ENM NODES (24353)
[1561707062] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH CM SUPERVISION (20468) IS BETWEEN 80% AND 95% OF ENM NODES (24353)
[1561707283] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (20468) IS BETWEEN 80% AND 95% OF ENM NODES (24353)
[1561708252] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.98%
[1561708403] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.19%
[1561708588] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708588] SERVICE ALERT: ieatnetsimv5116-25;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708588] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708588] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708588] SERVICE ALERT: ieatnetsimv5116-58;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708589] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708589] SERVICE ALERT: ieatnetsimv5116-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.95%
[1561708591] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708591] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708591] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708739] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561708739] SERVICE ALERT: ieatnetsimv5116-25;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561708739] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561708739] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561708739] SERVICE ALERT: ieatnetsimv5116-58;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561708740] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561708740] SERVICE ALERT: ieatnetsimv5116-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.97%
[1561708742] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561708744] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561708745] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561708890] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561708890] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561708890] SERVICE ALERT: ieatnetsimv5116-25;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561708890] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561708890] SERVICE ALERT: ieatnetsimv5116-58;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561708891] SERVICE ALERT: ieatnetsimv5116-26;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708891] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561708891] SERVICE ALERT: ieatnetsimv5116-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.58%
[1561708893] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561708897] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561708898] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561708899] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561708942] SERVICE ALERT: ieatnetsimv7035-19;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.41%
[1561708976] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (14316) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20429)
[1561709042] SERVICE ALERT: ieatnetsimv5116-26;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.82%
[1561709057] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561709192] SERVICE ALERT: ieatnetsimv5116-26;check netsim cpu;OK;HARD;3;OK- CPU IS USED 76.21%
[1561709194] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (14316) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20429)
[1561709210] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561709409] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (14428) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20429)
[1561710071] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1) IS LOWER THAN 80% OF ENM NODES (24353)
[1561710262] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1561710359] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (6664) IS LOWER THAN 80% OF ENM NODES (24353)
[1561710449] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1561710512] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (1) IS LOWER THAN 80% OF ENM COMECIM NODES (11144)
[1561710637] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1561710665] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;2;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (1) IS LOWER THAN 80% OF ENM COMECIM NODES (11144)
[1561710672] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM CPP NODES (7110)
[1561710813] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM CPP NODES (7110)
[1561710817] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (1) IS LOWER THAN 80% OF ENM COMECIM NODES (11144)
[1561710953] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM CPP NODES (7110)
[1561712520] SERVICE ALERT: ieatnetsimv5116-25;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.96%
[1561712521] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;OK;HARD;3;OK- CPU IS USED 41.37%
[1561712521] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.21%
[1561712521] SERVICE ALERT: ieatnetsimv5116-58;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.89%
[1561712522] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.19%
[1561712523] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.46%
[1561712528] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;OK;HARD;3;OK- CPU IS USED 44.23%
[1561713057] SERVICE ALERT: ieatlms3901;check pm collection statistical;OK;HARD;3;OK- NUMBER OF PM FILES (10881) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1561714088] SERVICE ALERT: ieatnetsimv7035-12;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1561716120] SERVICE ALERT: ieatnetsimv5116-23;check netsim cpu;OK;HARD;3;OK- CPU IS USED 22.26%
[1561716130] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 53.93%
[1561717066] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1561717251] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1561717436] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1561717889] SERVICE FLAPPING ALERT: ieatnetsimv7035-30;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1561719736] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.88%
[1561719887] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.4%
[1561720038] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.07%
[1561720042] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 60.86%
[1561720947] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.9%
[1561721013] SERVICE ALERT: ieatnetsimv5116-85;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561721062] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (1308) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1561721098] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 26.73%
[1561721101] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1561721127] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (24328) IS HIGHER THAN 95% OF ENM NODES (24353)
[1561721165] SERVICE ALERT: ieatnetsimv5116-85;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1561721315] SERVICE ALERT: ieatnetsimv5116-85;check netsim cpu;OK;HARD;3;OK- CPU IS USED 66.79%
[1561721722] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (5679) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1561721774] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (7100) IS HIGHER THAN 95% OF ENM CPP NODES (7110)
[1561722492] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4635) IS FROM 80% TO 95% OF ENM SNMP NODES (4893)
[1561722955] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.55%
[1561722988] SERVICE ALERT: ieatnetsimv5116-01;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561723106] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.18%
[1561723138] SERVICE ALERT: ieatnetsimv5116-01;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.68%
[1561723209] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1561723258] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.03%
[1561723330] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1561723409] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.39%
[1561723451] SERVICE ALERT: ieatnetsimv5116-18;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1561723454] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (11139) IS HIGHER THAN 95% OF ENM COMECIM NODES (11144)
[1561723563] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.96%
[1561723700] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.29%
[1561723851] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 9.01%
[1561724313] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4878) IS HIGHER THAN 95% OF ENM SNMP NODES (4893)
[1561724767] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1561724912] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (21818) IS BETWEEN 80% AND 95% OF ENM NODES (24353)
[1561724952] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1561725137] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1561725223] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH CM SUPERVISION (21818) IS BETWEEN 80% AND 95% OF ENM NODES (24353)
[1561725537] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (21818) IS BETWEEN 80% AND 95% OF ENM NODES (24353)
[1561725759] SERVICE ALERT: ieatlms3901;check pm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (24328) IS HIGHER THAN 95% OF ENM NODES (24353)
[1561726131] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (2368) IS LOWER THAN 80% OF ENM SNMP NODES (4893)
[1561726271] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (2368) IS LOWER THAN 80% OF ENM SNMP NODES (4893)
[1561726414] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1561726414] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (2368) IS LOWER THAN 80% OF ENM SNMP NODES (4893)
[1561727194] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;OK;HARD;3;OK- CPU IS USED 74.27%
[1561727524] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (11) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21945)
[1561727720] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (11) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21945)
[1561727908] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (11) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21945)
[1561728660] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1561728789] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1561728803] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1561728803] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1561728919] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1561730189] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1561730340] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.84%
[1561730553] SERVICE ALERT: ieatnetsimv5116-65;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1561730561] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.48%
[1561730700] SERVICE ALERT: ieatnetsimv5116-65;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1561730712] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 16.75%
[1561730938] SERVICE ALERT: ieatnetsimv5116-65;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1561734329] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1561734405] SERVICE ALERT: ieatnetsimv5116-86;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.39%
[1561734556] SERVICE ALERT: ieatnetsimv5116-86;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 4.09%
[1561735177] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1561735177] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (20239) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21945)
[1561736022] SERVICE ALERT: ieatnetsimv7035-26;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.57%
[1561736175] SERVICE ALERT: ieatnetsimv7035-26;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.68%
[1561736327] SERVICE ALERT: ieatnetsimv7035-26;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1561736429] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1561736429] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1561736429] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1561736478] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1561736606] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1561736735] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1561738847] SERVICE ALERT: ieatlms3901;check pm collection statistical;OK;HARD;3;OK- NUMBER OF PM FILES (21048) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (21945)
[1561739958] SERVICE ALERT: ieatnetsimv7035-26;check netsim cpu;OK;HARD;3;OK- CPU IS USED 71.41%
[1561741539] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1561741668] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1561741798] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1561745270] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.61%
[1561752010] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1561752131] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1561752251] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1561753005] HOST ALERT: ieatnetsimv5116-06;DOWN;SOFT;1;PING NOK
[1561753067] HOST ALERT: ieatnetsimv5116-06;UP;SOFT;1;PING OK
[1561756005] HOST ALERT: ieatnetsimv5116-68;DOWN;SOFT;1;PING NOK
[1561756067] HOST ALERT: ieatnetsimv5116-68;UP;SOFT;1;PING OK
[1561756101] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.83%
[1561756252] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.13%
[1561756402] SERVICE ALERT: ieatnetsimv7035-22;check netsim cpu;OK;HARD;3;OK- CPU IS USED 81.63%
[1561759625] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.33%
[1561759776] SERVICE ALERT: ieatnetsimv7035-25;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 67.49%
[1561760633] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1561761268] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)