[1576454933] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12930) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576455066] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (14015) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576457453] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1576457798] SERVICE ALERT: ieatnetsimv9034-109;check netsim disk;OK;HARD;3;OK
[1576460986] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576461123] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN DSC NODES !
[1576463287] SERVICE ALERT: ieatlms3901;check fm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (150) IS LOWER THAN 80% OF ENM NODES (300)
[1576463399] HOST ALERT: ieatnetsimv9041-22;DOWN;SOFT;1;PING NOK
[1576463415] SERVICE ALERT: ieatlms3901;check fm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1576463461] HOST ALERT: ieatnetsimv9041-22;UP;SOFT;1;PING OK
[1576464633] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576464765] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576466661] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1622) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576466786] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1570) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576466914] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1558) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576470519] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1748) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1576475689] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576475814] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (12) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576475912] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576475938] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576476097] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576476283] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576477188] HOST ALERT: ieatnetsimv7035-13;DOWN;SOFT;1;PING NOK
[1576477199] HOST ALERT: ieatnetsimv7035-41;DOWN;SOFT;1;PING NOK
[1576477201] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1200)
[1576477251] HOST ALERT: ieatnetsimv7035-13;UP;SOFT;1;PING OK
[1576477261] HOST ALERT: ieatnetsimv7035-41;UP;SOFT;1;PING OK
[1576477326] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1576479053] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576479174] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576479294] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576479364] SERVICE ALERT: ieatlms7435;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1576481929] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576486755] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (14) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576486793] HOST ALERT: ieatnetsimv7035-88;DOWN;SOFT;1;PING NOK
[1576486855] HOST ALERT: ieatnetsimv7035-88;UP;SOFT;1;PING OK
[1576486883] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (23) IS HIGHER THAN 20% OF ENM NODES (100)
[1576487010] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (25) IS HIGHER THAN 20% OF ENM NODES (100)
[1576487030] SERVICE ALERT: ieatnetsimv9034-54;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92%
[1576487180] SERVICE ALERT: ieatnetsimv9034-54;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.88%
[1576489801] HOST ALERT: ieatnetsimv9041-26;DOWN;SOFT;1;PING NOK
[1576489863] HOST ALERT: ieatnetsimv9041-26;UP;SOFT;1;PING OK
[1576490613] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576493993] HOST ALERT: ieatnetsimv7035-119;DOWN;SOFT;1;PING NOK
[1576493993] HOST ALERT: ieatnetsimv9041-08;DOWN;SOFT;1;PING NOK
[1576494055] HOST ALERT: ieatnetsimv9041-08;UP;SOFT;1;PING OK
[1576494055] HOST ALERT: ieatnetsimv7035-119;UP;SOFT;1;PING OK
[1576494202] HOST ALERT: ieatnetsimv017-34;DOWN;SOFT;1;PING NOK
[1576494221] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (10) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576494264] HOST ALERT: ieatnetsimv017-34;UP;SOFT;1;PING OK
[1576494348] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576494355] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (32) IS HIGHER THAN 20% OF ENM NODES (100)
[1576494479] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (27) IS HIGHER THAN 20% OF ENM NODES (100)
[1576495599] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576495786] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576495841] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (80) IS BETWEEN 80% AND 95% OF ENM NODES (93)
[1576495931] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576495966] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (80) IS BETWEEN 80% AND 95% OF ENM NODES (93)
[1576495998] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576496056] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576496111] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH PM SUPERVISION (80) IS BETWEEN 80% AND 95% OF ENM NODES (93)
[1576496113] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (80) IS BETWEEN 80% AND 95% OF ENM NODES (93)
[1576496236] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH PM SUPERVISION (80) IS BETWEEN 80% AND 95% OF ENM NODES (93)
[1576496261] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (15) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576496339] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (299) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576496367] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (80) IS BETWEEN 80% AND 95% OF ENM NODES (93)
[1576496376] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576496384] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (19) IS HIGHER THAN 20% OF ENM NODES (80)
[1576496472] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (16) IS LOWER THAN 5% OF ENM NODES (1200)
[1576496501] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576496515] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (24) IS HIGHER THAN 20% OF ENM NODES (80)
[1576496624] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1576496968] SERVICE ALERT: ieatnetsimv9041-06;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576496991] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1576497012] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1576497118] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1576497138] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1576497245] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1576497277] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1576497379] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1576497402] SERVICE ALERT: ieatnetsimv9041-06;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576497825] SERVICE ALERT: ieatnetsimv9041-06;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576497920] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1576497966] SERVICE ALERT: ieatlms7435;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1576498020] SERVICE ALERT: ieatlms7435;check node fs;OK;HARD;3;OK
[1576498073] SERVICE ALERT: ieatnetsimv9041-06;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1576498084] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1576498084] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576498176] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (71) IS HIGHER THAN 95% OF ENM NODES (72)
[1576498216] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1576498319] SERVICE ALERT: ieatnetsimv9041-06;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1576498674] HOST ALERT: ieatnetsimv7102-52;UP;HARD;1;PING OK
[1576498675] HOST ALERT: ieatnetsimv7102-11;UP;HARD;1;PING OK
[1576498683] HOST ALERT: ieatnetsimv7102-62;UP;HARD;1;PING OK
[1576498684] HOST ALERT: ieatnetsimv7102-80;UP;HARD;1;PING OK
[1576498692] HOST ALERT: ieatnetsimv7102-66;UP;HARD;1;PING OK
[1576498693] HOST ALERT: ieatnetsimv7102-70;UP;HARD;1;PING OK
[1576498694] HOST ALERT: ieatnetsimv7102-01;UP;HARD;1;PING OK
[1576498702] HOST ALERT: ieatnetsimv7102-81;UP;HARD;1;PING OK
[1576498703] HOST ALERT: ieatnetsimv7102-07;UP;HARD;1;PING OK
[1576498708] HOST ALERT: ieatnetsimv7102-45;UP;HARD;1;PING OK
[1576498708] HOST ALERT: ieatnetsimv7102-12;UP;HARD;1;PING OK
[1576498715] HOST ALERT: ieatnetsimv7102-78;UP;HARD;1;PING OK
[1576498715] HOST ALERT: ieatnetsimv7102-72;UP;HARD;1;PING OK
[1576498725] HOST ALERT: ieatnetsimv7102-43;UP;HARD;1;PING OK
[1576498728] HOST ALERT: ieatnetsimv7102-33;UP;HARD;1;PING OK
[1576498735] HOST ALERT: ieatnetsimv7102-06;UP;HARD;1;PING OK
[1576498739] HOST ALERT: ieatnetsimv7102-63;UP;HARD;1;PING OK
[1576498754] HOST ALERT: ieatnetsimv7102-04;UP;HARD;1;PING OK
[1576498783] HOST ALERT: ieatnetsimv7102-13;UP;HARD;1;PING OK
[1576498784] HOST ALERT: ieatnetsimv7102-57;UP;HARD;1;PING OK
[1576498785] HOST ALERT: ieatwlvm7102;UP;HARD;1;PING OK
[1576498794] HOST ALERT: ieatnetsimv7102-24;UP;HARD;1;PING OK
[1576498798] HOST ALERT: ieatnetsimv7102-16;UP;HARD;1;PING OK
[1576498799] HOST ALERT: ieatnetsimv7102-76;UP;HARD;1;PING OK
[1576498801] HOST ALERT: ieatnetsimv7102-19;UP;HARD;1;PING OK
[1576498806] HOST ALERT: ieatnetsimv7102-30;UP;HARD;1;PING OK
[1576498812] HOST ALERT: ieatnetsimv7102-37;UP;HARD;1;PING OK
[1576498812] HOST ALERT: ieatnetsimv7102-39;UP;HARD;1;PING OK
[1576498818] HOST ALERT: ieatnetsimv7102-58;UP;HARD;1;PING OK
[1576498819] HOST ALERT: ieatnetsimv7102-75;UP;HARD;1;PING OK
[1576498820] HOST ALERT: ieatnetsimv7102-27;UP;HARD;1;PING OK
[1576498820] HOST ALERT: ieatnetsimv7102-49;UP;HARD;1;PING OK
[1576498823] HOST ALERT: ieatnetsimv7102-54;UP;HARD;1;PING OK
[1576498825] HOST ALERT: ieatnetsimv7102-14;UP;HARD;1;PING OK
[1576498825] HOST ALERT: ieatnetsimv7102-05;UP;HARD;1;PING OK
[1576498825] HOST ALERT: ieatnetsimv7102-34;UP;HARD;1;PING OK
[1576498826] HOST ALERT: ieatnetsimv7102-02;UP;HARD;1;PING OK
[1576498826] HOST ALERT: ieatnetsimv7102-20;UP;HARD;1;PING OK
[1576498827] HOST ALERT: ieatnetsimv7102-08;UP;HARD;1;PING OK
[1576498827] HOST ALERT: ieatnetsimv7102-26;UP;HARD;1;PING OK
[1576498827] HOST ALERT: ieatnetsimv7102-17;UP;HARD;1;PING OK
[1576498828] HOST ALERT: ieatnetsimv7102-23;UP;HARD;1;PING OK
[1576498828] HOST ALERT: ieatnetsimv7102-60;UP;HARD;1;PING OK
[1576498828] HOST ALERT: ieatnetsimv7102-32;UP;HARD;1;PING OK
[1576498828] HOST ALERT: ieatnetsimv7102-29;UP;HARD;1;PING OK
[1576498828] HOST ALERT: ieatnetsimv7102-35;UP;HARD;1;PING OK
[1576498830] HOST ALERT: ieatnetsimv7102-38;UP;HARD;1;PING OK
[1576498830] HOST ALERT: ieatnetsimv7102-77;UP;HARD;1;PING OK
[1576498833] HOST ALERT: ieatnetsimv7102-47;UP;HARD;1;PING OK
[1576498834] HOST ALERT: ieatnetsimv7102-68;UP;HARD;1;PING OK
[1576498834] HOST ALERT: ieatnetsimv7102-59;UP;HARD;1;PING OK
[1576498834] HOST ALERT: ieatnetsimv7102-53;UP;HARD;1;PING OK
[1576498836] HOST ALERT: ieatnetsimv7102-71;UP;HARD;1;PING OK
[1576498836] HOST ALERT: ieatnetsimv7102-67;UP;HARD;1;PING OK
[1576498836] HOST ALERT: ieatnetsimv7102-83;UP;HARD;1;PING OK
[1576498838] HOST ALERT: ieatnetsimv7102-56;UP;HARD;1;PING OK
[1576498840] HOST ALERT: ieatnetsimv7102-82;UP;HARD;1;PING OK
[1576498844] HOST ALERT: ieatnetsimv7102-84;UP;HARD;1;PING OK
[1576498851] HOST ALERT: ieatnetsimv7102-41;UP;HARD;1;PING OK
[1576498852] HOST ALERT: ieatnetsimv7102-61;UP;HARD;1;PING OK
[1576498852] HOST ALERT: ieatnetsimv7102-31;UP;HARD;1;PING OK
[1576498853] HOST ALERT: ieatnetsimv7102-44;UP;HARD;1;PING OK
[1576498854] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (43) IS BETWEEN 80% AND 95% OF ENM NODES (52)
[1576498854] HOST ALERT: ieatnetsimv7102-55;UP;HARD;1;PING OK
[1576498856] HOST ALERT: ieatnetsimv7102-50;UP;HARD;1;PING OK
[1576498857] HOST ALERT: ieatnetsimv7102-15;UP;HARD;1;PING OK
[1576498859] HOST ALERT: ieatnetsimv7102-18;UP;HARD;1;PING OK
[1576498863] HOST ALERT: ieatnetsimv7102-42;UP;HARD;1;PING OK
[1576498874] HOST ALERT: ieatnetsimv7102-03;UP;HARD;1;PING OK
[1576498875] HOST ALERT: ieatnetsimv7102-25;UP;HARD;1;PING OK
[1576498875] HOST ALERT: ieatnetsimv7102-64;UP;HARD;1;PING OK
[1576498879] HOST ALERT: ieatnetsimv7102-51;UP;HARD;1;PING OK
[1576498881] HOST ALERT: ieatnetsimv7102-73;UP;HARD;1;PING OK
[1576498937] HOST ALERT: ieatnetsimv7102-74;UP;HARD;1;PING OK
[1576498938] HOST ALERT: ieatnetsimv7102-65;UP;HARD;1;PING OK
[1576498951] HOST ALERT: ieatnetsimv7102-48;UP;HARD;1;PING OK
[1576498951] HOST ALERT: ieatnetsimv7102-10;UP;HARD;1;PING OK
[1576498954] HOST ALERT: ieatnetsimv7102-09;UP;HARD;1;PING OK
[1576498960] HOST ALERT: ieatnetsimv7102-28;UP;HARD;1;PING OK
[1576498961] HOST ALERT: ieatnetsimv7102-22;UP;HARD;1;PING OK
[1576498966] HOST ALERT: ieatnetsimv7102-79;UP;HARD;1;PING OK
[1576498966] HOST ALERT: ieatnetsimv7102-40;UP;HARD;1;PING OK
[1576498969] HOST ALERT: ieatnetsimv7102-36;UP;HARD;1;PING OK
[1576498969] HOST ALERT: ieatnetsimv7102-21;UP;HARD;1;PING OK
[1576498971] HOST ALERT: ieatnetsimv7102-46;UP;HARD;1;PING OK
[1576498971] HOST ALERT: ieatnetsimv7102-69;UP;HARD;1;PING OK
[1576498978] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH CM SUPERVISION (40) IS BETWEEN 80% AND 95% OF ENM NODES (45)
[1576499110] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (37) IS HIGHER THAN 95% OF ENM NODES (38)
[1576499124] SERVICE ALERT: ieatwlvm7102;check workload profiles;CRITICAL;SOFT;1;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1576499287] SERVICE ALERT: ieatwlvm7102;check workload profiles;CRITICAL;SOFT;2;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1576499450] SERVICE ALERT: ieatwlvm7102;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1576499665] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576500122] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576500203] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (410) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576500231] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (11) IS HIGHER THAN 20% OF ENM NODES (50)
[1576500327] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (30) IS LOWER THAN 5% OF ENM NODES (1200)
[1576500355] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (50)
[1576500464] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (75) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576500481] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (14) IS HIGHER THAN 20% OF ENM NODES (50)
[1576500636] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (14) IS HIGHER THAN 20% OF ENM NODES (50)
[1576500695] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (79) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576500761] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (17) IS HIGHER THAN 20% OF ENM NODES (50)
[1576500823] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (205) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576500877] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8) IS HIGHER THAN 20% OF ENM NODES (30)
[1576500893] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576500914] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (3) IS LOWER THAN 80% OF ENM NODES (7)
[1576501001] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1576501020] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576501038] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (3) IS LOWER THAN 80% OF ENM NODES (7)
[1576501129] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576501147] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576501162] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (3) IS LOWER THAN 80% OF ENM NODES (7)
[1576501255] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576501380] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576501614] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576501799] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576501984] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576502068] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1576502093] HOST ALERT: ieatnetsimv7035-118;DOWN;SOFT;1;PING NOK
[1576502155] HOST ALERT: ieatnetsimv7035-118;UP;SOFT;1;PING OK
[1576502579] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576502822] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576502924] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576503317] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (0)!
[1576503529] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576503535] SERVICE ALERT: ieatlms3901;check number nodes RBS - DEMO;CRITICAL;SOFT;1;(No output on stdout) stderr: /usr/local/nagios/libexec/s4/check_number_nodes.sh: line 69: syntax error near unexpected token `)'
[1576503541] SERVICE ALERT: ieatlms3901;check number nodes RBS - DEMO;CRITICAL;SOFT;2;(No output on stdout) stderr: /usr/local/nagios/libexec/s4/check_number_nodes.sh: line 69: syntax error near unexpected token `)'
[1576503585] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576503605] SERVICE ALERT: ieatlms3901;check number nodes RBS - DEMO;CRITICAL;HARD;3;ENM TOTAL 11547 - 1903
[1576503727] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576503773] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576503849] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576503849] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576503851] SERVICE ALERT: ieatlms3901;check number nodes BSC - DEMO;CRITICAL;SOFT;1;ENM TOTAL 288 - 96
[1576503855] SERVICE ALERT: ieatlms3901;check number nodes BSC - DEMO;CRITICAL;SOFT;2;ENM TOTAL 288 - 96
[1576503996] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576503996] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576504055] SERVICE ALERT: ieatlms3901;check number nodes BSC - DEMO;CRITICAL;HARD;3;ENM TOTAL 288 - 96
[1576504176] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3201) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576504310] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (14035)
[1576504421] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576504421] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576504432] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (420) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576504442] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (2973) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576504577] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (530) IS LOWER THAN 5% OF ENM NODES (14035)
[1576504617] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576504617] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576504844] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576504844] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576504923] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576504923] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576505106] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576505106] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576505225] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1576505225] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576505253] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576505294] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (25) IS HIGHER THAN 20% OF ENM NODES (100)
[1576505302] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576505373] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576505420] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (25) IS HIGHER THAN 20% OF ENM NODES (100)
[1576505426] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576505494] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576505510] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576505546] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (29.4% change >= 20.0% threshold)
[1576505546] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (16) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576506075] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (29.4% change >= 20.0% threshold)
[1576506075] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576506233] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (29.4% change >= 20.0% threshold)
[1576506233] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576506310] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (29.4% change >= 20.0% threshold)
[1576506310] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576506496] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1576506519] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (29.4% change >= 20.0% threshold)
[1576506519] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576506643] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (29.4% change >= 20.0% threshold)
[1576506643] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576506841] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (29.4% change >= 20.0% threshold)
[1576506841] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576506909] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (29.4% change >= 20.0% threshold)
[1576506909] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576506995] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1576507147] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576507199] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1576507267] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576507388] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576508005] HOST ALERT: ieatnetsimv5116-118;DOWN;SOFT;1;PING NOK
[1576508005] HOST ALERT: ieatnetsimv5116-116;DOWN;SOFT;1;PING NOK
[1576508040] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1576508067] HOST ALERT: ieatnetsimv5116-118;UP;SOFT;1;PING OK
[1576508067] HOST ALERT: ieatnetsimv5116-116;UP;SOFT;1;PING OK
[1576508307] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3305) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576508352] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (7) IS HIGHER THAN 20% OF ENM NODES (30)
[1576508401] HOST ALERT: ieatnetsimv7035-78;DOWN;SOFT;1;PING NOK
[1576508436] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (2652) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576508463] HOST ALERT: ieatnetsimv7035-78;UP;SOFT;1;PING OK
[1576508583] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1717) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576509049] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (750) IS BETWEEN 80% AND 95% OF ENM NODES (923)
[1576509084] SERVICE ALERT: ieatlms7435;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1576509565] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (750) IS BETWEEN 80% AND 95% OF ENM NODES (889)
[1576509574] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1576509903] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1576510170] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1576510971] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1576510988] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;OK;HARD;3;OK
[1576511176] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1576511380] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1576511399] HOST ALERT: ieatnetsimv7035-73;DOWN;SOFT;1;PING NOK
[1576511461] HOST ALERT: ieatnetsimv7035-73;UP;SOFT;1;PING OK
[1576511955] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576512271] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (776)
[1576512656] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (767)
[1576512694] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576512749] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (100)
[1576513839] SERVICE ALERT: ieatnetsimv7102-38;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.01%
[1576513989] SERVICE ALERT: ieatnetsimv7102-38;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 1.14%
[1576514533] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (10) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576514657] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576514780] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576515065] SERVICE FLAPPING ALERT: ieatlms7435;check system service hc;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1576515065] SERVICE ALERT: ieatlms7435;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1576515169] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1576515283] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (181) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576515407] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (57) IS LOWER THAN 5% OF ENM NODES (1200)
[1576515783] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576515794] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (14035)
[1576518393] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (13) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576518521] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576519133] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (232) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576519259] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (21) IS LOWER THAN 5% OF ENM NODES (1200)
[1576519498] HOST ALERT: ieatnetsimv9041-19;DOWN;SOFT;1;PING NOK
[1576519560] HOST ALERT: ieatnetsimv9041-19;UP;SOFT;1;PING OK
[1576519952] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (17) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576520080] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (10) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576520204] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (4) IS LOWER THAN 5% OF ENM NODES (100)
[1576520527] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576520647] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576520695] HOST ALERT: ieatnetsimv7035-95;DOWN;SOFT;1;PING NOK
[1576520757] HOST ALERT: ieatnetsimv7035-95;UP;SOFT;1;PING OK
[1576521598] HOST ALERT: ieatnetsimv7035-49;DOWN;SOFT;1;PING NOK
[1576521660] HOST ALERT: ieatnetsimv7035-49;UP;SOFT;1;PING OK
[1576522247] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576522372] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576522376] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576522500] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576522514] HOST ALERT: ieatnetsimv9041-35;DOWN;SOFT;1;PING NOK
[1576522576] HOST ALERT: ieatnetsimv9041-35;UP;SOFT;1;PING OK
[1576522624] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576522983] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (312) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576522989] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (7) IS HIGHER THAN 20% OF ENM NODES (30)
[1576523004] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2789) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576523108] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (54) IS LOWER THAN 5% OF ENM NODES (1200)
[1576523117] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (7) IS HIGHER THAN 20% OF ENM NODES (30)
[1576523142] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (2913) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576523242] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1576523273] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (382) IS LOWER THAN 5% OF ENM NODES (14035)
[1576523808] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576523931] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (35) IS HIGHER THAN 20% OF ENM NODES (100)
[1576524056] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576526229] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1576526229] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1576526358] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1576526358] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576526697] HOST ALERT: ieatnetsimv7035-96;DOWN;SOFT;1;PING NOK
[1576526697] HOST ALERT: ieatnetsimv9041-15;DOWN;SOFT;1;PING NOK
[1576526697] HOST ALERT: ieatnetsimv9041-17;DOWN;SOFT;1;PING NOK
[1576526759] HOST ALERT: ieatnetsimv9041-15;UP;SOFT;1;PING OK
[1576526759] HOST ALERT: ieatnetsimv7035-96;UP;SOFT;1;PING OK
[1576526759] HOST ALERT: ieatnetsimv9041-17;UP;SOFT;1;PING OK
[1576526845] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1576526845] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN DSC NODES !
[1576526908] HOST ALERT: ieatnetsimv5116-78;DOWN;SOFT;1;PING NOK
[1576526970] HOST ALERT: ieatnetsimv5116-78;UP;SOFT;1;PING OK
[1576527095] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576527215] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576527231] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1576527374] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576527424] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1576527617] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1576528445] HOST ALERT: ieatnetsimv9034-93;DOWN;SOFT;1;PING NOK
[1576528445] HOST ALERT: ieatnetsimv9034-55;DOWN;SOFT;1;PING NOK
[1576528507] HOST ALERT: ieatnetsimv9034-93;UP;SOFT;1;PING OK
[1576528507] HOST ALERT: ieatnetsimv9034-55;UP;SOFT;1;PING OK
[1576528989] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576529112] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576529233] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576529697] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576529713] HOST ALERT: ieatnetsimv7035-66;DOWN;SOFT;1;PING NOK
[1576529775] HOST ALERT: ieatnetsimv7035-66;UP;SOFT;1;PING OK
[1576529822] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (80)
[1576529962] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576530089] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576530212] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576530432] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (111) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576530489] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1097) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576530563] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (93) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576530618] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (941) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576530687] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (17) IS LOWER THAN 5% OF ENM NODES (1200)
[1576530746] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (92) IS LOWER THAN 5% OF ENM NODES (14035)
[1576530850] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576530971] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576531025] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1576531092] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576531146] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1576531267] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1576531524] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1576531645] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1576531766] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1576531850] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576531872] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;SOFT;1;CRITICAL- san_alert_healthcheck FAILED
[1576531973] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576531993] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;SOFT;2;CRITICAL- san_alert_healthcheck FAILED
[1576532097] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576532113] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;HARD;3;CRITICAL- san_alert_healthcheck FAILED
[1576532118] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1576532167] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1576532238] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1576532288] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1576532359] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1576532536] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1576532970] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576533090] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1576533129] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1576533205] HOST ALERT: ieatnetsimv5116-47;DOWN;SOFT;1;PING NOK
[1576533267] HOST ALERT: ieatnetsimv5116-47;UP;SOFT;1;PING OK
[1576533456] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576533627] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1576533753] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (27) IS HIGHER THAN 20% OF ENM NODES (80)
[1576533761] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576533824] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1576533911] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576533931] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576534037] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (29) IS HIGHER THAN 20% OF ENM NODES (80)
[1576534166] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.21%
[1576534207] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (24) IS HIGHER THAN 20% OF ENM NODES (80)
[1576534222] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1576534291] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (592) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576534317] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 51.45%
[1576534505] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (11) IS LOWER THAN 5% OF ENM NODES (1200)
[1576534529] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576534574] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1576534574] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576534836] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576534927] SERVICE ALERT: ieatnetsimv5116-92;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576534991] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1576535021] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576535048] SERVICE ALERT: ieatnetsimv5116-92;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1576535172] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1576535206] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576535354] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1576535700] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1576535741] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1576535818] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1576536194] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1576536204] HOST ALERT: ieatnetsimv5116-111;DOWN;SOFT;1;PING NOK
[1576536266] HOST ALERT: ieatnetsimv5116-111;UP;SOFT;1;PING OK
[1576536372] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1576536434] SERVICE FLAPPING ALERT: ieatnetsimv7035-28;check netsim disk;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1576536434] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;OK;HARD;3;OK
[1576536551] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1576536727] HOST ALERT: ieatnetsimv5116-100;DOWN;SOFT;1;PING NOK
[1576536727] HOST ALERT: ieatnetsimv017-19;DOWN;SOFT;1;PING NOK
[1576536789] HOST ALERT: ieatnetsimv017-19;UP;SOFT;1;PING OK
[1576536789] HOST ALERT: ieatnetsimv5116-100;UP;SOFT;1;PING OK
[1576536909] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576537072] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576537196] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576537300] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576537424] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576537430] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576537555] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576538527] HOST ALERT: ieatnetsimv017-25;DOWN;SOFT;1;PING NOK
[1576538590] HOST ALERT: ieatnetsimv017-25;UP;SOFT;1;PING OK
[1576538697] HOST ALERT: ieatnetsimv7035-19;DOWN;SOFT;1;PING NOK
[1576538759] HOST ALERT: ieatnetsimv7035-19;UP;SOFT;1;PING OK
[1576539000] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576539016] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1576539016] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1576539045] HOST ALERT: ieatnetsimv7035-110;DOWN;SOFT;1;PING NOK
[1576539104] HOST ALERT: ieatnetsimv7035-110;UP;SOFT;1;PING OK
[1576539374] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1576540685] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)