[1568934411] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;OK;HARD;3;OK
[1568937363] SERVICE ALERT: ieatnetsimv5116-03;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.02%
[1568937514] SERVICE ALERT: ieatnetsimv5116-03;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.53%
[1568939033] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568948812] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1568948932] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1568949053] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1568958834] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1568959018] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1568959204] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1568959853] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;OK;HARD;3;OK
[1568963056] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.45%
[1568963207] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 28.91%
[1568965881] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568965884] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568966032] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568966035] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568966183] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1568966186] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1568966254] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568966273] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568966279] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568966282] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.5%
[1568966299] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568966405] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568966424] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568966433] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568966442] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.59%
[1568966451] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.99%
[1568966556] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1568966575] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 90.53%
[1568966584] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 71.6%
[1568966593] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.13%
[1568966601] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;OK;HARD;3;OK- CPU IS USED 78.59%
[1568966628] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.17%
[1568966656] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.58%
[1568966778] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 36.04%
[1568966806] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 42.89%
[1568966863] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (4209) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1568967275] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1568967468] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1568967540] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568967570] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568967660] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1568968256] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4093) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1568969196] SERVICE ALERT: ieatnetsimv5116-29;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568969477] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1568969656] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1568969814] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.22%
[1568969816] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.64%
[1568969835] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1568970186] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;OK;HARD;3;OK- CPU IS USED 31.81%
[1568970205] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;OK;HARD;3;OK- CPU IS USED 27.4%
[1568970223] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;OK;HARD;3;OK- CPU IS USED 47.48%
[1568970654] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1568970774] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1568970895] SERVICE FLAPPING ALERT: ieatnetsimv5116-96;check netsim disk;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1568970895] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1568971558] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1568971628] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2129) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568971858] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2216) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568971869] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1568972034] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2278) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568972664] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568972670] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (19614) IS BETWEEN 80% AND 95% OF ENM NODES (22115)
[1568972823] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568972962] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568973105] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568973105] SERVICE ALERT: ieatnetsimv5116-36;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.03%
[1568973256] SERVICE ALERT: ieatnetsimv5116-36;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 32.71%
[1568973257] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.52%
[1568973408] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.81%
[1568973852] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1568974158] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.37%
[1568974170] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1568974309] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.9%
[1568974352] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1568975091] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568975230] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568975534] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1568975720] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1568975905] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1568976064] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1568976192] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1568976320] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1568976335] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1568976518] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1568976700] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.34%
[1568976851] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.89%
[1568977001] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;OK;HARD;3;OK- CPU IS USED 14.43%
[1568977415] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.07%
[1568977565] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 31.08%
[1568977584] SERVICE ALERT: ieatnetsimv5116-37;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568977716] SERVICE ALERT: ieatnetsimv5116-37;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568977848] SERVICE ALERT: ieatnetsimv5116-37;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568978501] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1568978684] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1568982467] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1568982649] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1568982869] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1872) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1568983005] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1872) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1568983142] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1872) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1568984632] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1568984815] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1568984961] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1568984961] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1568985259] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.2%
[1568985410] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 52.41%
[1568986170] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568986304] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568986440] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568986547] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568986677] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568986786] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1568986797] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1568986917] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1568986981] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1568987128] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1568987164] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1568987255] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1568987340] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568987399] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1568987481] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568987620] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568987747] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1253) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1568987885] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1253) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1568988017] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1253) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1568988570] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1568988570] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1568989207] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1568989327] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1568990381] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2334) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568990527] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2392) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1568991231] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568991571] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1568991689] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3960) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1568991828] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3960) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1568991975] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3960) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1568992356] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (3972) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1568992839] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (3972) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1568993258] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (3972) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569000731] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (1694) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569004657] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (3972) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569007161] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90%
[1569007311] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 28.77%
[1569012035] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1569017698] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)