[1572131581] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572131706] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572131834] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572138943] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1572139040] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572139070] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1572139165] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572139197] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1572139291] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572140143] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;OK;HARD;3;OK- NUMBER OF PM FILES (2172) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1572142804] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;OK;HARD;3;OK- NUMBER OF PM FILES (4500) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (20)
[1572142838] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1216) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1572142897] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572143023] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572143150] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (5118) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572143750] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (1086) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1572143878] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (1086) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1572144006] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1086) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1572146443] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1572146757] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1572154305] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1572158354] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0)
[1572165564] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572165694] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572165823] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572168357] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572168489] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572168617] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572172225] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572172306] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572173387] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572173391] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572173428] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572173429] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572173451] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572173454] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572173539] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572173542] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572173580] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572173581] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572173603] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572173606] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572173689] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 11.53%
[1572173693] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.66%
[1572173733] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572173734] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572173755] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572173759] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572173761] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572173765] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572173772] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572173773] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572173914] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572173918] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572173924] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572173924] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572174069] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572174074] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572174078] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572174078] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572174294] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572175912] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572176038] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572176164] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572176644] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13770) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572177364] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.34%
[1572177364] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.87%
[1572177385] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.74%
[1572177389] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.59%
[1572177700] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.96%
[1572177705] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.85%
[1572177711] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 16.23%
[1572177712] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.92%
[1572179438] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572179566] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572179770] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572179896] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572180022] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572180254] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572180383] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572180512] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572183293] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572183419] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572183545] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572186327] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1572186547] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;SOFT;2;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1572187227] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572187354] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572191080] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572191205] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572191331] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572194942] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572195068] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572195193] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572195901] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572196026] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572196151] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572202133] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13770) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572205742] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572205884] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572206026] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14011)
[1572209598] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572209723] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572209849] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572210556] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572210681] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572210806] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572210932] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572211058] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572213454] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572213580] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572213706] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572213831] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572213956] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572214081] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572214208] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572214664] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572214789] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572214913] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572217814] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572218099] SERVICE ALERT: ieatwlvm5116;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE