[1567810942] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79.74%
[1567815955] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567816106] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.68%
[1567818382] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.47%
[1567818533] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.88%
[1567818683] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.11%
[1567818834] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.59%
[1567818838] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM CPP NODES (0)
[1567819769] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567821438] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567822358] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567823278] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized cpp;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1567823278] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567827169] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567827640] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567830554] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1567833154] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567834074] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;(Service check timed out after 800.00 seconds)
[1567834995] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1567835680] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.38%
[1567835831] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.94%
[1567835981] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.94%
[1567836132] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.73%
[1567836283] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.16%
[1567836985] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.09%
[1567837135] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.18%
[1567838539] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.09%
[1567838689] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.44%
[1567839331] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1567840885] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.35%
[1567841036] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.01%
[1567841931] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567842851] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567843771] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1567843771] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567850713] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.19%
[1567850864] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.43%
[1567851720] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1567854320] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567855240] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567856160] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567858278] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized cpp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567858708] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1567859186] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.33%
[1567859337] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.99%
[1567859488] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.81%
[1567859638] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.65%
[1567861309] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.38 seconds)
[1567862229] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567863149] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567867437] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1567870037] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567870958] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1567871878] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567875892] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.44%
[1567876043] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.86%
[1567883415] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.33%
[1567883565] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.52%
[1567883717] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.78%
[1567883867] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.5%
[1567884018] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.87%
[1567884170] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.46%
[1567888257] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.81%
[1567888408] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 3.18%
[1567888559] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.22%
[1567888709] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 43.43%
[1567892188] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.73%
[1567892339] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.14%
[1567895890] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567896011] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567896071] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567896192] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567896252] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567896372] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED