[1567380747] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (16930) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567381703] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.23%
[1567381853] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 4.18%
[1567382099] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567382249] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79.23%
[1567382351] SERVICE FLAPPING ALERT: ieatnetsimv5116-68;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567388081] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1567388081] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17113) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567390762] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1567391123] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1567391454] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1567391773] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (16890) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567391817] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567391956] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567392096] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567392850] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567393001] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.24%
[1567394009] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.05%
[1567394159] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 33.32%
[1567394514] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567395472] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17123) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567397042] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.73%
[1567397114] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567397192] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79.86%
[1567397596] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567399166] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567399698] SERVICE ALERT: ieatnetsimv5116-73;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.02%
[1567399849] SERVICE ALERT: ieatnetsimv5116-73;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.48%
[1567402431] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1567402492] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1567402570] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567403999] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.92%
[1567404150] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.63%
[1567406462] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17249) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567407905] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567407923] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;OK;HARD;3;OK
[1567408059] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567408209] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.45%
[1567410166] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (17018) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567410255] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.38%
[1567410406] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 36.14%
[1567411314] SERVICE ALERT: ieatnetsimv5116-74;check netsim disk;OK;HARD;3;OK
[1567411333] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;OK;HARD;3;OK
[1567411340] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;OK;HARD;3;OK
[1567411840] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.88%
[1567412917] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1567413148] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567413381] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567413849] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17329) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567416214] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567416365] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567416516] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.47%
[1567417439] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567417534] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (17053) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567418195] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567418384] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567418597] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567418754] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567420147] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;HARD;3;OK- CPU IS USED 20.82%
[1567420488] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567421220] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17205) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567422948] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.23 seconds)
[1567423339] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567424259] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567425179] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.79 seconds)
[1567425292] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.69%
[1567425443] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.2%
[1567425594] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96%
[1567426099] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567426372] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567426571] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567426769] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567426817] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567426902] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1567426973] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567427158] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567427344] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567427822] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567427929] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6729) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1567428078] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6527) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1567428116] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567428153] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567428222] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6613) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1567428742] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567429210] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567429225] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;OK;HARD;3;OK- CPU IS USED 55.52%
[1567430045] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (7120) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1567430445] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1567432679] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567432690] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.73%
[1567432842] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.09%
[1567432994] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.62%
[1567433145] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.52%
[1567433299] SERVICE ALERT: ieatnetsimv5116-63;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.01%
[1567433910] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.05%
[1567434060] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 30.73%
[1567434676] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567434862] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567435050] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567435279] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567435732] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567436076] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.93%
[1567436227] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 27.45%
[1567436378] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.48%
[1567436529] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 47.89%
[1567436653] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1567437573] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567437811] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.13%
[1567437961] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 12.04%
[1567438031] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567438716] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1567438716] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567440631] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567440970] SERVICE ALERT: ieatnetsimv5116-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.78%
[1567441120] SERVICE ALERT: ieatnetsimv5116-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 70.98%
[1567441551] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.31 seconds)
[1567442007] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567443353] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.14 seconds)
[1567444144] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.88%
[1567444271] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.18%
[1567444273] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567444294] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79.88%
[1567444421] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.15%
[1567444572] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 91.88%
[1567445193] SERVICE FLAPPING ALERT: ieatnetsimv5116-62;check netsim nodes;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1567445193] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567448168] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.05%
[1567448207] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567448318] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 42.92%
[1567448487] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567449407] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567449854] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567450775] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1567451695] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.05 seconds)
[1567451802] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.68%
[1567452069] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.67%
[1567452219] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.09%
[1567452370] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.84%
[1567452615] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.02 seconds)
[1567452627] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567453873] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.63%
[1567454023] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.17%
[1567454519] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567454678] SERVICE FLAPPING ALERT: ieatnetsimv5116-68;check netsim nodes;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1567454678] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567454680] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.46%
[1567454700] SERVICE ALERT: ieatlms5735-1;check node fs;OK;SOFT;2;OK
[1567454831] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.45%
[1567454981] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.22%
[1567455132] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 28.23%
[1567457503] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567457684] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567457865] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567459036] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567459187] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.7%
[1567459337] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.47%
[1567459726] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567460628] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1567461544] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567462968] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.21%
[1567463200] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.15%
[1567463351] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 70.64%
[1567464804] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.63%
[1567464954] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 23.05%