[1568156873] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.99%
[1568157024] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.59%
[1568157176] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99%
[1568157998] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.25%
[1568158148] SERVICE FLAPPING ALERT: ieatnetsimv7035-13;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1568158148] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.99%
[1568158299] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.64%
[1568158450] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.71%
[1568158600] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.02%
[1568159726] HOST ALERT: ieatrcxb3839;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.45 ms
[1568159769] HOST ALERT: ieatrcxb2446;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.16 ms
[1568159772] HOST ALERT: ieatrcxb3777;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.30 ms
[1568159773] HOST ALERT: ieatrcxb3778;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.74 ms
[1568159780] HOST ALERT: ieatrcxb3911;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.30 ms
[1568159792] HOST ALERT: ieatrcxb2445;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.27 ms
[1568159801] HOST ALERT: ieatrcxb3805;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.38 ms
[1568159801] HOST ALERT: ieatrcxb3807;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1568159804] HOST ALERT: ieatrcxb3804;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.19 ms
[1568159810] HOST ALERT: ieatrcxb3914;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.17 ms
[1568159812] HOST ALERT: ieatrcxb4416;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.33 ms
[1568159831] HOST ALERT: ieatrcxb6446;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.80 ms
[1568159923] HOST ALERT: ieatrcxb3806;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.14 ms
[1568159966] HOST ALERT: ieatrcxb3776;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 1.23 ms
[1568160807] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 82.37%
[1568163116] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.51%
[1568163274] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.43%
[1568163424] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.58%
[1568165830] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.75%
[1568165981] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568166132] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.04%
[1568168659] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.36%
[1568168810] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.13%
[1568169762] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.4%
[1568169913] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.3%
[1568170063] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.74%
[1568170655] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.33%
[1568170806] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.52%
[1568172561] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.32%
[1568172712] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568172863] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 91.7%
[1568174559] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.67%
[1568174709] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.83%
[1568174860] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.14%
[1568175010] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.01%
[1568175238] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.72%
[1568175388] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.87%
[1568175539] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.09%
[1568176493] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.98%
[1568178760] SERVICE FLAPPING ALERT: ieatnetsimv7035-08;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1568179329] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3907) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1568179461] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3907) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1568179594] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3907) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1568180124] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.07%
[1568180274] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.62%
[1568180425] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.94%
[1568180894] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568181044] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.3%
[1568181195] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.02%
[1568181811] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2362) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1568181944] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2362) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1568182070] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2362) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1568182361] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.59%
[1568182511] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.09%
[1568182662] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.4%
[1568186370] SERVICE FLAPPING ALERT: ieatnetsimv7035-12;check netsim cpu;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1568186370] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 78.07%
[1568187014] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (1296) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1568187071] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1568187140] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (1296) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1568187265] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1296) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1568187655] SERVICE FLAPPING ALERT: ieatnetsimv7035-10;check netsim cpu;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1568187655] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.83%
[1568188988] SERVICE ALERT: ieatnetsimv7035-87;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568190871] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1458) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1568195626] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.48%
[1568195776] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.51%
[1568202085] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.12%
[1568202236] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.67%
[1568202386] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.78%
[1568202537] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.81%
[1568204292] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.59%
[1568204443] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.17%
[1568204594] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 91.27%
[1568206726] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.88%
[1568206880] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.91%
[1568207030] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 91.59%
[1568207985] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1568208170] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1568208224] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.37%
[1568208356] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1568208721] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1568208849] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1568208977] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1568209405] SERVICE ALERT: ieatlms3901;check snapshots;OK;HARD;3;OK- NO ENM SNAPSHOTS ARE PRESENT
[1568209888] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.58%
[1568210038] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.56%
[1568210189] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.67%
[1568210661] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.78%
[1568211603] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.97%
[1568211756] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.19%
[1568211906] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.86%
[1568212585] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1568213519] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.47%
[1568213670] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.98%
[1568213820] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568213971] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.74%
[1568214122] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.63%
[1568214272] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.22%
[1568214292] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1568214423] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.18%
[1568214442] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.53%
[1568214573] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.09%
[1568214724] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.38%
[1568215454] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.89%
[1568215605] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.73%
[1568215620] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1568215755] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.6%
[1568217486] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1568217672] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1568217857] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1568217993] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1568218121] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1568218249] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1568221793] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.91%
[1568221943] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.55%
[1568222094] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568222245] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.81%
[1568222395] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.03%
[1568222986] SERVICE FLAPPING ALERT: ieatnetsimv7035-08;check netsim cpu;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1568222986] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.35%
[1568224622] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.96%
[1568224772] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.17%
[1568224923] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.55%
[1568225073] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.19%
[1568225555] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.48%
[1568226026] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.5%
[1568226177] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.38%
[1568226327] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.57%
[1568226337] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.89%
[1568226487] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.93%
[1568226616] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.22%
[1568226639] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.79%
[1568226767] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.8%
[1568226789] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.3%
[1568226918] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.41%
[1568226940] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.49%
[1568227068] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.47%
[1568227090] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.83%
[1568227219] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.32%
[1568227369] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.5%
[1568227520] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.83%
[1568227670] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.79%
[1568227821] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 91.84%
[1568229958] SERVICE FLAPPING ALERT: ieatnetsimv7035-13;check netsim cpu;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1568229958] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.99%
[1568233588] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.59%
[1568233739] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.71%
[1568233891] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.82%
[1568234041] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.26%
[1568234195] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.94%
[1568234346] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.99%
[1568234441] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.51%
[1568234460] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.62%
[1568234498] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568234592] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.25%
[1568234611] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 42.08%
[1568234648] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.94%
[1568234743] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.21%
[1568234762] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.66%
[1568234799] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568234912] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 43.7%
[1568234949] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.05%
[1568235052] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.24%
[1568242301] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.95%
[1568242451] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.86%