[1567292779] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567294399] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (5030) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1567294408] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.65%
[1567294546] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6699) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1567294559] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.38%
[1567294693] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (6943) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1567294710] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.38%
[1567297725] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567297914] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567298093] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567299209] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.5%
[1567299359] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82%
[1567299660] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567300552] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567300779] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567301495] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567301658] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567301699] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.02 seconds)
[1567301735] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567301816] SERVICE FLAPPING ALERT: ieatnetsimv5116-68;check netsim nodes;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1567301816] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567301941] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 66.62%
[1567302619] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567302952] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.91%
[1567303103] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.83%
[1567303555] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.21%
[1567303706] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 71.36%
[1567305207] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567306690] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1567306811] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1567306932] SERVICE ALERT: ieatnetsimv5116-81;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1567307807] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567308728] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1567309648] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567310454] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.03%
[1567310604] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 40.86%
[1567312200] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567312772] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.98%
[1567312923] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.83%
[1567314202] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.48%
[1567314355] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.07%
[1567314514] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 64.97%
[1567314952] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567315133] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567315314] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567315720] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.24%
[1567315870] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.87%
[1567316022] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;HARD;3;OK- CPU IS USED 26.86%
[1567317680] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567317835] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.47%
[1567317859] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567317935] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567317989] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567318034] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567318116] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567318139] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 79.73%
[1567318172] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567318265] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.98%
[1567318297] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567318416] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 74.19%
[1567318977] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;HARD;3;OK- CPU IS USED 61.1%
[1567320154] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567320338] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567320584] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567321775] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567321927] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.18%
[1567321958] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567322080] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567322515] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567322575] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1567323874] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.27%
[1567324024] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 44.39%
[1567324473] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567324603] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567324645] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567324710] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.63%
[1567324746] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567324785] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567324860] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.21%
[1567324861] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567324928] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567324968] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567325011] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.78%
[1567325109] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567325711] SERVICE FLAPPING ALERT: ieatnetsimv7035-12;check netsim cpu;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1567325711] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.87%
[1567329346] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567329496] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.52%
[1567333250] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567333410] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567333562] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.65%
[1567337165] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567338085] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567339005] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567340792] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.11%
[1567342117] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567342268] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.57%
[1567344435] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567344585] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.57%
[1567344738] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567344892] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.54%
[1567345043] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.82%
[1567345206] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567347038] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567347189] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567347350] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567348684] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567348711] SERVICE FLAPPING ALERT: ieatnetsimv5116-66;check netsim nodes;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1567348834] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.97%
[1567355652] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.88%
[1567355803] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.28%
[1567356186] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.07%
[1567356336] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.92%
[1567356488] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 81.88%
[1567356575] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.78%
[1567356726] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 31.97%
[1567366167] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (16966) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567369857] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17277) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567370619] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567370800] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567370924] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567370981] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567371075] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.66%
[1567371234] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567373954] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.19%
[1567374104] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.27%
[1567374836] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567374864] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.29%
[1567375017] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567375197] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567376442] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567377058] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK