[1567036606] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.09%
[1567036698] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.56%
[1567036709] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.01%
[1567036757] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.23%
[1567036848] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.84%
[1567036860] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.42%
[1567037643] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.92%
[1567037794] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.38%
[1567037944] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.76%
[1567038095] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 44.84%
[1567038678] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567038835] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567038993] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567039561] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567040481] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567041401] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567041642] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1567045446] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.46%
[1567045597] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 11.47%
[1567049323] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.98%
[1567049473] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.84%
[1567049624] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.72%
[1567049695] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567049699] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567049723] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized nodes cpp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567049784] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized nodes com/ecim;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567049786] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1567052948] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.72%
[1567053100] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567053251] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 34.16%
[1567053254] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.48%
[1567053405] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.88%
[1567054892] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567055086] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567055286] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567055986] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567056899] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567056914] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1567057020] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567057035] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1567057141] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567057156] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1567057777] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567057898] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567058019] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567058406] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1567058439] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567058526] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1567058560] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567058586] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567058647] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1567058681] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1567059506] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567059884] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567060426] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567060543] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567060759] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1567060801] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1567060801] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567060855] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1567060855] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1567060964] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567062325] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1567062763] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2467) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1567064159] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13259) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567064346] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.68%
[1567064496] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.91%
[1567064885] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2178) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1567066568] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567066755] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567066943] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567067648] SERVICE ALERT: ieatnetsimv5116-85;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567067789] SERVICE ALERT: ieatnetsimv5116-85;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567067934] SERVICE ALERT: ieatnetsimv5116-85;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567068529] SERVICE ALERT: ieatnetsimv5116-85;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567068673] SERVICE FLAPPING ALERT: ieatnetsimv5116-71;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567068680] SERVICE ALERT: ieatnetsimv5116-85;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.71%
[1567068812] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567068900] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (5002) IS HIGHER THAN 95% OF ENM SNMP NODES (5018)
[1567068949] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567069135] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567069320] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567069761] SERVICE ALERT: ieatnetsimv5116-85;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567071532] SERVICE FLAPPING ALERT: ieatnetsimv5116-70;check netsim nodes;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1567071532] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567071851] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567072002] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.45%
[1567072152] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 81.45%
[1567074476] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.12%
[1567074627] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.5%
[1567074785] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567075023] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.27%
[1567075156] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.51%
[1567075306] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.06%
[1567076686] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567078768] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1567078896] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1567079024] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1567079725] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1567079845] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1567079884] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567079966] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1567080290] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567080421] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567080552] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567081542] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1567081990] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1567081990] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567083567] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;OK;HARD;3;OK
[1567085616] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 62.79%
[1567085655] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567085805] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567085841] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567085987] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567086026] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567086113] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1567086169] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1567086292] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1567086470] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1567086582] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567086733] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.99%
[1567087399] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.28%
[1567087550] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.65%
[1567087701] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.13%
[1567087852] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 51.72%
[1567089254] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567089372] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567089390] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567089412] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567089505] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567089571] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567089643] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567089857] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1567089972] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2980) IS LOWER THAN 80% OF ENM CPP NODES (4224)
[1567090007] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.44%
[1567090091] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1567090236] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1567090310] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567090374] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1567091230] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.02 seconds)
[1567091631] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1567091631] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567091791] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (131) IS HIGHER THAN 95% OF ENM CPP NODES (0)
[1567093262] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567093608] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3688) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1567093755] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3688) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1567093902] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3688) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1567094080] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567094087] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.58%
[1567094168] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13334) IS FROM 80% TO 95% OF ENM COMECIM NODES (14049)
[1567094238] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.01%
[1567094268] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567094314] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13350) IS HIGHER THAN 95% OF ENM COMECIM NODES (14049)
[1567094389] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 81.4%
[1567094456] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567094613] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.27%
[1567094763] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 15.79%
[1567094821] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.81%
[1567094972] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.55%
[1567095734] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized cpp;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1567095734] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3226) IS LOWER THAN 80% OF ENM CPP NODES (4224)
[1567096325] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1567096325] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567096761] SERVICE FLAPPING ALERT: ieatnetsimv5116-70;check netsim nodes;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1567097017] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567097358] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.8%
[1567097509] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.83%
[1567097616] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3688) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1567097661] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.69%
[1567097811] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.67%
[1567097962] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.74%
[1567098020] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.78%
[1567098113] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.51%
[1567098171] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.26%
[1567098263] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.03%
[1567098321] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.38%
[1567098937] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567099491] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567099649] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567100402] SERVICE FLAPPING ALERT: ieatnetsimv5116-68;check netsim cpu;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1567100402] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 47.11%
[1567100576] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1879) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1567100708] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1879) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1567100843] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1961) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1567101657] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13337) IS FROM 80% TO 95% OF ENM COMECIM NODES (14049)
[1567101804] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13350) IS HIGHER THAN 95% OF ENM COMECIM NODES (14049)
[1567101894] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.12%
[1567101952] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.92%
[1567102044] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.84%
[1567102110] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.19%
[1567102260] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 82.66%
[1567102658] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1567103493] SERVICE FLAPPING ALERT: ieatnetsimv5116-68;check netsim nodes;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1567103493] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567103757] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13284) IS FROM 80% TO 95% OF ENM COMECIM NODES (14049)
[1567103854] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.77%
[1567103949] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13350) IS HIGHER THAN 95% OF ENM COMECIM NODES (14049)
[1567104004] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 9.84%
[1567104035] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567104186] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 65.56%
[1567105436] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567105587] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567105777] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567105891] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567106226] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567106385] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567106544] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567107244] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567108164] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567108174] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1567109085] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1567109491] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.03%
[1567109641] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.57%
[1567109792] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.12%
[1567109942] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.94%
[1567110093] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 82.4%
[1567111718] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567111874] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567112030] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567112050] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3075) IS LOWER THAN 80% OF ENM CPP NODES (4224)
[1567113122] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.23%
[1567113273] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.25%
[1567113730] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567113868] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3638) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1567113880] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.13%
[1567113940] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.35 seconds)
[1567114860] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.17 seconds)
[1567115163] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.29%
[1567115314] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 52.6%
[1567115780] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567116915] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567117050] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567117202] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567117474] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567117481] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2956) IS LOWER THAN 80% OF ENM CPP NODES (4224)
[1567117535] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 800.97 seconds)
[1567117631] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.87%
[1567117638] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567117781] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.05%
[1567117950] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567117970] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567118100] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.47%
[1567118251] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.19%
[1567118300] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567118322] SERVICE ALERT: ieatlms3901;check unsync nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1567118394] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567118456] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1567118559] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1567118890] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567118908] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.77%
[1567119058] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.72%
[1567119220] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567119295] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3642) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1567119314] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567119376] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567119480] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)