[1567120120] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567120140] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567121041] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;SOFT;2;(Service check timed out after 801.02 seconds)
[1567121820] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.21%
[1567121961] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567121971] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.29%
[1567122125] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.21%
[1567122276] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.88%
[1567122293] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1567122426] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.95%
[1567123957] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1567127375] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567130108] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3366) IS LOWER THAN 80% OF ENM CPP NODES (4224)
[1567130594] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567131923] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3558) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1567135537] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3075) IS LOWER THAN 80% OF ENM CPP NODES (4224)
[1567136283] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567136435] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567136587] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567138196] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.49%
[1567138231] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567138347] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.93%
[1567138497] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.84%
[1567140010] SERVICE FLAPPING ALERT: ieatnetsimv5116-68;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567140217] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.44%
[1567142128] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.89%
[1567142279] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.68%
[1567142429] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.72%
[1567142631] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567143119] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567144039] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567144056] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.06%
[1567144207] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.17%
[1567144357] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.38%
[1567144508] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.61%
[1567144513] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567144672] SERVICE ALERT: ieatnetsimv5116-74;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1567144793] SERVICE ALERT: ieatnetsimv5116-74;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1567144913] SERVICE ALERT: ieatnetsimv5116-74;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1567147264] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13259) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567147507] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567147706] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567147903] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567147991] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2178) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1567149096] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13333) IS FROM 80% TO 95% OF ENM COMECIM NODES (14049)
[1567149243] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13333) IS FROM 80% TO 95% OF ENM COMECIM NODES (14049)
[1567149389] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13342) IS FROM 80% TO 95% OF ENM COMECIM NODES (14049)
[1567149424] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.54%
[1567149574] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79.21%
[1567149660] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.55%
[1567149781] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1567151221] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13350) IS HIGHER THAN 95% OF ENM COMECIM NODES (14049)
[1567151607] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2522) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1567152544] SERVICE ALERT: ieatlms3901;check unsync nodes;OK;HARD;3;OK
[1567153158] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1567153527] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM CPP NODES (0)
[1567153757] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM COMECIM NODES (0)
[1567154349] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1567154392] SERVICE ALERT: ieatlms3901;check number nodes;OK;HARD;3;OK- NUMBER OF NODES CONFIGURED IN ENM (0) IS HIGHER THAN 95% OF NETSIM NODES (22495)
[1567154540] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1567154645] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567154771] SERVICE ALERT: ieatlms3901;check pm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1567155805] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (1273) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (0)
[1567156236] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES W/O FM HB FAILURES (3114) IS HIGHER THAN 95% OF SNMP NODES WITH FM SUPERVISION ENABLED (0)
[1567156542] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567156573] SERVICE ALERT: ieatlms3901;check pm collection statistical;OK;HARD;3;OK- NUMBER OF PM FILES (17564) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1567156873] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567156891] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.05%
[1567157041] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.4%
[1567157072] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567157192] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.73%
[1567158337] SERVICE FLAPPING ALERT: ieatnetsimv5116-68;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1567158675] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1567159059] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.74%
[1567159210] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.79%
[1567159361] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.08%
[1567159515] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.46%
[1567159666] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.93%
[1567160394] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (17769) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567160464] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23837) IS BETWEEN 80% AND 95% OF ENM NODES (25422)
[1567160605] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (17769) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567160636] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567160717] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23837) IS BETWEEN 80% AND 95% OF ENM NODES (25422)
[1567160880] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17769) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21330)
[1567161100] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23837) IS BETWEEN 80% AND 95% OF ENM NODES (25422)
[1567161556] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567161609] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1567161833] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1567161984] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1567162476] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567162524] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567162708] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567162802] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11402) IS FROM 80% TO 95% OF ENM COMECIM NODES (12046)
[1567162892] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567162978] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (11682) IS HIGHER THAN 95% OF ENM COMECIM NODES (12046)
[1567162995] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567163183] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567163305] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567163357] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567163371] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567163456] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.99%
[1567163747] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (1301) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (7152)
[1567164125] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.41%
[1567164179] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1567164275] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.25%
[1567164358] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1567164537] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1567164668] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1567164755] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567164938] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567165121] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567165589] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1567166281] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (22799) IS BETWEEN 80% AND 95% OF ENM NODES (25422)
[1567166617] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (25384) IS HIGHER THAN 95% OF ENM NODES (25422)
[1567166846] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567166962] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567166984] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567167027] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567167062] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567167087] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567167167] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567167208] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567167211] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567167243] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567167349] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567167424] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567167950] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized cpp;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567168482] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.74%
[1567168632] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.45%
[1567169213] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567169340] SERVICE ALERT: ieatnetsimv5116-89;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567169395] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567169396] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4501) IS FROM 80% TO 95% OF ENM SNMP NODES (5018)
[1567169491] SERVICE ALERT: ieatnetsimv5116-89;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.82%
[1567169579] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1567169599] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4519) IS FROM 80% TO 95% OF ENM SNMP NODES (5018)
[1567169642] SERVICE ALERT: ieatnetsimv5116-89;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.83%
[1567169775] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4510) IS FROM 80% TO 95% OF ENM SNMP NODES (5018)
[1567169888] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1567170015] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1567170081] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1567170139] SERVICE ALERT: ieatnetsimv5116-79;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1567170202] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1567170322] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1567170869] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1567171084] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567171442] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567171597] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4816) IS HIGHER THAN 95% OF ENM SNMP NODES (5018)
[1567171622] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.38%
[1567171773] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.98%
[1567172736] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.4%
[1567172893] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.95%
[1567173048] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567173200] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567173273] SERVICE ALERT: ieatnetsimv5116-89;check netsim cpu;OK;HARD;3;OK- CPU IS USED 43.34%
[1567173305] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567173350] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 64.13%
[1567173489] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567173673] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567174233] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567175153] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.34 seconds)
[1567175524] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.98%
[1567175536] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567175588] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567175675] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.35%
[1567175719] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567175826] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.42%
[1567175902] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567175976] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.46%
[1567176986] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567177136] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 63.49%
[1567177534] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11346) IS FROM 80% TO 95% OF ENM COMECIM NODES (12046)
[1567177688] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11162) IS FROM 80% TO 95% OF ENM COMECIM NODES (12046)
[1567177765] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567177834] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567177845] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11349) IS FROM 80% TO 95% OF ENM COMECIM NODES (12046)
[1567177948] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567177996] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567178028] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567178222] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567178776] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567178916] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567178959] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED %
[1567179100] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;OK;HARD;3;OK
[1567179679] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (11575) IS HIGHER THAN 95% OF ENM COMECIM NODES (12046)
[1567179691] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567179727] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.07%
[1567179812] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567179832] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567179877] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.84%
[1567179931] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567179933] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567179953] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567180074] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1567180085] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567180114] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567180175] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1567180205] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567180296] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1567180297] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1567180326] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567180416] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1567180660] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4398) IS FROM 80% TO 95% OF ENM SNMP NODES (5018)
[1567180749] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567180840] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4433) IS FROM 80% TO 95% OF ENM SNMP NODES (5018)
[1567180870] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567180990] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567181013] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1567181013] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4440) IS FROM 80% TO 95% OF ENM SNMP NODES (5018)
[1567181126] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1567181246] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1567181367] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1567181799] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567181896] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1567181936] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567182161] SERVICE FLAPPING ALERT: ieatlms3901;check vcs cluster hc;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1567182161] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567182187] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1567182187] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567182842] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4912) IS HIGHER THAN 95% OF ENM SNMP NODES (5018)
[1567182853] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1567184117] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1567185059] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1567185216] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567185824] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567186007] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567188264] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4410) IS FROM 80% TO 95% OF ENM SNMP NODES (5018)
[1567188409] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4410) IS FROM 80% TO 95% OF ENM SNMP NODES (5018)
[1567188550] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4439) IS FROM 80% TO 95% OF ENM SNMP NODES (5018)
[1567188639] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567188823] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567188947] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567189073] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical DSC;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1567189073] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1567189169] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567189364] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567189557] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567189871] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1567190371] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4967) IS HIGHER THAN 95% OF ENM SNMP NODES (5018)
[1567190788] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567190828] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.98%
[1567190978] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.9%
[1567191248] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567191428] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;SOFT;2;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567191431] SERVICE FLAPPING ALERT: ieatlms5735-1;check system service hc;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1567191431] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1567191708] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567192628] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567195208] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567195389] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;SOFT;2;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567195570] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567195751] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567195931] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567196129] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567196310] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567196491] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567196542] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567199592] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567199640] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567199999] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567200152] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1567200152] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1567200318] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567200505] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567200699] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567200893] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567200943] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1567201186] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.94%
[1567201337] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.85%
[1567201488] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 27.2%
[1567201737] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567201858] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567201863] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567201979] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1567202118] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1567202239] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1567202326] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567202360] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1567202613] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567202654] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567202775] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567202783] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567202896] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567202950] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567203059] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1567203180] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1567203193] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567203262] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567203301] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1567203314] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567203400] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567203434] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567203521] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567203545] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1567203580] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567203628] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1567203641] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567203666] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1567203749] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1567203787] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1567203870] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1567205119] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.23%
[1567205270] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.3%
[1567205295] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567205507] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1567205507] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567205640] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1567205640] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED