[1570408172] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1570408969] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570409425] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570410168] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570418017] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1570418028] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570418442] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570418518] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570419251] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570421712] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1570425128] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570425889] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570429379] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570429559] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570429740] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570430246] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570431601] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570431781] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570431962] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1570432573] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570432753] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570432934] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570433522] SERVICE ALERT: ieatnetsimv7035-12;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570433651] SERVICE ALERT: ieatnetsimv7035-12;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570433780] SERVICE ALERT: ieatnetsimv7035-12;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570433903] SERVICE ALERT: ieatnetsimv7035-65;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570434031] SERVICE ALERT: ieatnetsimv7035-65;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570434160] SERVICE ALERT: ieatnetsimv7035-65;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570435595] SERVICE ALERT: ieatnetsimv7035-12;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570436099] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570436251] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570436403] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570437422] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570437775] SERVICE ALERT: ieatnetsimv7035-65;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570438079] SERVICE ALERT: ieatlms3901;check pm collection statistical;OK;HARD;3;OK- NUMBER OF PM FILES (21657) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570438396] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1570440034] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.14%
[1570441083] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570441263] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570441444] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570442058] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570442238] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570442420] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1570442420] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570442434] SERVICE ALERT: ieatnetsimv017-30;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570442582] SERVICE ALERT: ieatnetsimv017-30;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570442730] SERVICE ALERT: ieatnetsimv017-30;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570442792] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570442941] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570443086] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570443305] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570443485] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570443666] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1570443666] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1570444725] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570444910] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570445095] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570446961] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570447056] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1570447385] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1570447743] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1570448055] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (40) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570448181] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (40) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570448306] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570450432] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (111) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570450557] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (72) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570450681] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (72) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570450848] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (1293) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570450975] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (1494) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570451101] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1936) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570451577] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.19%
[1570451727] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 52.98%
[1570452094] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570452277] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1570453193] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570453332] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570453474] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570453518] SERVICE ALERT: ieatnetsimv5116-41;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570453666] SERVICE ALERT: ieatnetsimv5116-41;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570453803] SERVICE ALERT: ieatnetsimv5116-41;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570456026] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570456212] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570456397] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1570456397] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570456632] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.32%
[1570456646] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.82%
[1570456668] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570456781] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570456782] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.4%
[1570456797] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 70.22%
[1570456812] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570456954] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570457016] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570457183] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570457257] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570457431] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570457633] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570457772] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570458262] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570458982] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570459082] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570459182] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570459221] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570459365] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570459370] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570459614] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570459754] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570459829] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570459885] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570459988] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570460148] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570460566] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570460793] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570461165] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6075) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1570461311] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6075) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1570461453] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6075) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1570461697] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570461991] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570462667] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (5772) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570463425] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (5622) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570464177] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (5772) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570474327] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1570474881] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570475579] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570479895] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (5772) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570488560] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (19873) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570488765] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (19873) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570488963] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (19873) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21948)