[1570492527] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570492662] SERVICE ALERT: ieatlms3901;check pm collection statistical;OK;HARD;3;OK- NUMBER OF PM FILES (20907) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570498342] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1570498342] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (1923) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570498835] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570500650] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570500780] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570500911] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570502617] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (5772) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570506119] HOST ALERT: ieatnetsimv5116-99;DOWN;SOFT;1;PING NOK
[1570506182] HOST ALERT: ieatnetsimv5116-99;UP;SOFT;1;PING OK
[1570506310] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570508274] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (5015) IS LOWER THAN 80% OF ENM CPP NODES (6589)
[1570510097] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5771) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1570510316] SERVICE ALERT: ieatnetsimv5116-40;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.4%
[1570510467] SERVICE ALERT: ieatnetsimv5116-40;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.48%
[1570510556] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (5773) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570510618] SERVICE ALERT: ieatnetsimv5116-40;check netsim cpu;OK;HARD;3;OK- CPU IS USED 24.27%
[1570520428] SERVICE ALERT: ieatnetsimv5116-41;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570520590] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570522016] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570522154] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570522291] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570524113] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570524715] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570524847] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570524984] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570525025] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570525180] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570525332] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570525384] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570525516] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570525645] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570529885] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570530127] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570530307] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570530681] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1570530727] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1570535588] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570535769] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570535949] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1570535952] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (20043) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570536063] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (608) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1570536142] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570536200] SERVICE ALERT: ieatlms3901;check pm collection statistical;OK;SOFT;2;OK- NUMBER OF PM FILES (21090) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570536323] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570536504] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570537098] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1570538165] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1570538309] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1570538524] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1570538680] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (10000) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1570538823] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1570538868] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1570539840] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.25%
[1570539990] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 54.63%
[1570546868] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (300) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1570546993] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (300) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1570547120] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (300) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1570547890] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (77) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570547904] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570548015] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (77) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570548031] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570548140] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (77) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570548157] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570548313] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (1905) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570548439] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (1905) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570548565] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1905) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570549113] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570549240] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (65) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570549367] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (65) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570549500] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1570549630] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (6811) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1570549673] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1570549760] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (6811) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1570549800] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1570549925] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1570552974] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (60) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570554148] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.16%
[1570554163] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570554299] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 71.09%
[1570554449] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.15%
[1570554600] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79.99%
[1570558404] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (5957) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570565215] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (20060) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570565419] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (20060) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570565627] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (20061) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570569349] SERVICE ALERT: ieatlms3901;check pm collection statistical;OK;HARD;3;OK- NUMBER OF PM FILES (21096) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (21948)
[1570573459] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2111) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)