[1562369579] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.67%
[1562369734] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1562369888] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.67%
[1562369934] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.8%
[1562370085] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.62%
[1562370236] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.59%
[1562370386] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.9%
[1562370537] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.88%
[1562371915] SERVICE FLAPPING ALERT: ieatnetsimv7035-22;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562373518] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.25%
[1562376820] SERVICE FLAPPING ALERT: ieatnetsimv7035-19;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1562379434] SERVICE FLAPPING ALERT: ieatnetsimv5116-16;check netsim disk;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562379623] SERVICE FLAPPING ALERT: ieatnetsimv5116-18;check netsim disk;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1562388645] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.64%
[1562388797] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.52%
[1562388947] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.34%
[1562389098] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.41%
[1562389251] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.71%
[1562389401] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.65%
[1562395910] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1562396091] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1562396272] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1562396395] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562396578] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562396751] SERVICE FLAPPING ALERT: ieatnetsimv7035-29;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1562396759] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1562402968] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.08%
[1562403119] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.58%
[1562406870] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.52%
[1562407021] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.22%
[1562407171] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.29%
[1562409331] HOST ALERT: ieatnetsimv7035-80;DOWN;SOFT;1;PING NOK
[1562409393] HOST ALERT: ieatnetsimv7035-80;UP;SOFT;1;PING OK
[1562410802] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;HARD;3;OK- CPU IS USED 82.53%
[1562411153] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.72%
[1562411304] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.07%
[1562411457] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.72%
[1562411569] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1562411932] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1562412291] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1562415088] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;OK;HARD;3;OK- CPU IS USED 82.11%
[1562451119] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.31%
[1562451270] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.82%
[1562451424] SERVICE ALERT: ieatnetsimv7035-29;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!