[1571267374] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1571267553] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1571267732] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1571267784] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571267826] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571268169] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13933) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571268416] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571271457] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1571271959] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2710) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571272191] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571272658] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2590) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571278612] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571278642] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2334) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571278739] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571278774] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2358) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571278865] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571278901] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2358) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571278996] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12635) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571279142] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12635) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571279222] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571279306] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12635) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571279367] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571279503] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571279605] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (3846) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571283217] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (3481) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571283346] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (3512) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571283477] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (3512) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571284049] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1571284049] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571285370] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4000) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1571285504] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4237) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1571285576] SERVICE ALERT: ieatnetsimv5116-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.24%
[1571285727] SERVICE ALERT: ieatnetsimv5116-33;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 70.35%
[1571285811] SERVICE ALERT: ieatnetsimv7035-30;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571285812] SERVICE ALERT: ieatnetsimv7035-31;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.32%
[1571285817] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.95%
[1571285961] SERVICE ALERT: ieatnetsimv7035-30;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 59%
[1571285963] SERVICE ALERT: ieatnetsimv7035-31;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.29%
[1571285967] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 44.82%
[1571286071] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (148) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571286109] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2513) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571286532] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13689) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571287086] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571288622] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571288747] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571288872] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571292052] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571292191] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (1210) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571292318] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1830) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571292476] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571292568] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (108) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571292695] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (108) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571292819] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (62) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571292822] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (108) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571292880] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.74%
[1571292943] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571293030] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 10.35%
[1571293658] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1571293818] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13056) IS FROM 80% TO 95% OF ENM COMECIM NODES (14059)
[1571294023] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13056) IS FROM 80% TO 95% OF ENM COMECIM NODES (14059)
[1571294185] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13056) IS FROM 80% TO 95% OF ENM COMECIM NODES (14059)
[1571294414] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571294600] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571295924] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2556) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571296426] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571296668] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1571296877] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571296917] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2325) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571297003] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571297046] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2325) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571297129] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571297174] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2325) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571297374] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12518) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571297533] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12518) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571297695] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12518) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571297838] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (14055) IS HIGHER THAN 95% OF ENM COMECIM NODES (14059)
[1571299241] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1571299422] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1571299604] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1571299842] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1571300023] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1571300205] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1571301465] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1571302066] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1571302767] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13757) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14018)
[1571303284] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571303409] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571303535] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571303617] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571303707] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571303770] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571303893] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571303920] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;OK;HARD;3;OK- CPU IS USED 66.06%
[1571304078] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571304922] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13933) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571305943] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571307140] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571307935] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571307983] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571308565] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1571308746] SERVICE ALERT: ieatlms5735-1;check node fs;OK;SOFT;2;OK
[1571309528] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571310253] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571313459] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571313659] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571313858] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571315735] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571315748] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12770) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571315866] SERVICE ALERT: ieatnetsimv7035-63;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571315890] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571315995] SERVICE ALERT: ieatnetsimv7035-63;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571316038] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12770) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571316075] SERVICE ALERT: ieatnetsimv7035-50;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571316082] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571316202] SERVICE ALERT: ieatnetsimv7035-50;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571316210] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571316505] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (42.3% change >= 20.0% threshold)
[1571316505] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (42.3% change >= 20.0% threshold)
[1571316505] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1571316505] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RBS;STARTED; Service appears to have started flapping (35.3% change >= 20.0% threshold)
[1571316505] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (48.4% change >= 20.0% threshold)
[1571316505] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1571316804] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (42.3% change >= 20.0% threshold)
[1571316804] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (42.3% change >= 20.0% threshold)
[1571316804] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1571316804] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RBS;STARTED; Service appears to have started flapping (35.3% change >= 20.0% threshold)
[1571316804] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (48.4% change >= 20.0% threshold)
[1571316804] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1571317243] SERVICE ALERT: ieatlms3901;diff_rpm_iso_installed RPM DIFF;WARNING;SOFT;1;----------------------------
[1571317367] SERVICE ALERT: ieatlms3901;diff_rpm_iso_installed RPM DIFF;WARNING;SOFT;2;----------------------------
[1571317491] SERVICE ALERT: ieatlms3901;diff_rpm_iso_installed RPM DIFF;WARNING;HARD;3;----------------------------
[1571317636] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (42.3% change >= 20.0% threshold)
[1571317636] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (42.3% change >= 20.0% threshold)
[1571317636] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1571317636] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RBS;STARTED; Service appears to have started flapping (35.3% change >= 20.0% threshold)
[1571317636] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (48.4% change >= 20.0% threshold)
[1571317636] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1571317926] SERVICE ALERT: ieatlms3901;check rpm diff;WARNING;SOFT;1;----------------------------
[1571318050] SERVICE ALERT: ieatlms3901;check rpm diff;WARNING;SOFT;2;----------------------------
[1571318174] SERVICE ALERT: ieatlms3901;check rpm diff;WARNING;HARD;3;----------------------------
[1571318741] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571318791] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2344) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571318884] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571318919] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2344) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571319015] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571319048] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2325) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571321546] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571321691] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571321818] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571322143] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2589) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571322625] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (142) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571322675] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571322888] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (30) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571323264] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13246) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571325135] SERVICE ALERT: ieatnetsimv7035-57;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571325263] SERVICE ALERT: ieatnetsimv7035-57;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571326241] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571326371] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571326499] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571326890] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12530) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571327039] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12530) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571327065] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1571327065] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571327188] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12683) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571328411] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571328610] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571328807] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571330106] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (142) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571330266] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1571330387] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1571330508] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1571330722] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571330843] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571330964] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571331422] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571331465] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.00 seconds)
[1571331518] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1571331533] SERVICE ALERT: ieatlms3901;check san alert healthcheck;CRITICAL;SOFT;1;CRITICAL- san_alert_healthcheck FAILED
[1571331539] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1571331544] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571331640] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1571331655] SERVICE ALERT: ieatlms3901;check san alert healthcheck;CRITICAL;SOFT;2;CRITICAL- san_alert_healthcheck FAILED
[1571331660] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1571331665] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1571331761] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1571331776] SERVICE ALERT: ieatlms3901;check san alert healthcheck;CRITICAL;HARD;3;CRITICAL- san_alert_healthcheck FAILED
[1571331782] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1571332385] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1571332488] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1571332488] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571332738] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (37) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1571332744] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (458) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1571332888] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571332888] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (37) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1571333018] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (37) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1571333305] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1571333460] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1683) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571333518] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (66) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571333530] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571333645] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571333645] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1571333714] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571333772] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571333840] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571333878] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (6728) IS LOWER THAN 80% OF ENM COMECIM NODES (14059)
[1571333967] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571334039] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;2;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (9000) IS LOWER THAN 80% OF ENM COMECIM NODES (14059)
[1571334170] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1571334200] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (8627) IS LOWER THAN 80% OF ENM COMECIM NODES (14059)
[1571334757] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571334944] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571335132] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571335438] SERVICE ALERT: ieatlms3901;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1571335621] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1571336165] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571336306] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12167) IS FROM 80% TO 95% OF ENM COMECIM NODES (14059)
[1571336316] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1571336361] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1042) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1571336363] SERVICE ALERT: ieatlms3901;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571336495] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1571336675] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1571336740] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (100) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (100)
[1571336812] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1571336814] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1571336853] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1571337154] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (26964) IS HIGHER THAN 95% OF ENM NODES (26965)
[1571337193] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571337376] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571337379] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571337559] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1571337861] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (22697) IS BETWEEN 80% AND 95% OF ENM NODES (26965)
[1571337912] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23698) IS BETWEEN 80% AND 95% OF ENM NODES (26965)
[1571338003] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (119) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571338147] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13479) IS HIGHER THAN 95% OF ENM COMECIM NODES (14059)
[1571338808] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1571338808] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571340669] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2318) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571340677] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571340867] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571340989] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (78) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571341061] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571341119] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6559) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571341632] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13058) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571343022] SERVICE FLAPPING ALERT: ieatlms3901;check vcs cluster hc;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1571343022] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571343428] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (17) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571343556] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF PM FILES (24) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571344828] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (150) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571344885] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571345068] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571345394] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571345653] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571347051] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571347239] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571347422] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1571349887] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571351085] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571351479] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2438) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)