[1573430852] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1573452146] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573452146] HOST ALERT: ieatnetsimv5116-66;DOWN;SOFT;1;PING NOK
[1573452208] HOST ALERT: ieatnetsimv5116-66;UP;SOFT;1;PING OK
[1573452266] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1573452850] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573453035] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1573453220] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1573458686] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573459846] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.61%
[1573459996] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 50.17%
[1573461743] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12)
[1573461882] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12)
[1573462008] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12)
[1573462812] SERVICE ALERT: ieatnetsimv5116-15;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462814] SERVICE ALERT: ieatnetsimv5116-18;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462815] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462818] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462819] SERVICE ALERT: ieatnetsimv5116-27;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462825] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573462842] SERVICE ALERT: ieatnetsimv5116-29;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462845] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462847] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462895] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462898] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462901] SERVICE ALERT: ieatnetsimv5116-81;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462902] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462930] SERVICE ALERT: ieatnetsimv5116-91;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573462954] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573462964] SERVICE ALERT: ieatnetsimv5116-15;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573462965] SERVICE ALERT: ieatnetsimv5116-18;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573462967] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573462970] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573462970] SERVICE ALERT: ieatnetsimv5116-27;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573462993] SERVICE ALERT: ieatnetsimv5116-29;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573462996] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573462999] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463028] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (161) IS BETWEEN 5% AND 20% OF ENM NODES (2560)
[1573463048] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463051] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463053] SERVICE ALERT: ieatnetsimv5116-81;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463055] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463083] SERVICE ALERT: ieatnetsimv5116-91;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463100] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573463115] SERVICE ALERT: ieatnetsimv5116-15;check netsim cpu;OK;HARD;3;OK- CPU IS USED 15.16%
[1573463116] SERVICE ALERT: ieatnetsimv5116-18;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.18%
[1573463118] SERVICE ALERT: ieatnetsimv5116-21;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.99%
[1573463121] SERVICE ALERT: ieatnetsimv5116-27;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.38%
[1573463121] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.42%
[1573463144] SERVICE ALERT: ieatnetsimv5116-29;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.99%
[1573463147] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.06%
[1573463150] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;OK;HARD;3;OK- CPU IS USED 11.58%
[1573463186] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (160) IS BETWEEN 5% AND 20% OF ENM NODES (2560)
[1573463200] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463203] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463206] SERVICE ALERT: ieatnetsimv5116-81;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463207] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463225] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573463235] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573463235] SERVICE ALERT: ieatnetsimv5116-91;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463297] SERVICE ALERT: ieatnetsimv5116-88;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573463364] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (160) IS BETWEEN 5% AND 20% OF ENM NODES (2560)
[1573463379] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463391] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463412] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2400) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573463424] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573463449] SERVICE ALERT: ieatnetsimv5116-88;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463482] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573463489] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573463499] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573463531] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463546] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463602] SERVICE ALERT: ieatnetsimv5116-88;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463611] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573463634] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463642] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463651] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463663] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1573463786] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463794] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463803] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573463804] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573463901] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1573463955] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573463964] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573463972] SERVICE ALERT: ieatlms3901;check fm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573463985] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (2400) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573464089] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573464107] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573464116] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573464179] SERVICE ALERT: ieatlms3901;check fm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573464229] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573464241] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573464261] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (2400) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573464268] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573464314] SERVICE ALERT: ieatlms3901;check fm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1573464325] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573464343] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12)
[1573464347] SERVICE ALERT: ieatlms3901;check fm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1573464349] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1573464367] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573464381] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573464392] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573464472] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12)
[1573464516] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573464518] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573464533] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573464578] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573464582] SERVICE ALERT: ieatlms3901;check number nodes;OK;SOFT;2;OK- NUMBER OF NODES CONFIGURED IN ENM (24470) IS HIGHER THAN 95% OF NETSIM NODES (17154)
[1573464586] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573464598] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12)
[1573464667] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573464670] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573464678] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573464730] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573464738] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573464818] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573464830] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573464881] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573464890] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573464982] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573465128] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573465134] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573465279] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573465285] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573465430] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573465437] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573465952] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573466137] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1573466288] SERVICE ALERT: ieatnetsimv5116-01;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573466290] SERVICE ALERT: ieatnetsimv5116-04;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573466292] SERVICE ALERT: ieatnetsimv5116-03;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573466322] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1573466399] SERVICE ALERT: ieatnetsimv5116-02;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573466440] SERVICE ALERT: ieatnetsimv5116-01;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573466445] SERVICE ALERT: ieatnetsimv5116-03;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573466448] SERVICE ALERT: ieatnetsimv5116-04;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573466550] SERVICE ALERT: ieatnetsimv5116-02;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573466591] SERVICE ALERT: ieatnetsimv5116-01;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573466596] SERVICE ALERT: ieatnetsimv5116-03;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573466599] SERVICE ALERT: ieatnetsimv5116-04;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573466701] SERVICE ALERT: ieatnetsimv5116-02;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573466709] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2445) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573466752] SERVICE ALERT: ieatnetsimv5116-24;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.97%
[1573466830] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 31.55%
[1573466834] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 40.17%
[1573466837] SERVICE ALERT: ieatnetsimv5116-81;check netsim cpu;OK;HARD;3;OK- CPU IS USED 63.91%
[1573466837] SERVICE ALERT: ieatnetsimv5116-84;check netsim cpu;OK;HARD;3;OK- CPU IS USED 55.09%
[1573466875] SERVICE ALERT: ieatnetsimv5116-91;check netsim cpu;OK;HARD;3;OK- CPU IS USED 46.2%
[1573466970] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN ERBS NODES !
[1573467161] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.85%
[1573467177] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;HARD;3;OK- CPU IS USED 11.67%
[1573467232] SERVICE ALERT: ieatnetsimv5116-88;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.1%
[1573467868] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2550) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1573468248] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.08%
[1573468253] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1573468380] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1573468398] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.58%
[1573468509] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1573470144] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573470189] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573470222] SERVICE ALERT: ieatnetsimv5116-01;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.48%
[1573470284] SERVICE ALERT: ieatlms3901;check fm supervision nodes MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (150) IS LOWER THAN 80% OF ENM NODES (300)
[1573470316] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1573470412] SERVICE ALERT: ieatlms3901;check fm supervision nodes MGW;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1573473144] SERVICE ALERT: ieatlms3901;check pm supervision nodes DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (30) IS LOWER THAN 80% OF ENM NODES (60)
[1573473250] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573473268] SERVICE ALERT: ieatlms3901;check pm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1573473552] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1573473745] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573473827] SERVICE ALERT: ieatnetsimv5116-03;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.12%
[1573474184] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573474366] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573474487] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573474547] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573474668] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573474757] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (746) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573474778] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (24)
[1573474849] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573474972] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573475405] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (72) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (12)
[1573475447] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13484) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573476414] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-IS;OK;HARD;3;OK- NUMBER OF PM FILES (72) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (12)
[1573476913] SERVICE ALERT: ieatlms3901;check pm collection statistical CISCO-ASR900;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573477239] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;OK;HARD;3;OK- NUMBER OF PM FILES (50) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (50)
[1573477429] SERVICE ALERT: ieatnetsimv5116-04;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.33%
[1573477670] SERVICE ALERT: ieatnetsimv5116-81;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.31%
[1573477821] SERVICE ALERT: ieatnetsimv5116-81;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 75.81%
[1573477828] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;OK;HARD;3;OK- NUMBER OF PM FILES (1200) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (1200)
[1573477974] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;OK;HARD;3;OK- NUMBER OF PM FILES (1000) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (1000)
[1573478217] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;OK;HARD;3;OK- CPU IS USED 15.84%
[1573478225] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 80.34%
[1573478698] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;OK;HARD;3;OK- CPU IS USED 14.72%
[1573478823] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 22.85%
[1573479248] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.57%
[1573480972] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1573480986] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573481118] SERVICE ALERT: ieatnetsimv5116-02;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7%
[1573481124] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1573481252] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1573481380] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1573482094] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.57%
[1573485564] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (690) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573485698] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF PM FILES (745) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573485727] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573485880] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573486032] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573486512] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.11%
[1573486521] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 11.11%
[1573486613] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.96%
[1573487061] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.06%
[1573489033] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.83%
[1573492069] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.18%
[1573492220] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 25.78%
[1573493364] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 17.24%
[1573493500] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 18.64%
[1573495434] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573495447] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573495615] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573495628] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573495798] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573495811] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573496853] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.97%
[1573497004] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.9%
[1573497155] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.83%
[1573499253] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1573499373] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1573499494] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1573500785] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 77.46%
[1573504661] SERVICE ALERT: ieatlms3901;check unsync nodes;OK;HARD;3;OK
[1573504992] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0)
[1573505161] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0)
[1573505422] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (48) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1573505913] SERVICE ALERT: ieatlms3901;check pm collection statistical BSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0)
[1573507087] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0)
[1573507171] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0)
[1573507248] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573507434] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573507620] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1573507662] SERVICE FLAPPING ALERT: ieatnetsimv5116-75;check netsim cpu;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1573507662] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.4%
[1573508309] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573508460] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1573508490] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573508671] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573509345] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1573509596] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1573509848] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1573510321] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573510502] SERVICE ALERT: ieatlms3901;check node fs;OK;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1573510534] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1573510684] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573510865] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573510937] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;OK;HARD;3;OK- CPU IS USED 56.37%
[1573511046] SERVICE FLAPPING ALERT: ieatlms3901;check node fs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1573511046] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573512395] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573512576] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573512756] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573512907] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1573513581] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1573513932] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1573514183] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1573514572] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.94%
[1573514723] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 9.38%