[1576715827] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1576716393] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized EPG;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576717216] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (106) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576717342] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (83) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576717467] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (83) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576719371] SERVICE FLAPPING ALERT: ieatlms7435;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1576720228] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized ERBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576720247] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MTAS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576721078] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1576724350] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576724479] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576724600] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576726322] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1486) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576726478] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1465) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576726603] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1459) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576728301] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1576728505] SERVICE ALERT: ieatlms7435;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1576728792] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1576730509] SERVICE FLAPPING ALERT: ieatlms7435;check system service hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576733068] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576733253] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576733439] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576734931] SERVICE FLAPPING ALERT: ieatlms7435;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576736534] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576737107] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576737430] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576737550] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576737671] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576738345] SERVICE FLAPPING ALERT: ieatlms7435;check hw resources hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1576738682] SERVICE ALERT: ieatnetsimv9041-05;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576738904] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576739457] SERVICE ALERT: ieatnetsimv9041-05;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1576739692] SERVICE ALERT: ieatnetsimv9041-05;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1576740038] SERVICE ALERT: ieatnetsimv7035-89;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576740189] SERVICE ALERT: ieatnetsimv7035-89;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1576740300] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;OK;HARD;3;OK- CPU IS USED 74.79%
[1576740335] SERVICE ALERT: ieatnetsimv7035-89;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1576740777] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576740880] SERVICE ALERT: ieatnetsimv9041-06;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1576740890] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1576740963] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576741011] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1576741148] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576742095] SERVICE ALERT: ieatnetsimv9041-06;check netsim cpu;OK;HARD;3;OK- CPU IS USED 1.09%
[1576742379] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-6352;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MINI-LINK-6352 NODES !
[1576742601] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576742721] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576742842] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576744385] SERVICE ALERT: ieatlms7435;check fm supervision nodes SBG-IS;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1576744979] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (20060) IS HIGHER THAN 95% OF ENM NODES (20060)
[1576745962] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576745962] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (26.7% change >= 20.0% threshold)
[1576745962] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1576745984] SERVICE ALERT: ieatlms3901;copy nodes workload - DEMO;CRITICAL;SOFT;1;CRITICAL -  scp ok - nodes exists
[1576746104] SERVICE ALERT: ieatlms3901;copy nodes workload - DEMO;CRITICAL;SOFT;2;CRITICAL -  scp ok - nodes exists
[1576746225] SERVICE ALERT: ieatlms3901;copy nodes workload - DEMO;CRITICAL;HARD;3;CRITICAL -  scp ok - nodes exists
[1576746359] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576746359] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (26.7% change >= 20.0% threshold)
[1576746360] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1576746442] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576746946] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576746946] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576746946] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (29.6% change >= 20.0% threshold)
[1576746955] SERVICE ALERT: ieatlms3901;copy nodes workload - DEMO;WARNING;HARD;3;WARNING -  scp ok - nodes exists
[1576747023] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576747023] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576747024] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (29.6% change >= 20.0% threshold)
[1576747030] SERVICE ALERT: ieatlms3901;copy nodes workload - DEMO;OK;HARD;3;OK -  scp ok - nodes exists
[1576747230] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576747230] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576747230] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (29.6% change >= 20.0% threshold)
[1576747266] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576747266] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576747266] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (29.6% change >= 20.0% threshold)
[1576747341] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576747341] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576747341] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (29.6% change >= 20.0% threshold)
[1576747913] SERVICE ALERT: ieatnetsimv7035-31;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.49%
[1576748063] SERVICE ALERT: ieatnetsimv7035-31;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 69.49%
[1576748414] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1576748414] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576748878] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1576748878] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (22.3% change >= 20.0% threshold)
[1576748878] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1576748878] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (29.6% change >= 20.0% threshold)
[1576751973] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1576751973] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1576751973] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1576752350] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1576752350] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576752350] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1576752408] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1576752408] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576752409] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1576752473] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1576752473] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576752473] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1576753084] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1576753084] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1576753084] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1576753483] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (315) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576753613] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (327) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576753760] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (17) IS HIGHER THAN 20% OF ENM NODES (50)
[1576753765] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1576753765] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (336) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576753885] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (18) IS HIGHER THAN 20% OF ENM NODES (50)
[1576753897] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576754009] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (18) IS HIGHER THAN 20% OF ENM NODES (50)
[1576754023] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576754149] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576754599] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576754724] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576754851] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576755319] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (17) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576755443] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (16) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576755566] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.1% change >= 20.0% threshold)
[1576755566] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (22) IS HIGHER THAN 20% OF ENM NODES (100)
[1576755771] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (70) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576755897] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (69) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576755964] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.1% change >= 20.0% threshold)
[1576755964] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.0% change >= 20.0% threshold)
[1576755964] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1576755964] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1576755964] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1576755994] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576756002] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.1% change >= 20.0% threshold)
[1576756002] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.0% change >= 20.0% threshold)
[1576756002] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1576756002] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1576756002] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1576756022] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (69) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576756112] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.1% change >= 20.0% threshold)
[1576756112] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.0% change >= 20.0% threshold)
[1576756112] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1576756112] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1576756112] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1576756119] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576756147] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.1% change >= 20.0% threshold)
[1576756147] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.0% change >= 20.0% threshold)
[1576756147] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1576756147] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1576756147] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1576756231] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.1% change >= 20.0% threshold)
[1576756231] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.0% change >= 20.0% threshold)
[1576756231] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1576756231] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1576756231] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1576756264] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576757235] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576757256] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (743) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576757284] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.1% change >= 20.0% threshold)
[1576757284] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (21.0% change >= 20.0% threshold)
[1576757284] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1576757284] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1576757284] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1576757391] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (755) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576757430] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576757521] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (474) IS LOWER THAN 5% OF ENM NODES (14035)
[1576757624] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576757760] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1576757760] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576757827] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (79) IS HIGHER THAN 95% OF ENM  NODES (80)
[1576758069] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1576758173] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.1% change >= 20.0% threshold)
[1576758173] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1576758173] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1576758173] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1576758173] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1576758454] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1576758454] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN DSC NODES !
[1576758499] SERVICE ALERT: ieatlms7435;check rpm diff;CRITICAL;HARD;3;ATTENTION!!! ISO version is not present
[1576758546] SERVICE ALERT: ieatlms3901;check rpm diff;CRITICAL;HARD;3;ATTENTION!!! ISO version is not present
[1576758656] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1576758656] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.1% change >= 20.0% threshold)
[1576758656] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1576758656] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1576758656] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1576758656] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1576758817] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1576758817] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.1% change >= 20.0% threshold)
[1576758817] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1576758817] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1576758817] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1576758817] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1576758889] SERVICE ALERT: ieatlms5735-1;check rpm diff;CRITICAL;HARD;3;ATTENTION!!! ISO version is not present !!!
[1576759171] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576759279] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576759438] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1576759438] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (28.7% change >= 20.0% threshold)
[1576759438] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1576759438] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1576759438] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1576759619] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576759644] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1576759644] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (28.7% change >= 20.0% threshold)
[1576759644] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1576759644] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1576759644] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1576759723] SERVICE ALERT: ieatlms3901;check list nodes;CRITICAL;SOFT;1;CRITICAL - some  type nodes in management server (141.137.208.23) are different (28)
[1576759845] SERVICE ALERT: ieatlms3901;check list nodes;CRITICAL;SOFT;2;CRITICAL - some  type nodes in management server (141.137.208.23) are different (28)
[1576759966] SERVICE ALERT: ieatlms3901;check list nodes;CRITICAL;HARD;3;CRITICAL - some  type nodes in management server (141.137.208.23) are different (28)
[1576760285] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1576760285] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (28.7% change >= 20.0% threshold)
[1576760285] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1576760285] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1576760285] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1576760865] SERVICE ALERT: ieatlms7435;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576760968] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1576761213] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1576762843] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (27.9% change >= 20.0% threshold)
[1576762843] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (26.1% change >= 20.0% threshold)
[1576762843] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1576764424] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576764443] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576764544] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576764563] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576764665] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576764684] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1576764684] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576765305] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576765492] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576765658] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576765795] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576766030] HOST ALERT: ieatnetsimv5116-25;DOWN;SOFT;1;PING NOK
[1576766092] HOST ALERT: ieatnetsimv5116-25;UP;SOFT;1;PING OK
[1576766525] SERVICE ALERT: ieatwlvm7102;check workload profiles;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1576766647] SERVICE ALERT: ieatwlvm7102;check workload profiles;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1576766770] SERVICE ALERT: ieatwlvm7102;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1576766894] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1576767015] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1576767136] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1576767185] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1576767201] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1576767306] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1576767322] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1576767414] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576767427] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1576767443] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1576767535] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576767656] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576767819] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1576767843] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1576767940] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1576767964] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1576768023] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1576768061] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1576768086] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1576768145] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1576768266] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1576768535] SERVICE ALERT: ieatlms3901;check san alert healthcheck;CRITICAL;SOFT;1;CRITICAL- san_alert_healthcheck FAILED
[1576768657] SERVICE ALERT: ieatlms3901;check san alert healthcheck;CRITICAL;SOFT;2;CRITICAL- san_alert_healthcheck FAILED
[1576768713] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1576768723] HOST ALERT: ieatnetsimv5116-108;DOWN;SOFT;1;PING NOK
[1576768728] HOST ALERT: ieatnetsimv5116-46;DOWN;SOFT;1;PING NOK
[1576768778] SERVICE ALERT: ieatlms3901;check san alert healthcheck;CRITICAL;HARD;3;CRITICAL- san_alert_healthcheck FAILED
[1576768785] HOST ALERT: ieatnetsimv5116-108;UP;SOFT;1;PING OK
[1576768790] HOST ALERT: ieatnetsimv5116-46;UP;SOFT;1;PING OK
[1576768835] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1576768956] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1576769164] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1576769285] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1576769442] SERVICE ALERT: ieatlms3901;check cm supervision nodes BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (96)
[1576769537] SERVICE ALERT: ieatnetsimv7102-81;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576769541] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576769568] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576769585] SERVICE ALERT: ieatlms3901;check cm supervision nodes BSC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (96) IS HIGHER THAN 95% OF ENM NODES (96)
[1576769589] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1576769687] SERVICE ALERT: ieatnetsimv7102-81;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.28%
[1576769692] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 4.89%
[1576769925] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1576769964] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1576770358] SERVICE ALERT: ieatnetsimv7102-33;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576770509] SERVICE ALERT: ieatnetsimv7102-33;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.6%
[1576770938] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1576771089] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1576771104] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1576771754] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1576771814] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 132 processes with STATE = RSZDT
[1576771866] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;OK;HARD;3;OK
[1576771885] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576771924] SERVICE ALERT: ieatlms3901;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1576772440] SERVICE ALERT: ieatlms3901;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1576772626] SERVICE ALERT: ieatlms3901;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1576772922] HOST ALERT: ieatnetsimv5116-116;DOWN;SOFT;1;PING NOK
[1576772952] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1576772984] HOST ALERT: ieatnetsimv5116-116;UP;SOFT;1;PING OK
[1576773134] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1576773236] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576773248] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1576773317] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1576773424] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576773427] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1576773606] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1576773611] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576777279] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1576777279] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576778505] SERVICE ALERT: ieatnetsimv7102-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576778655] SERVICE ALERT: ieatnetsimv7102-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.16%
[1576779149] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576779338] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576779527] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576780404] SERVICE ALERT: ieatlms3901;check synchronized RBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (3205) IS FROM 80% TO 95% OF ENM  NODES (3849)
[1576780534] SERVICE ALERT: ieatlms3901;check synchronized RBS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (3849)!
[1576782380] SERVICE FLAPPING ALERT: ieatlms3901;check vcs cluster hc;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1576782380] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1576782666] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576782787] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576782907] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576784784] SERVICE ALERT: ieatlms3901;check pm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1576784908] SERVICE ALERT: ieatlms3901;check pm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1576785590] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1576785716] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1576785824] HOST ALERT: ieatnetsimv5116-13;DOWN;SOFT;1;PING NOK
[1576785886] HOST ALERT: ieatnetsimv5116-13;UP;SOFT;1;PING OK
[1576786043] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1576786071] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (587) IS LOWER THAN 80% OF ENM  NODES (767)
[1576786225] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1576786286] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576786406] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576786526] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576788009] SERVICE ALERT: ieatnetsimv7035-19;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576788129] SERVICE ALERT: ieatnetsimv7035-19;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576788207] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1576788250] SERVICE ALERT: ieatnetsimv7035-19;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576788390] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1576788521] HOST ALERT: ieatnetsimv5116-114;DOWN;SOFT;1;PING NOK
[1576788574] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1576788583] HOST ALERT: ieatnetsimv5116-114;UP;SOFT;1;PING OK
[1576790181] SERVICE ALERT: ieatnetsimv7102-77;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.46%
[1576790331] SERVICE ALERT: ieatnetsimv7102-77;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.44%
[1576791221] HOST ALERT: ieatnetsimv5116-56;DOWN;SOFT;1;PING NOK
[1576791221] HOST ALERT: ieatnetsimv5116-57;DOWN;SOFT;1;PING NOK
[1576791239] HOST ALERT: ieatnetsimv5116-56;UP;SOFT;1;PING OK
[1576791283] HOST ALERT: ieatnetsimv5116-57;UP;SOFT;1;PING OK
[1576791850] SERVICE ALERT: ieatnetsimv7035-19;check netsim disk;OK;HARD;3;OK
[1576793743] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576794037] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1576794548] SERVICE ALERT: ieatlms3901;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (41) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576794679] SERVICE ALERT: ieatlms3901;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576794804] SERVICE ALERT: ieatlms3901;check synchronized Router6274;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576796608] SERVICE ALERT: ieatlms3901;check synchronized Router6274;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1576797700] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1576797883] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1576799329] HOST ALERT: ieatnetsimv5116-73;DOWN;SOFT;1;PING NOK
[1576799329] HOST ALERT: ieatnetsimv5116-74;DOWN;SOFT;1;PING NOK
[1576799391] HOST ALERT: ieatnetsimv5116-74;UP;SOFT;1;PING OK
[1576799391] HOST ALERT: ieatnetsimv5116-73;UP;SOFT;1;PING OK
[1576799865] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED