[1574035326] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (40) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1574035452] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (40) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1574036270] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574036419] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574036545] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574037001] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574037131] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF PM FILES (24) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (24)
[1574037910] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (609) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574038009] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;OK;HARD;3;OK
[1574038195] SERVICE ALERT: ieatlms3901;check pm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574038279] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574038321] SERVICE ALERT: ieatlms3901;check pm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574038408] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2558) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1574038442] SERVICE ALERT: ieatlms3901;check pm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (2560)
[1574038493] SERVICE ALERT: ieatlms3901;check cm supervision nodes DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (14) IS LOWER THAN 80% OF ENM NODES (30)
[1574038572] SERVICE ALERT: ieatlms3901;check pm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574038618] SERVICE ALERT: ieatlms3901;check cm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574038753] SERVICE ALERT: ieatlms3901;check pm supervision nodes ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1000)
[1574038879] SERVICE ALERT: ieatlms3901;check pm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1574039059] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (39) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1574039084] SERVICE ALERT: ieatlms3901;check pm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1574039209] SERVICE ALERT: ieatlms3901;check pm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1574039575] SERVICE ALERT: ieatlms3901;check cm supervision nodes RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (3849)
[1574039664] SERVICE ALERT: ieatlms3901;check cm supervision nodes DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (2) IS LOWER THAN 80% OF ENM NODES (30)
[1574039705] SERVICE ALERT: ieatlms3901;check cm supervision nodes RBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (3849) IS HIGHER THAN 95% OF ENM NODES (3849)
[1574039750] SERVICE ALERT: ieatlms3901;check pm supervision nodes MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574039788] SERVICE ALERT: ieatlms3901;check cm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574039875] SERVICE ALERT: ieatlms3901;check pm supervision nodes MGW;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574040151] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;OK;HARD;3;OK- NUMBER OF PM FILES (40) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574040628] SERVICE ALERT: ieatlms3901;check pm supervision nodes MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (6) IS LOWER THAN 80% OF ENM NODES (12)
[1574040744] SERVICE ALERT: ieatlms3901;check cm supervision nodes MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574040755] SERVICE ALERT: ieatlms3901;check pm supervision nodes MSC-BC-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1574040787] SERVICE ALERT: ieatlms3901;check pm supervision nodes MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1574040870] SERVICE ALERT: ieatlms3901;check cm supervision nodes MGW;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574040914] SERVICE ALERT: ieatlms3901;check pm supervision nodes MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574040974] SERVICE ALERT: ieatlms3901;check fm supervision nodes Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (26) IS LOWER THAN 80% OF ENM NODES (50)
[1574041099] SERVICE ALERT: ieatlms3901;check fm supervision nodes Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1574041517] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (554) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574041724] SERVICE ALERT: ieatlms3901;check fm supervision nodes SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (80)
[1574041850] SERVICE ALERT: ieatlms3901;check fm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1574042531] SERVICE ALERT: ieatlms3901;check fm supervision nodes MSC-BC-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (12)
[1574042559] SERVICE ALERT: ieatlms3901;check pm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (24)
[1574042655] SERVICE ALERT: ieatlms3901;check fm supervision nodes MSC-BC-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1574042665] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (40) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1574042683] SERVICE ALERT: ieatlms3901;check pm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1574042706] SERVICE ALERT: ieatlms3901;check pm supervision nodes MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574042804] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (40) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1574042831] SERVICE ALERT: ieatlms3901;check pm supervision nodes MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574042931] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1574042931] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (40) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1574042962] SERVICE ALERT: ieatlms3901;check pm supervision nodes MGW;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574043758] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574043895] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574043907] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1574044022] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574044033] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1574047038] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1601) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574047167] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1581) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574047292] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1555) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574047629] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;OK;HARD;3;OK- NUMBER OF PM FILES (40) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574048292] SERVICE ALERT: ieatlms3901;check fm supervision nodes SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (80)
[1574048353] SERVICE ALERT: ieatlms3901;check pm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574048418] SERVICE ALERT: ieatlms3901;check fm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (0)
[1574048480] SERVICE ALERT: ieatlms3901;check pm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574048605] SERVICE ALERT: ieatlms3901;check pm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (2560)
[1574048736] SERVICE ALERT: ieatlms3901;check pm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574049495] SERVICE ALERT: ieatlms3901;check pm supervision nodes RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (3849)
[1574049535] SERVICE ALERT: ieatlms3901;check cm supervision nodes EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (50)
[1574049625] SERVICE ALERT: ieatlms3901;check pm supervision nodes RBS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (3849) IS HIGHER THAN 95% OF ENM NODES (3849)
[1574049660] SERVICE ALERT: ieatlms3901;check cm supervision nodes EPG;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1574050395] SERVICE ALERT: ieatlms3901;check pm supervision nodes EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (50)
[1574050448] SERVICE ALERT: ieatlms3901;check pm supervision nodes RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (14003)
[1574050521] SERVICE ALERT: ieatlms3901;check pm supervision nodes EPG;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1574050592] SERVICE ALERT: ieatlms3901;check pm supervision nodes RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (13974) IS HIGHER THAN 95% OF ENM NODES (14002)
[1574050928] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1662) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574051237] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574051364] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574051490] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (36) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574051638] SERVICE ALERT: ieatlms3901;check fm supervision nodes BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (96)
[1574051701] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1750)
[1574051763] SERVICE ALERT: ieatlms3901;check fm supervision nodes BSC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (96) IS HIGHER THAN 95% OF ENM NODES (96)
[1574051783] SERVICE ALERT: ieatlms3901;check fm supervision nodes MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (10) IS BETWEEN 80% AND 95% OF ENM NODES (12)
[1574051829] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1750) IS HIGHER THAN 95% OF ENM NODES (0)
[1574051908] SERVICE ALERT: ieatlms3901;check cm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1574051909] SERVICE ALERT: ieatlms3901;check fm supervision nodes MSC-BC-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1574052032] SERVICE ALERT: ieatlms3901;check cm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1574052878] SERVICE ALERT: ieatlms3901;check cm supervision nodes RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (14029)
[1574053020] SERVICE ALERT: ieatlms3901;check cm supervision nodes RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (13973) IS HIGHER THAN 95% OF ENM NODES (14029)
[1574053473] SERVICE ALERT: ieatlms3901;check fm supervision nodes Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1200)
[1574053603] SERVICE ALERT: ieatlms3901;check fm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1574053742] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (38) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1574053801] SERVICE ALERT: ieatlms3901;check cm supervision nodes Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1200)
[1574053936] SERVICE ALERT: ieatlms3901;check cm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1574054458] SERVICE ALERT: ieatlms3901;check fm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574054584] SERVICE ALERT: ieatlms3901;check fm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574054653] SERVICE ALERT: ieatlms3901;check pm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1574054659] SERVICE ALERT: ieatlms3901;check fm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (10) IS LOWER THAN 80% OF ENM NODES (24)
[1574054729] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574054765] SERVICE ALERT: ieatlms3901;check pm supervision nodes DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1574054778] SERVICE ALERT: ieatlms3901;check pm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1574054783] SERVICE ALERT: ieatlms3901;check fm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1574054857] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1746) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574054890] SERVICE ALERT: ieatlms3901;check pm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574054964] SERVICE ALERT: ieatlms3901;check fm supervision nodes RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (3) IS LOWER THAN 80% OF ENM NODES (6)
[1574055092] SERVICE ALERT: ieatlms3901;check fm supervision nodes RNC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (6) IS HIGHER THAN 95% OF ENM NODES (6)
[1574055134] SERVICE ALERT: ieatlms3901;check cm supervision nodes RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (3849)
[1574055244] SERVICE ALERT: ieatlms3901;check fm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (2560)
[1574055264] SERVICE ALERT: ieatlms3901;check cm supervision nodes RBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (3849) IS HIGHER THAN 95% OF ENM NODES (3849)
[1574055374] SERVICE ALERT: ieatlms3901;check fm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574055932] SERVICE ALERT: ieatlms3901;check synchronized RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (3849)
[1574056069] SERVICE ALERT: ieatlms3901;check synchronized RBS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (3849)!
[1574056154] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (12)
[1574056281] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-BC-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1574057011] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-BC-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (5) IS LOWER THAN 80% OF ENM NODES (12)
[1574057068] SERVICE ALERT: ieatlms3901;check pm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1574057092] SERVICE ALERT: ieatlms3901;check pm supervision nodes RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (6)
[1574057135] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-BC-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1574057203] SERVICE ALERT: ieatlms3901;check pm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1574057220] SERVICE ALERT: ieatlms3901;check pm supervision nodes RNC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (6) IS HIGHER THAN 95% OF ENM NODES (6)
[1574057359] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (40) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1574057486] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (40) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1574057613] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (40) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100)
[1574057773] SERVICE ALERT: ieatlms3901;check cm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574057900] SERVICE ALERT: ieatlms3901;check cm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574058107] SERVICE ALERT: ieatlms3901;check synchronized MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (150)
[1574058181] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574058258] SERVICE ALERT: ieatlms3901;check synchronized MGW;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1574058361] SERVICE ALERT: ieatlms5735-1;check node fs;OK;SOFT;2;OK
[1574058374] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1750)
[1574058512] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1750) IS HIGHER THAN 95% OF ENM NODES (1750)
[1574058542] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574058617] SERVICE ALERT: ieatlms3901;check pm supervision nodes DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (22) IS LOWER THAN 80% OF ENM NODES (30)
[1574058723] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574058745] SERVICE ALERT: ieatlms3901;check pm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574058903] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1574059445] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574059679] SERVICE ALERT: ieatlms3901;check pm supervision nodes Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (29) IS LOWER THAN 80% OF ENM NODES (50)
[1574059687] SERVICE ALERT: ieatlms3901;check cm supervision nodes EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (35) IS LOWER THAN 80% OF ENM NODES (50)
[1574059805] SERVICE ALERT: ieatlms3901;check pm supervision nodes Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1574059813] SERVICE ALERT: ieatlms3901;check cm supervision nodes EPG;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1574060403] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (240) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574060529] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (316) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574060656] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (459) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574061254] SERVICE ALERT: ieatwlvm5116;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1574061526] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574061700] SERVICE ALERT: ieatlms3901;check cm supervision nodes MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574061825] SERVICE ALERT: ieatlms3901;check cm supervision nodes MGW;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574062706] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (12)
[1574062768] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1750)
[1574062830] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-BC-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1574062897] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1750) IS HIGHER THAN 95% OF ENM NODES (1750)
[1574063070] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1574063209] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574063306] SERVICE ALERT: ieatlms3901;check pm supervision nodes ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1000)
[1574063330] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;OK;SOFT;2;OK
[1574063428] SERVICE ALERT: ieatlms3901;check fm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1574063434] SERVICE ALERT: ieatlms3901;check pm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1574063554] SERVICE ALERT: ieatlms3901;check fm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1574063679] SERVICE ALERT: ieatlms3901;check fm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1574063804] SERVICE ALERT: ieatlms3901;check fm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1574063837] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (24)
[1574063978] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (0)
[1574063986] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (4) IS LOWER THAN 80% OF ENM NODES (20)
[1574063988] SERVICE ALERT: ieatlms3901;check cm supervision nodes RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (14017)
[1574064050] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (750)
[1574064097] SERVICE ALERT: ieatlms3901;check synchronized Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (1200)
[1574064110] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1574064141] SERVICE ALERT: ieatlms3901;check cm supervision nodes RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (14016) IS HIGHER THAN 95% OF ENM NODES (14017)
[1574064180] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1574064246] SERVICE ALERT: ieatlms3901;check synchronized Router6672;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1200)!
[1574064263] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574064287] SERVICE ALERT: ieatlms3901;check pm supervision nodes DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1574064307] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (750)
[1574064414] SERVICE ALERT: ieatlms3901;check pm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574064436] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1574064572] SERVICE ALERT: ieatlms3901;check pm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1574064696] SERVICE ALERT: ieatlms3901;check pm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1574065233] SERVICE ALERT: ieatlms3901;check cm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574065359] SERVICE ALERT: ieatlms3901;check cm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574065550] SERVICE ALERT: ieatlms3901;check cm supervision nodes MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574065676] SERVICE ALERT: ieatlms3901;check cm supervision nodes MGW;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574065831] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13695) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1574065896] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0)
[1574065942] SERVICE ALERT: ieatlms3901;check pm supervision nodes BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (96)
[1574066086] SERVICE ALERT: ieatlms3901;check pm supervision nodes BSC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (96) IS HIGHER THAN 95% OF ENM NODES (96)
[1574066198] SERVICE ALERT: ieatlms3901;check fm supervision nodes MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574066274] SERVICE ALERT: ieatlms3901;check fm supervision nodes RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (3849)
[1574066300] SERVICE ALERT: ieatlms3901;check fm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (2560)
[1574066323] SERVICE ALERT: ieatlms3901;check fm supervision nodes MGW;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574066357] SERVICE ALERT: ieatlms3901;check pm supervision nodes RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (6)
[1574066404] SERVICE ALERT: ieatlms3901;check fm supervision nodes RBS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (3849) IS HIGHER THAN 95% OF ENM NODES (3849)
[1574066430] SERVICE ALERT: ieatlms3901;check fm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574066486] SERVICE ALERT: ieatlms3901;check pm supervision nodes RNC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (6) IS HIGHER THAN 95% OF ENM NODES (0)
[1574066949] SERVICE ALERT: ieatlms3901;check fm supervision nodes RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (2) IS LOWER THAN 80% OF ENM NODES (6)
[1574067051] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574067077] SERVICE ALERT: ieatlms3901;check fm supervision nodes RNC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (6) IS HIGHER THAN 95% OF ENM NODES (6)
[1574067171] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574067203] SERVICE ALERT: ieatlms3901;check fm supervision nodes RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (2) IS LOWER THAN 80% OF ENM NODES (6)
[1574067291] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574067330] SERVICE ALERT: ieatlms3901;check fm supervision nodes RNC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (6) IS HIGHER THAN 95% OF ENM NODES (6)
[1574067584] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574067715] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1746) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574067840] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1574067887] SERVICE ALERT: ieatlms3901;check pm supervision nodes MSC-BC-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (12)
[1574067964] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1574068013] SERVICE ALERT: ieatlms3901;check pm supervision nodes MSC-BC-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1574068091] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (11) IS LOWER THAN 80% OF ENM NODES (20)
[1574068177] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (178) IS BETWEEN 5% AND 20% OF ENM NODES (2560)
[1574068218] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (0)
[1574068303] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (240) IS BETWEEN 5% AND 20% OF ENM NODES (2560)
[1574068373] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (750)
[1574068450] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (240) IS BETWEEN 5% AND 20% OF ENM NODES (2560)
[1574068521] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1574069136] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2326) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574069312] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2558) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1574069527] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574069654] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574069784] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574070146] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2421) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574070289] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2421) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574070418] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2421) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574072085] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN ERBS NODES !
[1574072094] SERVICE ALERT: ieatlms3901;check cm supervision nodes Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1200)
[1574072220] SERVICE ALERT: ieatlms3901;check cm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1574072267] SERVICE ALERT: ieatlms3901;check pm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (2560)
[1574072419] SERVICE ALERT: ieatlms3901;check synchronized RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (3849)
[1574072424] SERVICE ALERT: ieatlms3901;check pm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574072550] SERVICE ALERT: ieatlms3901;check synchronized RBS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (3849)!
[1574073133] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (14) IS LOWER THAN 80% OF ENM NODES (24)
[1574073259] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (16) IS LOWER THAN 80% OF ENM NODES (24)
[1574073383] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1574073394] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0)
[1574073769] SERVICE ALERT: ieatlms3901;check fm supervision nodes SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (80)
[1574073895] SERVICE ALERT: ieatlms3901;check fm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1574073921] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574074028] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574074996] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574075182] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574075369] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574075731] SERVICE ALERT: ieatwlvm5116;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1574077000] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574077127] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574077254] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical EPG;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1574077254] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574078451] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2400) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574078592] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (2400) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574078614] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574078719] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (2400) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574079290] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (160) IS BETWEEN 5% AND 20% OF ENM NODES (2560)
[1574079416] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (160) IS BETWEEN 5% AND 20% OF ENM NODES (2560)
[1574079562] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (160) IS BETWEEN 5% AND 20% OF ENM NODES (2560)
[1574080369] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574080460] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574080586] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574080712] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574081239] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574081370] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574081382] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574081508] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574081533] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 73.04%
[1574081677] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.75%
[1574081827] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 16.07%
[1574082340] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574083897] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (41) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574084724] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (668) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574084878] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2444) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1574085116] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1574085116] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2480) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574085134] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1652) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574085268] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1657) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574085396] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1657) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574086835] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures ERBS;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1574086835] SERVICE ALERT: ieatlms3901;check fm hb failures ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (41) IS LOWER THAN 5% OF ENM NODES (2560)
[1574087252] HOST ALERT: ieatnetsimv5116-115;DOWN;SOFT;1;PING NOK
[1574087314] HOST ALERT: ieatnetsimv5116-115;UP;SOFT;1;PING OK
[1574087504] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;OK;HARD;3;OK- NUMBER OF PM FILES (53) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (50)
[1574088331] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (737) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574088473] SERVICE ALERT: ieatnetsimv9041-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574088562] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;OK;HARD;3;OK- NUMBER OF PM FILES (1200) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (1200)
[1574088680] SERVICE ALERT: ieatnetsimv9041-08;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574088889] SERVICE ALERT: ieatnetsimv9041-08;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574089609] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (2560) IS LOWER THAN 80% OF ENM NODES (5120)
[1574089735] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574090046] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (107) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (100)
[1574090859] SERVICE ALERT: ieatnetsimv7035-84;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574091016] SERVICE ALERT: ieatnetsimv7035-84;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574091125] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574091179] SERVICE ALERT: ieatnetsimv7035-84;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574091269] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574091309] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574091349] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1) IS HIGHER THAN 20% OF ENM NODES (4)
[1574091373] SERVICE ALERT: ieatnetsimv7035-84;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574091393] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574091451] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3) IS HIGHER THAN 20% OF ENM NODES (13)
[1574091476] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1) IS HIGHER THAN 20% OF ENM NODES (4)
[1574091518] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574091521] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (37) IS LOWER THAN 80% OF ENM  NODES (98)
[1574091575] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574091605] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1) IS HIGHER THAN 20% OF ENM NODES (4)
[1574091630] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (5462) IS LOWER THAN 80% OF ENM  NODES (14004)
[1574091677] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (94) IS HIGHER THAN 95% OF ENM  NODES (98)
[1574091699] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574091736] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (2) IS LOWER THAN 80% OF ENM  NODES (4)
[1574091763] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12103) IS FROM 80% TO 95% OF ENM  NODES (14004)
[1574091869] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (2) IS LOWER THAN 80% OF ENM  NODES (4)
[1574091926] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12063) IS FROM 80% TO 95% OF ENM  NODES (14004)
[1574091938] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (621) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574091995] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (2) IS LOWER THAN 80% OF ENM  NODES (4)
[1574092073] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (621) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574092201] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (621) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574092491] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;OK;HARD;3;OK
[1574093322] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1574093739] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13964) IS HIGHER THAN 95% OF ENM  NODES (14004)
[1574093800] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (4)!
[1574095211] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN RNC NODES !
[1574096201] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1672) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574099409] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (572) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574107384] SERVICE ALERT: ieatnetsimv7035-105;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574107546] SERVICE ALERT: ieatnetsimv7035-105;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574108363] SERVICE ALERT: ieatnetsimv7035-95;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574108523] SERVICE ALERT: ieatnetsimv7035-95;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574108550] SERVICE ALERT: ieatlms3901;check cm supervision nodes EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (50) IS LOWER THAN 80% OF ENM NODES (150)
[1574108683] SERVICE ALERT: ieatlms3901;check cm supervision nodes EPG;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1574108686] SERVICE ALERT: ieatnetsimv7035-95;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574109072] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13179) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14004)
[1574109213] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13179) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14004)
[1574109238] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574109364] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13179) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14004)
[1574109390] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574109543] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574110491] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574110525] SERVICE ALERT: ieatnetsimv7035-95;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574110557] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574110612] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574110732] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574110779] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574111118] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574112987] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13726) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14004)
[1574113111] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574115114] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574115496] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574115750] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574117668] SERVICE FLAPPING ALERT: ieatnetsimv7035-51;check netsim nodes;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1574117668] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574121016] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (600) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)