[1567206069] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1567206069] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1567206375] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1567206559] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs llt heartbeat hc;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1567206559] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1567206977] SERVICE FLAPPING ALERT: ieatlms5735-1;check storagepool hc;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1567206977] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1567207156] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567207337] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567207413] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567207518] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567207528] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1567207593] SERVICE ALERT: ieatlms5735-1;check node fs;OK;SOFT;2;OK
[1567207623] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1567207774] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567207955] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567208136] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567209261] SERVICE ALERT: ieatnetsimv5116-79;check netsim cpu;OK;HARD;3;OK- CPU IS USED 23.76%
[1567209997] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1567210306] SERVICE FLAPPING ALERT: ieatnetsimv5116-66;check netsim nodes;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1567210306] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567210403] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.78%
[1567210554] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 73.01%
[1567211057] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567211179] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567211253] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567211452] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567212907] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1567213363] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567214990] SERVICE FLAPPING ALERT: ieatlms3901;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567215133] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1567215312] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1567215560] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;SOFT;2;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1567216329] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.15%
[1567216480] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.43%
[1567217258] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567217439] SERVICE ALERT: ieatlms5735-1;check node fs;OK;SOFT;2;OK
[1567219420] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567219601] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567219643] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12884) IS FROM 80% TO 95% OF ENM COMECIM NODES (14049)
[1567219782] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567219790] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12884) IS FROM 80% TO 95% OF ENM COMECIM NODES (14049)
[1567219938] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12884) IS FROM 80% TO 95% OF ENM COMECIM NODES (14049)
[1567220240] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567220421] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567220603] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567221504] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.43%
[1567221643] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1567221655] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 21.35%
[1567222886] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12793) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567223028] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12793) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567223169] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12793) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567224514] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.89%
[1567224594] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567224664] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 21.04%
[1567226883] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567227803] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567227864] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567228044] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567228225] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567228723] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567228904] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567229084] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567229265] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567231886] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567232925] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1567234887] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.81%
[1567235041] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567235195] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567236253] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567237049] SERVICE FLAPPING ALERT: ieatlms5735-1;check system service hc;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1567238102] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567238853] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567239022] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs llt heartbeat hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567239215] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.13%
[1567239365] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 30.62%
[1567239773] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567239773] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567239805] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.29%
[1567239956] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 27.05%
[1567240107] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.82%
[1567240257] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 60.18%
[1567240693] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567240947] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567241127] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567241308] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567241985] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567242166] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567242347] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567242427] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 77.05%
[1567242803] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567243169] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567243349] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567243530] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567244648] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (23765) IS BETWEEN 80% AND 95% OF ENM NODES (25422)
[1567244964] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH CM SUPERVISION (23479) IS BETWEEN 80% AND 95% OF ENM NODES (25422)
[1567245321] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (23108) IS BETWEEN 80% AND 95% OF ENM NODES (25422)
[1567245416] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5951) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1567245559] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5798) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1567245702] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (5673) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1567246058] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.87%
[1567246208] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 55.82%
[1567247203] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.31 seconds)
[1567248123] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567248952] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.38%
[1567249044] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.49 seconds)
[1567249102] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 59.99%
[1567251534] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567252591] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567252651] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567252853] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.57%
[1567253004] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.17%
[1567253015] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567253208] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1567253214] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567253414] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567253916] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.67%
[1567254067] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 11.46%
[1567254217] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.96%
[1567254368] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 11.58%
[1567254488] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567254668] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567255308] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567255345] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567255352] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (20114) IS LOWER THAN 80% OF ENM NODES (25422)
[1567255496] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567256618] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.23%
[1567256769] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.65%
[1567257161] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567257188] SERVICE ALERT: ieatnetsimv7035-11;check netsim disk;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567257313] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.28%
[1567257372] SERVICE ALERT: ieatnetsimv7035-11;check netsim disk;OK;SOFT;2;OK
[1567257469] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567257478] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical DSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567257621] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.83%
[1567257734] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567257778] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567258655] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1567259575] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567260692] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1567261634] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567261690] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567261784] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.2%
[1567261935] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.51%
[1567262086] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.71%
[1567264069] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567264250] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567264431] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567265251] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567265432] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567265613] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567265848] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.22%
[1567265999] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.16%
[1567267474] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567268091] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1567268981] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (20564) IS BETWEEN 80% AND 95% OF ENM NODES (25422)
[1567269690] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567270084] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567270272] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567270459] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567270601] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567272327] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567273355] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.73%
[1567273505] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.79%
[1567274529] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5953) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1567274569] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567275121] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567275452] SERVICE FLAPPING ALERT: ieatlms5735-1;check storagepool hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567276041] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567276294] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (24182) IS HIGHER THAN 95% OF ENM NODES (25422)
[1567276360] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (7142) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1567276962] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1567277256] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97%
[1567277406] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.46%
[1567277562] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567277712] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.4%
[1567279066] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567281465] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567281615] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.18%
[1567281771] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567287993] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567288866] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567289786] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1567289916] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567290088] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567290274] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567290706] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567291821] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.9%
[1567291971] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.73%
[1567292122] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 56.5%