[1569453075] SERVICE FLAPPING ALERT: ieatlms3901;check vcs sg;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1569473233] SERVICE FLAPPING ALERT: ieatnetsimv5116-96;check netsim disk;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1569478274] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1569480084] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1569480214] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1569480344] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1569480433] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;OK;HARD;3;OK
[1569480642] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1569480828] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1569481013] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1569482878] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1569484033] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1569484154] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1569484274] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1569485754] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1569486831] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (24450) IS BETWEEN 80% AND 95% OF ENM NODES (25989)
[1569487154] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH CM SUPERVISION (24450) IS BETWEEN 80% AND 95% OF ENM NODES (26966)
[1569487542] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (24450) IS BETWEEN 80% AND 95% OF ENM NODES (26966)
[1569487788] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (2499) IS LOWER THAN 80% OF ENM SNMP NODES (5015)
[1569487929] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (2825) IS LOWER THAN 80% OF ENM SNMP NODES (5015)
[1569488071] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1569488071] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (3680) IS LOWER THAN 80% OF ENM SNMP NODES (5015)
[1569488085] HOST ALERT: ieatrcxb3804;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1569488559] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (26965) IS HIGHER THAN 95% OF ENM NODES (26966)
[1569489266] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1569489662] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1569489893] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4889) IS HIGHER THAN 95% OF ENM SNMP NODES (5015)
[1569490094] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1)
[1569490144] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1569490330] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1569490500] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (914) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1569490626] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (1388) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1569490752] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1569491272] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1569493698] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1569493823] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1569493955] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1569494116] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1569494301] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1569494486] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1569494983] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1569495142] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1569495289] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1569500124] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1569500304] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1569500339] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1569500483] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1569500533] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1569500720] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1569500722] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1569502534] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1569502717] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1569502900] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1569504315] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4216) IS FROM 80% TO 95% OF ENM SNMP NODES (5015)
[1569504484] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4228) IS FROM 80% TO 95% OF ENM SNMP NODES (5015)
[1569504638] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4328) IS FROM 80% TO 95% OF ENM SNMP NODES (5015)
[1569505874] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;OK;HARD;3;OK
[1569506461] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1569506461] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4886) IS HIGHER THAN 95% OF ENM SNMP NODES (5015)
[1569507152] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1569507152] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1569507365] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1569507494] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1569507624] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1569507809] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1569530617] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1569534050] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1569534947] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6593) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1569536200] SERVICE ALERT: ieatnetsimv7035-21;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1569536331] SERVICE ALERT: ieatnetsimv7035-21;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1569536464] SERVICE ALERT: ieatnetsimv7035-21;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1569538278] SERVICE ALERT: ieatnetsimv7035-21;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES