[1566860955] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.63%
[1566861106] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.81%
[1566861698] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 69.65%
[1566867139] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.39%
[1566867290] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.14%
[1566870534] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.57%
[1566870685] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 63.28%
[1566870836] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.21%
[1566870986] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 69.2%
[1566871040] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.27%
[1566871191] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.8%
[1566871342] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.67%
[1566871493] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.68%
[1566872529] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.07%
[1566872680] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 73.91%
[1566874277] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.69%
[1566874428] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.95%
[1566874512] SERVICE ALERT: ieatnetsimv5116-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566874633] SERVICE ALERT: ieatnetsimv5116-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566874753] SERVICE ALERT: ieatnetsimv5116-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566877555] SERVICE ALERT: ieatnetsimv5116-16;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.03%
[1566877705] SERVICE ALERT: ieatnetsimv5116-16;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.55%
[1566878772] SERVICE ALERT: ieatnetsimv5116-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.36%
[1566878923] SERVICE ALERT: ieatnetsimv5116-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.88%
[1566879257] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566879407] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.12%
[1566879558] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.93%
[1566879708] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.16%
[1566883458] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.32%
[1566883609] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.87%
[1566883760] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.16%
[1566885378] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.3%
[1566885459] SERVICE ALERT: ieatnetsimv5116-69;check netsim disk;OK;HARD;3;OK
[1566885480] SERVICE ALERT: ieatnetsimv5116-77;check netsim disk;OK;HARD;3;OK
[1566885501] SERVICE ALERT: ieatnetsimv5116-85;check netsim disk;OK;HARD;3;OK
[1566885529] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.52%
[1566885554] SERVICE ALERT: ieatnetsimv5116-12;check netsim disk;OK;HARD;3;OK
[1566885679] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.59%
[1566885689] SERVICE ALERT: ieatnetsimv5116-64;check netsim disk;OK;HARD;3;OK
[1566885830] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.17%
[1566886044] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.87%
[1566886194] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.97%
[1566887390] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.84%
[1566887541] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.38%
[1566887692] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.64%
[1566891322] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.83%
[1566891473] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.84%
[1566891623] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 93.18%
[1566892737] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.97%
[1566892740] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.47%
[1566892744] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566892888] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 55.16%
[1566892890] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 56.22%
[1566892896] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566893041] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.76%
[1566893047] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566893192] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 54.12%
[1566894713] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1566894864] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566895015] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566895255] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.57%
[1566895346] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566896248] SERVICE ALERT: ieatnetsimv7035-85;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566896383] SERVICE ALERT: ieatnetsimv7035-85;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566896447] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566896513] SERVICE ALERT: ieatnetsimv7035-85;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566896767] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566896814] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566896960] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566896962] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1502) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566897109] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566897788] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566897932] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4068) IS HIGHER THAN 95% OF ENM CPP NODES (4224)
[1566898209] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566898328] SERVICE ALERT: ieatnetsimv7035-85;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566898399] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566898425] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1566898552] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1566898588] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566898693] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1566898751] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566898934] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566899603] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566899698] SERVICE ALERT: ieatnetsimv5116-99;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.14%
[1566899732] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566899738] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566899873] SERVICE ALERT: ieatnetsimv5116-99;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566899918] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566900104] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566900278] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;OK;HARD;3;OK- CPU IS USED 76.69%
[1566900695] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1)
[1566900763] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1566900943] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1566901066] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1566901122] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1566901190] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1566901223] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1)
[1566901549] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1566901679] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1566901809] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1566902246] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 82.75%
[1566902719] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566902904] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566903892] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566904043] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566904167] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566904195] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.95%
[1566904293] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566904300] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1566904319] SERVICE ALERT: ieatlms3901;check pm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (23829) IS HIGHER THAN 95% OF ENM NODES (1476)
[1566904418] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566904424] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1566904551] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1566904682] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566904796] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566904838] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1566904843] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566904888] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566904941] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566904986] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1566905014] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566905071] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1566905074] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566905111] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1566905231] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.99%
[1566905256] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566905345] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23829) IS BETWEEN 80% AND 95% OF ENM NODES (25421)
[1566905382] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 66.89%
[1566905532] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.42%
[1566905574] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23829) IS BETWEEN 80% AND 95% OF ENM NODES (25421)
[1566905694] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566905802] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23829) IS BETWEEN 80% AND 95% OF ENM NODES (25421)
[1566905847] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566907119] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566907305] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1566907488] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566907581] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.38%
[1566907718] SERVICE ALERT: ieatnetsimv5116-59;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566907731] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.95%
[1566907825] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;OK;HARD;3;OK- CPU IS USED 36.05%
[1566907856] SERVICE ALERT: ieatnetsimv5116-59;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566907882] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.05%
[1566907912] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1566907937] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566907945] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.66%
[1566907987] SERVICE ALERT: ieatnetsimv5116-59;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566908032] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1566908032] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.34%
[1566908092] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1566908095] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.54%
[1566908118] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566908160] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1566908183] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.63%
[1566908246] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.31%
[1566908273] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1566908288] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1566908298] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566908333] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.67%
[1566909201] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566909351] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566909353] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566909477] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.38%
[1566909505] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566909535] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566909627] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.72%
[1566909940] SERVICE ALERT: ieatnetsimv5116-59;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566910133] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1566910219] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1566910285] SERVICE ALERT: ieatnetsimv5116-46;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566910385] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1566910419] SERVICE ALERT: ieatnetsimv5116-46;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566910533] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1566910548] SERVICE ALERT: ieatnetsimv5116-46;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566910623] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1566910756] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566910817] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1566910938] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566912341] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1566912479] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1566912608] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1566913127] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1566913286] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.62%
[1566913437] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.93%
[1566913588] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.56%
[1566914209] SERVICE ALERT: ieatnetsimv5116-46;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566914416] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1566914416] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1566915056] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.94%
[1566915121] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566915206] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 70.69%
[1566915234] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566915306] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1566915492] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566916224] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1566916352] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1566916480] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1566916732] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1566916858] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1566916987] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1566917048] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 801.02 seconds)
[1566917250] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566917356] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566917518] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566917540] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566917798] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566917839] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (16579) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21329)
[1566917968] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1566918888] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;HARD;3;(Service check timed out after 800.04 seconds)
[1566919076] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 82.08%
[1566919194] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1566919280] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566919320] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1566919375] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1566919417] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1566919461] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566919478] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1566919515] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1566919555] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1566919642] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566919994] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 19.89%
[1566920258] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566920336] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 11.81%
[1566921352] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3965) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1566921486] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4068) IS HIGHER THAN 95% OF ENM CPP NODES (4224)
[1566921490] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim nodes;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1566921490] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566921539] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17431) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21329)
[1566922558] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566922708] SERVICE ALERT: ieatnetsimv5116-96;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.59%
[1566923039] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566923084] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566923119] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1566923165] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566923209] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566923244] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1566923291] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized nodes cpp;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1566923291] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566923334] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized nodes com/ecim;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1566923334] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566923368] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1566923368] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1566923536] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566923916] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566924158] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566924294] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566926226] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1566926424] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1566926624] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1566927510] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED %
[1566928019] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.6%
[1566928169] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.94%
[1566930305] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1566930987] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.18%
[1566931138] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 62.22%
[1566931567] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566931945] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566932198] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566936047] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566936379] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1566937283] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.11%
[1566937434] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.92%
[1566937438] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566937569] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566937703] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES