[1570057598] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1570057632] HOST ALERT: ieatrcxb3807;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570057696] HOST ALERT: ieatrcxb3807;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1570057781] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1570057964] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1570058732] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4647) IS FROM 80% TO 95% OF ENM SNMP NODES (5015)
[1570058878] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4691) IS FROM 80% TO 95% OF ENM SNMP NODES (5015)
[1570059060] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4750) IS FROM 80% TO 95% OF ENM SNMP NODES (5015)
[1570060209] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570060299] HOST ALERT: ieatrcxb3778;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570060363] HOST ALERT: ieatrcxb3778;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1570060407] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1570060603] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570060887] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1570060887] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (5014) IS HIGHER THAN 95% OF ENM SNMP NODES (5015)
[1570061626] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1570063010] HOST ALERT: ieatrcxb3911;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1570063074] HOST ALERT: ieatrcxb3911;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1570063490] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1570063674] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1570063856] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1570064957] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570066352] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4471) IS FROM 80% TO 95% OF ENM SNMP NODES (5015)
[1570066534] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4571) IS FROM 80% TO 95% OF ENM SNMP NODES (5015)
[1570066679] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4631) IS FROM 80% TO 95% OF ENM SNMP NODES (5015)
[1570067678] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1570068511] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (5000) IS HIGHER THAN 95% OF ENM SNMP NODES (5015)
[1570073243] SERVICE FLAPPING ALERT: ieatlms5735-1;check system service hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570073516] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs llt heartbeat hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570073748] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570074433] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570074807] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2560) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570075165] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570075622] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3849) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570079743] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570080481] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570083028] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1570083691] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570084343] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570085497] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1570087208] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1570087273] HOST ALERT: ieatrcxb6296-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570087337] HOST ALERT: ieatrcxb6296-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1570087389] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1570087571] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1570087795] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1570089040] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1570089046] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570089228] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1570089255] SERVICE ALERT: ieatnetsimv017-30;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570089382] HOST ALERT: ieatrcxb6220-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570089412] SERVICE ALERT: ieatnetsimv017-30;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570089417] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1570089459] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570089472] HOST ALERT: ieatrcxb6220-1;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1570089536] HOST ALERT: ieatrcxb6220-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1570089564] SERVICE ALERT: ieatnetsimv017-30;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570089655] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1570089849] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570090906] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570091087] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570091268] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570091741] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570091934] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570092029] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570092216] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570092404] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570093033] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1570093128] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570093256] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570093309] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570093437] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570093490] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1570093617] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570093917] HOST ALERT: ieatrcxb6204;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570093950] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570094007] HOST ALERT: ieatrcxb6204;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1570094071] HOST ALERT: ieatrcxb6204;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1570094143] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1570094272] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570094337] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570094519] SERVICE FLAPPING ALERT: ieatlms3901;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570096233] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570096290] HOST ALERT: ieatrcxb6268-1;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1570096380] HOST ALERT: ieatrcxb6268-1;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1570096440] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1570096444] HOST ALERT: ieatrcxb6268-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1570096633] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570096695] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1570096877] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1570098330] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570099878] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2055) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1570100019] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2098) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1570100152] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2130) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1570100340] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1570100750] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570100878] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1570101965] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1570102105] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570102299] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1570102492] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1570102731] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized snmp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570103649] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570103798] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570103809] SERVICE FLAPPING ALERT: ieatlms3901;check hw resources hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1570103948] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570105140] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570105328] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570105416] SERVICE ALERT: ieatnetsimv017-14;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570105516] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570105575] SERVICE ALERT: ieatnetsimv017-14;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570105724] SERVICE ALERT: ieatnetsimv017-14;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570106168] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570107244] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570107402] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570107406] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1570107406] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570107555] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570109041] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (27818) IS HIGHER THAN 95% OF ENM NODES (27047)
[1570109382] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570109403] SERVICE ALERT: ieatnetsimv017-30;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570110084] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (22698) IS BETWEEN 80% AND 95% OF ENM NODES (27047)
[1570110430] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (22698) IS BETWEEN 80% AND 95% OF ENM NODES (27047)
[1570110653] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (22698) IS BETWEEN 80% AND 95% OF ENM NODES (27047)
[1570112311] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13236) IS FROM 80% TO 95% OF ENM COMECIM NODES (14053)
[1570112480] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13099) IS FROM 80% TO 95% OF ENM COMECIM NODES (14053)
[1570112635] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13139) IS FROM 80% TO 95% OF ENM COMECIM NODES (14053)
[1570112796] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570112981] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570113166] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570113539] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570113720] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570113901] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1570114468] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13482) IS HIGHER THAN 95% OF ENM COMECIM NODES (14053)
[1570115211] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570115392] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570115573] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1570115573] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1570115761] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570115943] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570116123] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1570116123] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570116295] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13321) IS FROM 80% TO 95% OF ENM COMECIM NODES (14053)
[1570116459] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13382) IS HIGHER THAN 95% OF ENM COMECIM NODES (14053)
[1570116832] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570118404] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13302) IS FROM 80% TO 95% OF ENM COMECIM NODES (14053)
[1570118563] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13289) IS FROM 80% TO 95% OF ENM COMECIM NODES (14053)
[1570118727] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13437) IS HIGHER THAN 95% OF ENM COMECIM NODES (14053)
[1570120583] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12976) IS FROM 80% TO 95% OF ENM COMECIM NODES (14053)
[1570120746] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12924) IS FROM 80% TO 95% OF ENM COMECIM NODES (14053)
[1570120905] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12901) IS FROM 80% TO 95% OF ENM COMECIM NODES (14053)
[1570122042] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570122737] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized comecim;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1570122737] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13605) IS HIGHER THAN 95% OF ENM COMECIM NODES (14053)
[1570124369] HOST ALERT: ieatnetsimv7035-102;DOWN;SOFT;1;PING NOK
[1570124431] HOST ALERT: ieatnetsimv7035-102;UP;SOFT;1;PING OK
[1570128161] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13340) IS FROM 80% TO 95% OF ENM COMECIM NODES (14053)
[1570128265] HOST ALERT: ieatnetsimv7035-26;DOWN;SOFT;1;PING NOK
[1570128312] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13776) IS HIGHER THAN 95% OF ENM COMECIM NODES (14053)
[1570128327] HOST ALERT: ieatnetsimv7035-26;UP;SOFT;1;PING OK
[1570137234] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1570137784] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1570139665] HOST ALERT: ieatnetsimv7035-31;DOWN;SOFT;1;PING NOK
[1570139665] HOST ALERT: ieatnetsimv7035-44;DOWN;SOFT;1;PING NOK
[1570139727] HOST ALERT: ieatnetsimv7035-31;UP;SOFT;1;PING OK
[1570139727] HOST ALERT: ieatnetsimv7035-44;UP;SOFT;1;PING OK
[1570141469] HOST ALERT: ieatnetsimv7035-76;DOWN;SOFT;1;PING NOK
[1570141531] HOST ALERT: ieatnetsimv7035-76;UP;SOFT;1;PING OK
[1570141673] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)