[1566949376] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566949699] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1566950429] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (17061) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21329)
[1566950688] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566951984] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.03%
[1566952135] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.15%
[1566952286] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.79%
[1566952440] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.12%
[1566952590] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.27%
[1566953653] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566953694] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1566953774] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566953895] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;OK;HARD;3;OK
[1566954115] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1566954115] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (17261) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21329)
[1566957495] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566957616] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566957737] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566960720] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.54%
[1566960858] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1566960871] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.21%
[1566961021] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.84%
[1566961041] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED %
[1566961222] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2924) IS LOWER THAN 80% OF ENM CPP NODES (4224)
[1566961355] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2909) IS LOWER THAN 80% OF ENM CPP NODES (4224)
[1566961498] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3144) IS LOWER THAN 80% OF ENM CPP NODES (4224)
[1566964652] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.6%
[1566964803] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.98%
[1566964953] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.65%
[1566965104] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.48%
[1566965128] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3687) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1566971786] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1566971917] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566971985] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1566972138] SERVICE ALERT: ieatnetsimv5116-09;check netsim disk;OK;HARD;3;OK
[1566972187] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1566972288] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1566972671] SERVICE ALERT: ieatnetsimv5116-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.29%
[1566972822] SERVICE ALERT: ieatnetsimv5116-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 63.98%
[1566972973] SERVICE ALERT: ieatnetsimv5116-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.45%
[1566973124] SERVICE ALERT: ieatnetsimv5116-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 60.12%
[1566974375] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM CPP NODES (0)
[1566974621] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1566975285] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566975328] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566975405] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566975416] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1566975448] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566975503] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1566975526] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566975539] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1566975569] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566975844] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1566975866] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1566976455] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1566976585] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1566977234] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (24197) IS HIGHER THAN 95% OF ENM NODES (25421)
[1566977429] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1566977429] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1566977459] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1566977459] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1566977695] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1566977823] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1566977951] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1566978877] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (4384) IS FROM 80% TO 95% OF ENM SNMP NODES (5018)
[1566979012] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566979209] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1566979342] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1566979370] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566979481] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1566979539] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2310) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566979568] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566979656] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.45%
[1566979711] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2310) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566979759] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1566979806] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.48%
[1566979854] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3687) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1566979878] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566979901] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2310) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566979957] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.67%
[1566980012] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3687) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1566980061] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566980107] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.82%
[1566980158] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized cpp;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1566980158] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3628) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1566980190] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566980194] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1566980258] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.41%
[1566980328] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1566980330] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566980455] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1566981037] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1566981045] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566981072] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1566981093] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566981164] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1566981176] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566981203] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1566981227] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566981290] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1566981309] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566981335] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1566981343] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1566981374] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1566981546] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1566981744] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1566981826] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1566981952] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1566982078] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2179)
[1566982727] SERVICE ALERT: ieatnetsimv5116-11;check netsim disk;OK;HARD;3;OK
[1566982770] SERVICE ALERT: ieatnetsimv5116-17;check netsim disk;OK;HARD;3;OK
[1566983890] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.97%
[1566984040] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.47%
[1566984109] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566984191] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.78%
[1566984261] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566984416] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566985143] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566985274] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (6953) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1566985295] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566985425] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1566985425] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1566985447] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566989169] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566989354] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566989493] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566991107] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.91%
[1566991258] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.45%
[1566992678] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.04%
[1566993835] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566994021] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566994206] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566996309] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566996460] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566996612] SERVICE FLAPPING ALERT: ieatnetsimv5116-71;check netsim cpu;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1566996612] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566996862] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566997725] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566997913] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566998100] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566998317] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566998457] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1566998498] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566998508] SERVICE ALERT: ieatwlvm5116;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1566998609] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.13%
[1566998641] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1566998680] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566998692] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566998759] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.21%
[1566998823] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1566998839] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1566999117] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1567000242] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.11%
[1567000667] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567003272] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567003719] SERVICE ALERT: ieatnetsimv5116-81;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.86%
[1567003869] SERVICE ALERT: ieatnetsimv5116-81;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 43.24%
[1567005137] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567005323] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567005510] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1567005510] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567006052] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1567006231] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1567006409] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1567006520] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567006702] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567006884] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1567008239] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.07%
[1567008390] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 13.76%
[1567008713] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1721) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1567008745] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567010528] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1567010528] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2310) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1567010607] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567010789] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567010795] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized cpp;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1567012141] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.85%
[1567012292] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 17.35%
[1567014571] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567014752] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567016912] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.79%
[1567017063] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.75%
[1567017214] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.05%
[1567017777] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.02%
[1567017905] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1567017928] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.14%
[1567017947] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 25.16%
[1567018078] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 80.99%
[1567020336] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567020518] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567021914] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.28%
[1567022064] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.89%
[1567022215] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.62%
[1567022365] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.89%
[1567022517] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.28%
[1567022599] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567022667] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.84%
[1567023576] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567024444] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.83%
[1567024513] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1567024595] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.98%
[1567024940] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1869) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1567025876] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567026057] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567026238] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567026752] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2310) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1567027421] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.08%
[1567027572] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.52%
[1567027764] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567027944] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1567028125] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE