[1575245299] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.63%
[1575245450] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.42%
[1575248310] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1575249720] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575249876] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575250033] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575252146] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575255469] SERVICE ALERT: ieatlms7555;check node fs;OK;HARD;3;OK
[1575257523] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1578) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1575257655] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1599) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1575257780] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1617) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1575259588] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1737) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1575260212] HOST ALERT: ieatnetsimv7048-20;DOWN;SOFT;1;PING NOK
[1575260274] HOST ALERT: ieatnetsimv7048-20;UP;SOFT;1;PING OK
[1575266980] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575267165] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575267350] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575269216] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575271705] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575271861] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575272018] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575272548] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2329) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1575272674] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2211) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1575272800] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (2152) IS BETWEEN 80% AND 95% OF ENM NODES (2560)
[1575274608] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (2539) IS HIGHER THAN 95% OF ENM NODES (2560)
[1575274983] SERVICE ALERT: ieatnetsimv9041-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575275175] SERVICE ALERT: ieatnetsimv9041-11;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1575275228] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.31%
[1575275299] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1575275357] SERVICE ALERT: ieatnetsimv9041-11;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1575275379] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.93%
[1575275428] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1575275529] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.97%
[1575275558] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1575275609] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (14) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575275655] SERVICE ALERT: ieatlms7555;check node fs;OK;HARD;3;OK
[1575275732] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575275875] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (14) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575276441] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.77%
[1575276592] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.96%
[1575276742] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.98%
[1575277259] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575277355] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575277410] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575277505] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.96%
[1575277561] SERVICE ALERT: ieatnetsimv9034-100;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575277577] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1575277656] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575277727] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.99%
[1575277730] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (141) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1575277856] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (142) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1575277878] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575277982] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (176) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1575278128] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575278270] SERVICE ALERT: ieatnetsimv9034-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575278278] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575278421] SERVICE ALERT: ieatnetsimv9034-71;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575278429] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575278574] SERVICE ALERT: ieatnetsimv9034-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575279164] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1575279482] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1575280977] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575281118] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575281169] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13030) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575281288] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575281302] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12867) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575281439] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12591) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575281464] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575281586] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (9) IS LOWER THAN 5% OF ENM NODES (1200)
[1575281591] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575281746] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575281752] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575281880] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575281882] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575282016] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575282067] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575282253] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575283133] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (50) IS LOWER THAN 80% OF ENM  NODES (80)
[1575283251] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (11144) IS LOWER THAN 80% OF ENM  NODES (14005)
[1575283361] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575283539] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575283672] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575283823] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575284196] SERVICE ALERT: ieatnetsimv9034-02;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.32%
[1575284347] SERVICE ALERT: ieatnetsimv9034-02;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.84%
[1575284933] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575284965] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1575285066] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13499) IS HIGHER THAN 95% OF ENM  NODES (14005)
[1575285069] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575285190] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575285192] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (115) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1575285193] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (4) IS LOWER THAN 5% OF ENM NODES (80)
[1575285319] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (68) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1575285327] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;SOFT;2;OK
[1575285415] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1575285454] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (85) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1575285463] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575285600] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575285632] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575285739] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;HARD;3;OK
[1575285856] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575285982] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575286390] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (50)
[1575286516] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (25) IS HIGHER THAN 20% OF ENM NODES (50)
[1575286728] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (27) IS HIGHER THAN 20% OF ENM NODES (50)
[1575286883] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12543) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575287015] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13598) IS HIGHER THAN 95% OF ENM  NODES (14005)
[1575287789] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575287789] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575288797] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (11) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575288960] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575289069] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575289069] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (691) IS HIGHER THAN 20% OF ENM NODES (1200)
[1575289103] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (19) IS HIGHER THAN 20% OF ENM NODES (80)
[1575289359] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575289495] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;SOFT;2;OK
[1575289755] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575289899] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1575290028] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575290157] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575290342] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575290361] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (21) IS HIGHER THAN 20% OF ENM NODES (100)
[1575290381] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575290491] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (21) IS HIGHER THAN 20% OF ENM NODES (100)
[1575290507] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (76) IS HIGHER THAN 95% OF ENM  NODES (80)
[1575290615] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575290821] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575290908] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575290908] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575290948] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575291032] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575291032] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575291185] SERVICE ALERT: ieatlms3901;check list nodes BSC;CRITICAL;SOFT;1;(Return code of 13 for service 'check list nodes BSC' on host 'ieatlms3901' was out of bounds)
[1575291305] SERVICE ALERT: ieatlms3901;check list nodes BSC;CRITICAL;SOFT;2;(Return code of 13 for service 'check list nodes BSC' on host 'ieatlms3901' was out of bounds)
[1575291386] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575291386] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575291393] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575291426] SERVICE ALERT: ieatlms3901;check list nodes BSC;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1575291517] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1575291596] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575291596] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1575291702] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575291702] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575291724] SERVICE ALERT: ieatlms3901;check list nodes BSC;OK;HARD;3;Files /usr/local/nagios/libexec/s4/list_check_number_nodes/check_ieatnetsimv5116.netsim and /usr/local/nagios/libexec/s4/list_check_number_nodes/host_ieatnetsimv5116.netsim differ
[1575291849] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575291849] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575292139] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575292139] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575292330] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1661) IS BETWEEN 5% AND 20% OF ENM NODES (14005)
[1575292462] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1576) IS BETWEEN 5% AND 20% OF ENM NODES (14005)
[1575292516] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575292551] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575292551] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575292593] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1851) IS BETWEEN 5% AND 20% OF ENM NODES (14005)
[1575292689] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1575292710] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575292710] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575292797] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575292797] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575292815] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (7) IS HIGHER THAN 20% OF ENM NODES (30)
[1575292857] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575292857] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575292875] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575292900] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575292900] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575292941] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575292984] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575292984] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1575293004] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (49) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575293072] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1575293231] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575293368] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;SOFT;2;OK
[1575293446] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575293446] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575293475] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13297) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575293504] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575293539] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575293539] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575293606] SERVICE ALERT: ieatnetsimv7048-01;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575293609] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13613) IS HIGHER THAN 95% OF ENM  NODES (14005)
[1575293635] SERVICE ALERT: ieatnetsimv9034-102;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.11%
[1575293640] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575293645] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.43%
[1575293651] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575293651] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575293684] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575293691] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575293756] SERVICE ALERT: ieatnetsimv7048-01;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 29.51%
[1575293777] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575293785] SERVICE ALERT: ieatnetsimv9034-102;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 57.96%
[1575293796] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.84%
[1575293834] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 39.81%
[1575293848] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575293944] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575293944] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575294005] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575294031] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294031] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575294039] SERVICE ALERT: ieatlms3901;check list nodes BSC - DEMO;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1575294120] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294120] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575294128] SERVICE ALERT: ieatlms3901;check list nodes BSC - DEMO;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1575294189] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294189] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575294197] SERVICE ALERT: ieatlms3901;check list nodes BSC - DEMO;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1575294220] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294220] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (100)
[1575294292] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294292] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294292] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575294361] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294361] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294361] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575294441] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294441] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294441] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575294517] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575294518] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (80)
[1575294602] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294602] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294602] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575294655] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575294822] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294822] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294822] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575294828] SERVICE ALERT: ieatlms3901;check list nodes BSC - DEMO;OK;HARD;3;OK - check_ieatnetsimv5116.netsim and host_ieatnetsimv5116.netsim are the same
[1575294888] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294889] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294889] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575294987] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294987] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575294987] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575295031] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575295031] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575295031] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575295087] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575295087] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575295087] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575295216] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575295216] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575295216] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1575295310] SERVICE ALERT: ieatnetsimv7102-83;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575295311] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.47%
[1575295461] SERVICE ALERT: ieatnetsimv7102-83;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.96%
[1575295461] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 26.5%
[1575295520] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575295520] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575295520] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575295543] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13125) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575295679] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13153) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575295826] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1575295826] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (13173) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575295953] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.72%
[1575295979] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575295979] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575295979] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575295979] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1575296016] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575296016] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575296016] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575296016] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1575296103] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.67%
[1575296133] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575296133] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575296133] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575296133] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1575296202] SERVICE ALERT: ieatlms3901;check pm supervision nodes MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1575296205] SERVICE ALERT: ieatlms3901;check cm supervision nodes RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (3849)
[1575296210] SERVICE ALERT: ieatnetsimv7102-50;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.72%
[1575296252] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575296252] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575296252] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575296252] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1575296299] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.97%
[1575296323] SERVICE ALERT: ieatlms3901;check pm supervision nodes MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1575296326] SERVICE ALERT: ieatlms3901;check cm supervision nodes RBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1575296360] SERVICE ALERT: ieatnetsimv7102-50;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 1.33%
[1575296486] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575296486] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1575296486] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575296486] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1575296531] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575296531] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1575296531] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575296531] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1575296628] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575296628] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1575296628] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575296628] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1575296667] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1575296667] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1575296667] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1575296667] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1575296676] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575296884] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1575297013] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575297145] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575297271] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575297626] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575297641] SERVICE FLAPPING ALERT: ieatlms7555;check node fs;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1575297641] SERVICE ALERT: ieatlms7555;check node fs;OK;HARD;3;OK
[1575297643] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13379) IS HIGHER THAN 95% OF ENM  NODES (14005)
[1575297777] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1575297826] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575297928] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1575297958] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (4) IS LOWER THAN 5% OF ENM NODES (100)
[1575298126] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575298260] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (4) IS LOWER THAN 5% OF ENM NODES (80)
[1575298449] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1575298902] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1575298902] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (50)
[1575299459] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1575299461] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13210) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575299596] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13467) IS HIGHER THAN 95% OF ENM  NODES (14005)
[1575300082] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575300264] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575300323] SERVICE ALERT: ieatnetsimv7035-105;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575300445] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575300477] SERVICE ALERT: ieatnetsimv7035-105;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575300714] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1575300895] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1575300994] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;HARD;3;OK
[1575301077] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1575301224] HOST ALERT: ieatnetsimv7035-40;DOWN;SOFT;1;PING NOK
[1575301226] HOST ALERT: ieatnetsimv7035-90;DOWN;SOFT;1;PING NOK
[1575301286] HOST ALERT: ieatnetsimv7035-40;UP;SOFT;1;PING OK
[1575301288] HOST ALERT: ieatnetsimv7035-90;UP;SOFT;1;PING OK
[1575301684] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575301832] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575301864] HOST ALERT: ieatnetsimv7102-48;DOWN;SOFT;1;PING NOK
[1575301926] HOST ALERT: ieatnetsimv7102-48;UP;SOFT;1;PING OK
[1575301957] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (11) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575301993] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575302119] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575302249] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (80)
[1575302528] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (15) IS HIGHER THAN 20% OF ENM NODES (50)
[1575302545] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575302652] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (17) IS HIGHER THAN 20% OF ENM NODES (50)
[1575302699] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.64%
[1575302780] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (14) IS HIGHER THAN 20% OF ENM NODES (50)
[1575302850] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.4%
[1575303125] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1575303310] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1575303496] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1575303496] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1575304222] HOST ALERT: ieatnetsimv7035-28;DOWN;SOFT;1;PING NOK
[1575304226] HOST ALERT: ieatnetsimv7035-73;DOWN;SOFT;1;PING NOK
[1575304226] HOST ALERT: ieatnetsimv7035-66;DOWN;SOFT;1;PING NOK
[1575304284] HOST ALERT: ieatnetsimv7035-28;UP;SOFT;1;PING OK
[1575304288] HOST ALERT: ieatnetsimv7035-73;UP;SOFT;1;PING OK
[1575304288] HOST ALERT: ieatnetsimv7035-66;UP;SOFT;1;PING OK
[1575304477] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1575304610] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1575304747] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;SOFT;2;OK
[1575305657] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575305784] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575305908] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575305910] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1575306032] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1575306152] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1575306556] SERVICE ALERT: ieatnetsimv7035-80;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575306705] SERVICE ALERT: ieatnetsimv7035-80;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575307090] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (113) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1575308631] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (69) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575308659] SERVICE ALERT: ieatnetsimv7035-80;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1575308730] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12342) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575308771] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (70) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575308810] SERVICE ALERT: ieatnetsimv7035-80;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1575308875] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13706) IS HIGHER THAN 95% OF ENM  NODES (14005)
[1575308901] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (79) IS HIGHER THAN 95% OF ENM  NODES (80)
[1575309174] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (5) IS LOWER THAN 5% OF ENM NODES (100)
[1575309461] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575309599] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575309643] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575309745] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575309775] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575309902] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (24) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575311013] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575311174] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1575311202] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575311363] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1575311392] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575311553] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1575311801] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575311951] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 11.82%
[1575312624] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (11649) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575312781] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (11669) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575312787] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575312791] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575312916] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13341) IS HIGHER THAN 95% OF ENM  NODES (14005)
[1575312921] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575312931] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575312938] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575313047] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575313057] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1575313064] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575313116] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1575313116] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN DSC NODES !
[1575313214] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575313347] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (42) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575313351] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1575313351] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (4) IS LOWER THAN 5% OF ENM NODES (80)
[1575313471] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (42) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575313518] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575313642] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575314202] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (630) IS LOWER THAN 5% OF ENM NODES (14005)
[1575314731] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13297) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575314852] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575314865] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13593) IS HIGHER THAN 95% OF ENM  NODES (14005)
[1575315279] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (49) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575315284] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575315408] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575315532] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575315677] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575315833] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575315990] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575316107] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575316235] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575316365] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575316678] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575316720] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575316817] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13031) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575316820] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575316847] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575316957] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12341) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575316963] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575316966] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575316976] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575317090] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575317096] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12580) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575317125] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575317215] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1575317251] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575317367] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (23) IS LOWER THAN 80% OF ENM  NODES (30)
[1575317394] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575317492] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575317520] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575317616] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (24) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575317831] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1459) IS BETWEEN 5% AND 20% OF ENM NODES (14005)
[1575317899] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (285) IS HIGHER THAN 20% OF ENM NODES (1200)
[1575317960] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (798) IS BETWEEN 5% AND 20% OF ENM NODES (14005)
[1575318090] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1355) IS BETWEEN 5% AND 20% OF ENM NODES (14005)
[1575318772] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MTAS;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1575318772] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575319136] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1575319136] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1575319431] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575319626] SERVICE ALERT: ieatlms7555;check node fs;OK;HARD;3;OK
[1575319972] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1575320584] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1575320724] HOST ALERT: ieatnetsimv7035-62;DOWN;SOFT;1;PING NOK
[1575320786] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1575320786] HOST ALERT: ieatnetsimv7035-62;UP;SOFT;1;PING OK
[1575321503] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (175) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1575322378] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575322519] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575322646] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575322770] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575323064] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575323189] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575323579] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (67) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575323722] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (76) IS HIGHER THAN 95% OF ENM  NODES (80)
[1575324393] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (13) IS HIGHER THAN 20% OF ENM NODES (50)
[1575325108] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (246) IS HIGHER THAN 20% OF ENM NODES (1200)
[1575325970] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.77%
[1575326120] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13512) IS HIGHER THAN 95% OF ENM  NODES (14005)
[1575326342] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575326469] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1575326597] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (7) IS HIGHER THAN 20% OF ENM NODES (30)
[1575326720] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1575327090] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.32%
[1575327464] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575327507] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1575327596] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (71) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575327725] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (77) IS HIGHER THAN 95% OF ENM  NODES (80)
[1575327790] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575327919] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575327934] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12757) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575328049] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575328069] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12795) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575328138] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1575328218] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12817) IS FROM 80% TO 95% OF ENM  NODES (14005)
[1575328298] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575328309] SERVICE ALERT: ieatnetsimv9034-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.87%
[1575328422] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575328546] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575328675] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575329061] SERVICE ALERT: ieatnetsimv9034-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.35%
[1575329553] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575329688] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (71) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575329820] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1575330393] SERVICE ALERT: ieatlms3901;check cm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1575330447] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575330519] SERVICE ALERT: ieatlms3901;check cm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1575330573] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575330697] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1575330773] SERVICE ALERT: ieatnetsimv9034-81;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.01%