[1560931726] HOST ALERT: localhost;DOWN;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_ping, ...) failed. errno is 2: No such file or directory
[1560931762] SERVICE ALERT: localhost;Current Load;CRITICAL;HARD;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_load, ...) failed. errno is 2: No such file or directory
[1560931786] HOST ALERT: localhost;DOWN;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_ping, ...) failed. errno is 2: No such file or directory
[1560931800] SERVICE ALERT: localhost;Current Users;CRITICAL;HARD;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_users, ...) failed. errno is 2: No such file or directory
[1560931837] SERVICE ALERT: localhost;HTTP;CRITICAL;HARD;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_http, ...) failed. errno is 2: No such file or directory
[1560931846] HOST ALERT: localhost;DOWN;SOFT;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_ping, ...) failed. errno is 2: No such file or directory
[1560931875] SERVICE ALERT: localhost;PING;CRITICAL;HARD;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_ping, ...) failed. errno is 2: No such file or directory
[1560931906] HOST ALERT: localhost;DOWN;SOFT;4;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_ping, ...) failed. errno is 2: No such file or directory
[1560931912] SERVICE ALERT: localhost;Root Partition;CRITICAL;HARD;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_disk, ...) failed. errno is 2: No such file or directory
[1560931954] HOST ALERT: localhost;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.12 ms
[1560932062] SERVICE ALERT: localhost;Current Load;OK;SOFT;1;OK - load average: 0.41, 0.30, 0.17
[1560932100] SERVICE ALERT: localhost;Current Users;OK;SOFT;1;USERS OK - 1 users currently logged in
[1560932137] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 5215 bytes in 0.003 second response time
[1560932179] SERVICE ALERT: localhost;PING;OK;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.05 ms
[1560932212] SERVICE ALERT: localhost;Root Partition;OK;SOFT;1;DISK OK - free space: / 42567 MB (83.18% inode=100%):
[1560932437] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 5215 bytes in 0.003 second response time
[1560932497] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 5215 bytes in 0.002 second response time
[1560932557] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5215 bytes in 0.002 second response time
[1560933800] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;1;(Service check timed out after 61.01 seconds)
[1560933809] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.38%
[1560933865] SERVICE ALERT: ieatlms3901;check snapshots;WARNING;SOFT;1;CRITICAL- ENM SNAPSHOTS ARE PRESENT
[1560933899] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 61.01 seconds)
[1560933900] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560933902] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560933903] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560933908] HOST ALERT: ieatrcxb3779;DOWN;SOFT;1;CRITICAL - Host Unreachable (141.137.208.10)
[1560933909] SERVICE ALERT: ieatlms3901;check unsync nodes;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560933911] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560933925] HOST ALERT: ieatrcxb3912;DOWN;SOFT;1;CRITICAL - Host Unreachable (141.137.208.16)
[1560933928] HOST ALERT: ieatrcxb3913;DOWN;SOFT;1;CRITICAL - Host Unreachable (141.137.208.17)
[1560933933] HOST ALERT: ieatrcxb3995;DOWN;SOFT;1;CRITICAL - Host Unreachable (141.137.208.20)
[1560933960] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.5%
[1560933980] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560933989] HOST ALERT: ieatrcxb3912;DOWN;SOFT;2;CRITICAL - Host Unreachable (141.137.208.16)
[1560933992] HOST ALERT: ieatrcxb3913;DOWN;SOFT;2;CRITICAL - Host Unreachable (141.137.208.17)
[1560933997] HOST ALERT: ieatrcxb3995;DOWN;SOFT;2;CRITICAL - Host Unreachable (141.137.208.20)
[1560933998] HOST ALERT: ieatrcxb3779;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1560934004] SERVICE ALERT: ieatlms3901;check snapshots;WARNING;SOFT;2;CRITICAL- ENM SNAPSHOTS ARE PRESENT
[1560934026] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1560934039] SERVICE ALERT: ieatwlvm5116;check workload profiles started;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1560934048] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_vcs_sg.sh, ...) failed. errno is 2: No such file or directory
[1560934054] HOST ALERT: ieatrcxb3912;DOWN;SOFT;3;CRITICAL - Host Unreachable (141.137.208.16)
[1560934057] HOST ALERT: ieatrcxb3913;DOWN;SOFT;3;CRITICAL - Host Unreachable (141.137.208.17)
[1560934062] HOST ALERT: ieatrcxb3779;DOWN;SOFT;3;CRITICAL - Host Unreachable (141.137.208.10)
[1560934080] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;2;(Service check timed out after 61.01 seconds)
[1560934081] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;SOFT;2;(Service check timed out after 61.01 seconds)
[1560934082] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934083] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934087] HOST ALERT: ieatrcxb3995;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1560934090] SERVICE ALERT: ieatlms3901;check unsync nodes;CRITICAL;SOFT;2;(Service check timed out after 60.31 seconds)
[1560934092] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;(Service check timed out after 61.01 seconds)
[1560934097] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 61.01 seconds)
[1560934098] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;(Service check timed out after 61.01 seconds)
[1560934099] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934101] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;SOFT;1;(Service check timed out after 60.02 seconds)
[1560934104] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;SOFT;1;(Service check timed out after 60.71 seconds)
[1560934106] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934110] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.74%
[1560934110] SERVICE ALERT: ieatnetsimv5116-24;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560934118] HOST ALERT: ieatrcxb3912;DOWN;SOFT;4;CRITICAL - Host Unreachable (141.137.208.16)
[1560934121] HOST ALERT: ieatrcxb3913;DOWN;SOFT;4;CRITICAL - Host Unreachable (141.137.208.17)
[1560934127] HOST ALERT: ieatrcxb3779;DOWN;SOFT;4;CRITICAL - Host Unreachable (141.137.208.10)
[1560934147] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1560934152] HOST ALERT: ieatrcxb3995;DOWN;SOFT;4;CRITICAL - Host Unreachable (141.137.208.20)
[1560934158] SERVICE ALERT: ieatlms3901;check snapshots;WARNING;HARD;3;CRITICAL- ENM SNAPSHOTS ARE PRESENT
[1560934160] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934167] SERVICE ALERT: ieatwlvm5116;check workload profiles started;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1560934168] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_vcs_sg.sh, ...) failed. errno is 2: No such file or directory
[1560934173] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1560934183] HOST ALERT: ieatrcxb3912;DOWN;SOFT;5;CRITICAL - Host Unreachable (141.137.208.16)
[1560934186] HOST ALERT: ieatrcxb3913;DOWN;SOFT;5;CRITICAL - Host Unreachable (141.137.208.17)
[1560934191] HOST ALERT: ieatrcxb3779;DOWN;SOFT;5;CRITICAL - Host Unreachable (141.137.208.10)
[1560934216] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934216] HOST ALERT: ieatrcxb3995;DOWN;SOFT;5;CRITICAL - Host Unreachable (141.137.208.20)
[1560934242] SERVICE ALERT: ieatnetsimv5116-24;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560934247] HOST ALERT: ieatrcxb3912;DOWN;SOFT;6;CRITICAL - Host Unreachable (141.137.208.16)
[1560934250] HOST ALERT: ieatrcxb3913;DOWN;SOFT;6;CRITICAL - Host Unreachable (141.137.208.17)
[1560934256] HOST ALERT: ieatrcxb3779;DOWN;SOFT;6;CRITICAL - Host Unreachable (141.137.208.10)
[1560934260] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934261] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 60.03 seconds)
[1560934262] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;(Service check timed out after 60.05 seconds)
[1560934263] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934267] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1560934270] SERVICE ALERT: ieatlms3901;check unsync nodes;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934272] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934277] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934278] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934279] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934281] HOST ALERT: ieatrcxb3995;DOWN;SOFT;6;CRITICAL - Host Unreachable (141.137.208.20)
[1560934281] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934284] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934286] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934288] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_vcs_sg.sh, ...) failed. errno is 2: No such file or directory
[1560934293] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934295] SERVICE ALERT: ieatwlvm5116;check workload profiles started;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1560934297] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934300] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 61.01 seconds)
[1560934300] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934301] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1560934312] HOST ALERT: ieatrcxb3912;DOWN;SOFT;7;CRITICAL - Host Unreachable (141.137.208.16)
[1560934315] HOST ALERT: ieatrcxb3913;DOWN;SOFT;7;CRITICAL - Host Unreachable (141.137.208.17)
[1560934320] HOST ALERT: ieatrcxb3779;DOWN;SOFT;7;CRITICAL - Host Unreachable (141.137.208.10)
[1560934345] HOST ALERT: ieatrcxb3995;DOWN;SOFT;7;CRITICAL - Host Unreachable (141.137.208.20)
[1560934376] HOST ALERT: ieatrcxb3912;DOWN;SOFT;8;CRITICAL - Host Unreachable (141.137.208.16)
[1560934379] HOST ALERT: ieatrcxb3913;DOWN;SOFT;8;CRITICAL - Host Unreachable (141.137.208.17)
[1560934385] HOST ALERT: ieatrcxb3779;DOWN;SOFT;8;CRITICAL - Host Unreachable (141.137.208.10)
[1560934396] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934410] HOST ALERT: ieatrcxb3995;DOWN;SOFT;8;CRITICAL - Host Unreachable (141.137.208.20)
[1560934429] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1560934441] HOST ALERT: ieatrcxb3912;DOWN;SOFT;9;CRITICAL - Host Unreachable (141.137.208.16)
[1560934444] HOST ALERT: ieatrcxb3913;DOWN;SOFT;9;CRITICAL - Host Unreachable (141.137.208.17)
[1560934449] HOST ALERT: ieatrcxb3779;DOWN;SOFT;9;CRITICAL - Host Unreachable (141.137.208.10)
[1560934457] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934458] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;(Service check timed out after 60.02 seconds)
[1560934459] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934461] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934465] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;HARD;3;(Service check timed out after 61.01 seconds)
[1560934466] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934473] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934474] HOST ALERT: ieatrcxb3995;DOWN;SOFT;9;CRITICAL - Host Unreachable (141.137.208.20)
[1560934477] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934480] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;SOFT;2;(Service check timed out after 60.01 seconds)
[1560934481] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;2;(Service check timed out after 61.01 seconds)
[1560934505] HOST ALERT: ieatrcxb3912;DOWN;HARD;10;CRITICAL - Host Unreachable (141.137.208.16)
[1560934508] HOST ALERT: ieatrcxb3913;DOWN;HARD;10;CRITICAL - Host Unreachable (141.137.208.17)
[1560934514] HOST ALERT: ieatrcxb3779;DOWN;HARD;10;CRITICAL - Host Unreachable (141.137.208.10)
[1560934539] HOST ALERT: ieatrcxb3995;DOWN;HARD;10;CRITICAL - Host Unreachable (141.137.208.20)
[1560934576] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934651] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (11015) IS HIGHER THAN 95% OF ENM COMECIM NODES (11144)
[1560934654] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;CRITICAL;HARD;3;(Service check timed out after 61.01 seconds)
[1560934657] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934660] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 60.01 seconds)
[1560934719] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560934844] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934846] SERVICE ALERT: ieatnetsimv5116-73;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934851] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560934852] SERVICE ALERT: ieatnetsimv5116-76;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 61.01 seconds)
[1560934853] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934855] SERVICE ALERT: ieatnetsimv5116-78;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934866] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934868] SERVICE ALERT: ieatnetsimv5116-85;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 60.01 seconds)
[1560934981] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560934999] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935001] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935002] SERVICE ALERT: ieatnetsimv5116-85;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935083] SERVICE ALERT: ieatnetsimv5116-76;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935124] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1560935129] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1560935135] SERVICE ALERT: ieatlms3901;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1560935139] SERVICE ALERT: ieatnetsimv5116-78;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935146] SERVICE ALERT: ieatnetsimv5116-73;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935161] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935185] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1560935253] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1560935255] SERVICE ALERT: ieatnetsimv5116-61;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.09%
[1560935319] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1560935323] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560935329] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560935341] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.92%
[1560935344] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560935357] SERVICE ALERT: ieatnetsimv5116-51;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560935358] SERVICE ALERT: ieatnetsimv5116-50;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560935361] SERVICE ALERT: ieatnetsimv5116-52;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560935406] SERVICE ALERT: ieatnetsimv5116-61;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 53.11%
[1560935423] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1560935456] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935462] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935468] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;OK;HARD;3;OK
[1560935477] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935486] SERVICE ALERT: ieatnetsimv5116-51;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560935491] SERVICE ALERT: ieatnetsimv5116-50;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560935492] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 15.37%
[1560935494] SERVICE ALERT: ieatnetsimv5116-52;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560935540] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1560935569] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1560935600] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (20429) IS BETWEEN 80% AND 95% OF ENM NODES (21970)
[1560935608] SERVICE ALERT: ieatlms3901;check number nodes;OK;HARD;3;OK- NUMBER OF NODES CONFIGURED IN ENM (21970) IS HIGHER THAN 95% OF NETSIM NODES (22495)
[1560935614] SERVICE ALERT: ieatnetsimv5116-51;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560935626] SERVICE ALERT: ieatnetsimv5116-52;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560935691] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1560935713] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.48%
[1560935754] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (19229) IS BETWEEN 80% AND 95% OF ENM NODES (21970)
[1560935755] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (21945) IS HIGHER THAN 95% OF ENM NODES (21970)
[1560935863] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.63%
[1560935907] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (10504) IS FROM 80% TO 95% OF ENM COMECIM NODES (11144)
[1560935924] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560935930] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560935931] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560935940] SERVICE ALERT: ieatnetsimv5116-43;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560935942] SERVICE ALERT: ieatnetsimv5116-44;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560936053] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560936061] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560936064] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560936069] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1560936070] SERVICE ALERT: ieatnetsimv5116-43;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560936071] SERVICE ALERT: ieatnetsimv5116-44;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560936089] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (10432) IS FROM 80% TO 95% OF ENM COMECIM NODES (11144)
[1560936131] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560936163] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.64%
[1560936182] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560936189] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1560936190] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560936197] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560936199] SERVICE ALERT: ieatnetsimv5116-43;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560936212] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560936224] SERVICE ALERT: ieatnetsimv5116-51;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560936226] SERVICE ALERT: ieatnetsimv5116-50;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560936241] SERVICE ALERT: ieatnetsimv5116-52;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560936267] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (10443) IS FROM 80% TO 95% OF ENM COMECIM NODES (11144)
[1560936310] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1560936310] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560936313] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 29.39%
[1560936330] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560936345] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560936359] SERVICE ALERT: ieatnetsimv5116-50;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560936373] SERVICE ALERT: ieatnetsimv5116-52;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560936462] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560936491] SERVICE ALERT: ieatnetsimv5116-50;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560936533] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560936536] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560936595] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560936666] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560936666] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560936800] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560936945] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (10717) IS HIGHER THAN 95% OF ENM COMECIM NODES (11144)
[1560937149] SERVICE ALERT: ieatnetsimv5116-47;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560937154] SERVICE ALERT: ieatnetsimv5116-48;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560937232] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 501.01 seconds)
[1560937278] SERVICE ALERT: ieatnetsimv5116-47;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560937286] SERVICE ALERT: ieatnetsimv5116-48;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560937844] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560937975] SERVICE ALERT: ieatnetsimv5116-60;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1560938007] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1560938126] SERVICE ALERT: ieatnetsimv5116-60;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 10.7%
[1560938194] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1560938357] SERVICE ALERT: ieatnetsimv5116-49;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560938382] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1560938464] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560938491] SERVICE ALERT: ieatnetsimv5116-49;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560938624] SERVICE ALERT: ieatnetsimv5116-49;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560938711] SERVICE FLAPPING ALERT: ieatnetsimv5116-93;check netsim disk;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1560938711] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;OK;HARD;3;OK
[1560938946] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560939076] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560939084] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 500.01 seconds)
[1560939100] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 500.01 seconds)
[1560939205] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560939242] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.14%
[1560939311] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1560939393] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 12.36%
[1560939432] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1560939553] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1560939655] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560939795] SERVICE FLAPPING ALERT: ieatlms3901;check pm supervision nodes;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1560939795] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (20429) IS BETWEEN 80% AND 95% OF ENM NODES (21970)
[1560939815] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560940755] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560940784] SERVICE ALERT: ieatnetsimv5116-63;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560940788] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560940793] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560940794] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560940819] SERVICE ALERT: ieatnetsimv5116-81;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560940922] SERVICE ALERT: ieatnetsimv5116-63;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560940924] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560940939] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560940941] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560940953] SERVICE ALERT: ieatnetsimv5116-81;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560941059] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560941086] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560941087] SERVICE ALERT: ieatnetsimv5116-81;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560941088] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560941376] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 500.63 seconds)
[1560941394] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560941396] SERVICE ALERT: ieatnetsimv5116-67;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560941402] SERVICE ALERT: ieatnetsimv5116-65;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560941411] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560941501] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.01%
[1560941551] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560941558] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560941568] SERVICE ALERT: ieatnetsimv5116-67;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560941626] SERVICE ALERT: ieatnetsimv5116-65;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560941651] SERVICE ALERT: ieatnetsimv5116-76;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.81%
[1560941676] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560941685] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560941700] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560941711] SERVICE ALERT: ieatnetsimv5116-67;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560941792] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (10500) IS FROM 80% TO 95% OF ENM COMECIM NODES (11144)
[1560941819] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560941997] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 501.01 seconds)
[1560942001] SERVICE ALERT: ieatnetsimv5116-69;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560942002] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560942010] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560942029] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (10412) IS FROM 80% TO 95% OF ENM COMECIM NODES (11144)
[1560942143] SERVICE ALERT: ieatnetsimv5116-69;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560942149] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560942157] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560942183] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (10295) IS FROM 80% TO 95% OF ENM COMECIM NODES (11144)
[1560942283] SERVICE ALERT: ieatnetsimv5116-69;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560942293] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560942299] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560942315] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560942318] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560942325] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560942331] SERVICE ALERT: ieatnetsimv5116-67;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560942443] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560942483] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560942486] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560942553] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560942553] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;OK;HARD;3;OK
[1560942589] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560942631] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560942631] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560942688] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560942732] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560942821] SERVICE ALERT: ieatnetsimv5116-79;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560942867] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.29%
[1560942875] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560942912] SERVICE ALERT: ieatnetsimv5116-69;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560942989] SERVICE ALERT: ieatnetsimv5116-65;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560943018] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.91%
[1560943056] SERVICE ALERT: ieatnetsimv5116-69;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560943070] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1560943070] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560943141] SERVICE ALERT: ieatnetsimv5116-65;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560943154] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1560943169] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 65.57%
[1560943198] SERVICE ALERT: ieatnetsimv5116-69;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560943249] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1560943255] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560943274] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1560943299] SERVICE ALERT: ieatnetsimv5116-65;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560943362] SERVICE ALERT: ieatnetsimv5116-61;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.52%
[1560943382] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.84%
[1560943395] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1560943402] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560943439] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1560943513] SERVICE ALERT: ieatnetsimv5116-61;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 13.26%
[1560943533] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 36.75%
[1560943550] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560943627] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1560943794] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560943857] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560943949] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560944026] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized comecim;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1560944026] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (10716) IS HIGHER THAN 95% OF ENM COMECIM NODES (11144)
[1560944100] SERVICE ALERT: ieatnetsimv5116-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560944101] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560944123] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560944170] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560944172] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560944315] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1560944315] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1560944328] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1560944400] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.82%
[1560944422] SERVICE ALERT: ieatnetsimv5116-69;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560944461] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1560944551] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 43.27%
[1560944612] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1560944702] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.33%
[1560944765] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1560944790] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 500.01 seconds)
[1560944852] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 33.11%
[1560945394] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560945410] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 500.01 seconds)
[1560945582] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1560945770] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1560945790] SERVICE ALERT: ieatnetsimv5116-65;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1560947343] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.42%
[1560947494] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 41.8%
[1560948357] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1560948546] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1560948604] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.54%
[1560948696] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 501.01 seconds)
[1560948709] SERVICE ALERT: ieatlms3901;check number nodes;CRITICAL;SOFT;1;(Service check timed out after 501.01 seconds)
[1560948726] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560948733] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1560948748] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;(Service check timed out after 501.01 seconds)
[1560948754] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.32%
[1560948855] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560948855] SERVICE ALERT: ieatlms3901;check fm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 500.01 seconds)
[1560948905] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 46.93%
[1560948920] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (11140) IS HIGHER THAN 95% OF ENM COMECIM NODES (11144)
[1560948928] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (7108) IS HIGHER THAN 95% OF ENM CPP NODES (7110)
[1560949147] SERVICE ALERT: ieatlms3901;check number nodes;OK;SOFT;2;OK- NUMBER OF NODES CONFIGURED IN ENM (21970) IS HIGHER THAN 95% OF NETSIM NODES (22495)
[1560949279] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (21945) IS HIGHER THAN 95% OF ENM NODES (21970)
[1560949396] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;OK;HARD;3;OK
[1560949404] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (20429) IS BETWEEN 80% AND 95% OF ENM NODES (21970)
[1560949420] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1560949632] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (19229) IS BETWEEN 80% AND 95% OF ENM NODES (21970)
[1560950105] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560950144] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.28%
[1560950294] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 58.22%
[1560950597] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1560950718] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1560950838] SERVICE ALERT: ieatnetsimv5116-93;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1560951048] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.71%
[1560951199] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 70.1%
[1560951806] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 501.01 seconds)
[1560951931] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.17%
[1560952082] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.62%
[1560952424] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560952661] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1560952812] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 27.95%
[1560953588] SERVICE ALERT: ieatnetsimv5116-36;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.61%
[1560953644] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560953739] SERVICE ALERT: ieatnetsimv5116-36;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 23.6%
[1560953749] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.92%
[1560953900] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 64.36%
[1560954227] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560954512] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized comecim;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1560954650] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.2%
[1560954801] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 67.45%
[1560955448] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.95 seconds)
[1560956022] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560957242] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.37 seconds)
[1560957353] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.75%
[1560957503] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 55.47%
[1560957844] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.17%
[1560957863] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 501.01 seconds)
[1560957995] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 16.64%
[1560958255] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.34%
[1560958406] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 63.51%
[1560958434] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560959346] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.42%
[1560959496] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.57%
[1560959647] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 59.27%
[1560960134] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.02 seconds)
[1560960754] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 500.01 seconds)
[1560960921] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560961012] SERVICE FLAPPING ALERT: ieatlms3901;check pm supervision nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1560961357] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560961638] SERVICE FLAPPING ALERT: ieatnetsimv5116-93;check netsim disk;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1560962458] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 501.01 seconds)
[1560962930] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1560963066] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6959) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560963127] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1560963192] SERVICE ALERT: ieatnetsimv5116-19;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.32%
[1560963324] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1560963343] SERVICE ALERT: ieatnetsimv5116-19;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 62.89%
[1560963686] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560964215] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.18%
[1560964301] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560964366] SERVICE ALERT: ieatnetsimv5116-20;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.11%
[1560964558] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1560964709] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 54.29%
[1560964911] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1560965097] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.45%
[1560965248] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 52.43%
[1560965462] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.53%
[1560965612] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.18%
[1560967001] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1560969721] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560970307] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560971528] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 501.01 seconds)
[1560972119] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560973339] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560973933] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560975877] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.97%
[1560976028] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 8.43%
[1560976953] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560977185] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1560977367] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1560977380] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1560977431] SERVICE ALERT: ieatnetsimv5116-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.18%
[1560977548] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1560977553] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560977562] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1560977583] SERVICE ALERT: ieatnetsimv5116-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.25%
[1560977743] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1560978173] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560978779] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560978934] SERVICE ALERT: ieatnetsimv5116-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.31%
[1560978966] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.24%
[1560979085] SERVICE ALERT: ieatnetsimv5116-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.43%
[1560979117] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.13%
[1560979268] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 38.27%
[1560980000] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.08 seconds)
[1560980621] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 501.01 seconds)
[1560981242] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 501.01 seconds)
[1560982339] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560983439] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 500.01 seconds)
[1560984060] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 501.01 seconds)
[1560984662] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7109) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5594)
[1560984699] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.49%
[1560984850] SERVICE ALERT: ieatnetsimv5116-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 64.39%