[1568768508] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.83%
[1568768658] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.33%
[1568772427] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2332) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568772559] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2438) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1568788446] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4018) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1568788585] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4011) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1568788720] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3992) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1568790536] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4223) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1568793290] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1568793472] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1568793653] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1568794340] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1568794520] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1568794701] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1568794913] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1568795389] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1568796407] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.73%
[1568796558] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 54.83%
[1568797313] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1568798362] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1568799808] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1489) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1568799997] SERVICE ALERT: ieatnetsimv5116-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.8%
[1568800147] SERVICE ALERT: ieatnetsimv5116-33;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 15.49%
[1568806374] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1568806555] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1568806736] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1568807424] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1568807604] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1568807785] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1568810382] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1974) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1568810535] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2180) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568810706] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2402) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1568810789] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.89%
[1568810792] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568810793] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.71%
[1568810827] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.83%
[1568810838] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568810846] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568810850] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568810879] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568810939] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 20.88%
[1568810942] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 52.16%
[1568810944] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.28%
[1568810978] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 4.67%
[1568810988] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.55%
[1568810999] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568811002] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568811029] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.64%
[1568811152] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1568811154] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1568811163] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568811315] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568811465] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.91%
[1568812568] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2344) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568812572] SERVICE ALERT: ieatnetsimv5116-100;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.23%
[1568812706] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2341) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568812722] SERVICE ALERT: ieatnetsimv5116-100;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.18%
[1568812842] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2365) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568812873] SERVICE ALERT: ieatnetsimv5116-100;check netsim cpu;OK;HARD;3;OK- CPU IS USED 36.37%
[1568814695] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2499) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1568814782] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 17.69%
[1568814785] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 14.95%
[1568815045] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1568815045] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1568815796] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1568815796] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1568816617] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568816638] SERVICE ALERT: ieatnetsimv5116-27;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568816679] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568816731] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568816740] SERVICE ALERT: ieatnetsimv5116-37;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568816760] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568816777] SERVICE ALERT: ieatnetsimv5116-27;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568816807] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13276) IS FROM 80% TO 95% OF ENM COMECIM NODES (14057)
[1568816814] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568816864] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568816881] SERVICE ALERT: ieatnetsimv5116-37;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568816899] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568816940] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568816949] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568816952] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568816964] SERVICE ALERT: ieatnetsimv5116-56;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568816984] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568816997] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12203) IS FROM 80% TO 95% OF ENM COMECIM NODES (14057)
[1568817001] SERVICE ALERT: ieatnetsimv5116-41;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817039] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817059] SERVICE ALERT: ieatnetsimv5116-37;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817072] SERVICE ALERT: ieatnetsimv5116-82;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817076] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817088] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817111] SERVICE ALERT: ieatnetsimv5116-56;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817120] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817123] SERVICE ALERT: ieatnetsimv5116-87;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817141] SERVICE ALERT: ieatnetsimv5116-41;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817156] SERVICE ALERT: ieatnetsimv5116-69;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817170] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817204] SERVICE ALERT: ieatnetsimv5116-82;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817212] SERVICE ALERT: ieatnetsimv5116-37;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817212] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817222] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817222] SERVICE ALERT: ieatlms3901;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (10425) IS LOWER THAN 80% OF ENM COMECIM NODES (14057)
[1568817249] SERVICE ALERT: ieatnetsimv5116-56;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817251] SERVICE ALERT: ieatnetsimv5116-75;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817265] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817278] SERVICE ALERT: ieatnetsimv5116-41;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817280] SERVICE ALERT: ieatnetsimv5116-87;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817293] SERVICE ALERT: ieatnetsimv5116-69;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817315] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817336] SERVICE ALERT: ieatnetsimv5116-89;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817340] SERVICE ALERT: ieatnetsimv5116-82;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568817348] SERVICE ALERT: ieatnetsimv5116-72;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817352] SERVICE ALERT: ieatnetsimv5116-37;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568817353] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817372] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817384] SERVICE ALERT: ieatnetsimv5116-75;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817394] SERVICE ALERT: ieatnetsimv5116-50;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817401] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817449] SERVICE ALERT: ieatnetsimv5116-15;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817449] SERVICE ALERT: ieatnetsimv5116-69;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568817467] SERVICE ALERT: ieatnetsimv5116-92;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817467] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817472] SERVICE ALERT: ieatnetsimv5116-54;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817473] SERVICE ALERT: ieatnetsimv5116-89;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817473] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817485] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817485] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817487] SERVICE ALERT: ieatnetsimv5116-86;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817504] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817509] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817509] SERVICE ALERT: ieatnetsimv5116-63;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817516] SERVICE ALERT: ieatnetsimv5116-75;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817517] SERVICE ALERT: ieatnetsimv5116-44;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817523] SERVICE ALERT: ieatnetsimv5116-50;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817538] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817538] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817549] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817581] SERVICE ALERT: ieatnetsimv5116-15;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817598] SERVICE ALERT: ieatnetsimv5116-92;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817599] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817601] SERVICE ALERT: ieatnetsimv5116-54;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817613] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817620] SERVICE ALERT: ieatnetsimv5116-86;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817621] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817639] SERVICE ALERT: ieatnetsimv5116-55;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817640] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817645] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817645] SERVICE ALERT: ieatnetsimv5116-63;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817649] SERVICE ALERT: ieatnetsimv5116-44;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817651] SERVICE ALERT: ieatnetsimv5116-50;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817656] SERVICE ALERT: ieatnetsimv5116-48;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817672] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817682] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817719] SERVICE ALERT: ieatnetsimv5116-22;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817732] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817735] SERVICE ALERT: ieatnetsimv5116-92;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817736] SERVICE ALERT: ieatnetsimv5116-57;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817737] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817753] SERVICE ALERT: ieatnetsimv5116-61;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568817757] SERVICE ALERT: ieatnetsimv5116-86;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568817764] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817768] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817781] SERVICE ALERT: ieatnetsimv5116-55;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817787] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817787] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817787] SERVICE ALERT: ieatnetsimv5116-44;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817788] SERVICE ALERT: ieatnetsimv5116-63;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817789] SERVICE ALERT: ieatnetsimv5116-48;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817813] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817815] SERVICE ALERT: ieatnetsimv5116-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568817854] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817858] SERVICE ALERT: ieatnetsimv5116-22;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817868] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817870] SERVICE ALERT: ieatnetsimv5116-57;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568817877] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817920] SERVICE ALERT: ieatnetsimv5116-48;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817921] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817926] SERVICE ALERT: ieatnetsimv5116-55;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817931] SERVICE ALERT: ieatnetsimv5116-63;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817932] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568817944] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817950] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568817983] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568817999] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1568818002] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.07%
[1568818004] SERVICE ALERT: ieatnetsimv5116-36;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.88%
[1568818012] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.33%
[1568818013] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568818017] SERVICE ALERT: ieatnetsimv5116-41;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.24%
[1568818055] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568818065] SERVICE ALERT: ieatnetsimv5116-63;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568818083] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568818098] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568818112] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568818123] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818149] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 48.56%
[1568818153] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.19%
[1568818154] SERVICE ALERT: ieatnetsimv5116-36;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 39.18%
[1568818161] SERVICE ALERT: ieatnetsimv5116-60;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818163] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.92%
[1568818168] SERVICE ALERT: ieatnetsimv5116-41;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.8%
[1568818184] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818253] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568818293] SERVICE ALERT: ieatnetsimv5116-60;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568818313] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568818314] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;OK;HARD;3;OK- CPU IS USED 23.67%
[1568818355] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818384] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568818422] SERVICE ALERT: ieatnetsimv5116-60;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568818442] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568818447] SERVICE ALERT: ieatnetsimv5116-73;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818486] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568818576] SERVICE ALERT: ieatnetsimv5116-73;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568818618] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568818671] SERVICE ALERT: ieatnetsimv5116-32;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818705] SERVICE ALERT: ieatnetsimv5116-73;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568818763] SERVICE ALERT: ieatnetsimv5116-80;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818798] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818800] SERVICE ALERT: ieatnetsimv5116-32;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568818841] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818855] SERVICE ALERT: ieatnetsimv5116-58;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818892] SERVICE ALERT: ieatnetsimv5116-80;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568818928] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568818931] SERVICE ALERT: ieatnetsimv5116-32;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568818971] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568818983] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568818988] SERVICE ALERT: ieatnetsimv5116-58;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568819024] SERVICE ALERT: ieatnetsimv5116-80;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568819055] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11931) IS FROM 80% TO 95% OF ENM COMECIM NODES (14057)
[1568819058] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568819063] SERVICE ALERT: ieatnetsimv5116-56;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568819107] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568819118] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568819121] SERVICE ALERT: ieatnetsimv5116-58;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568819180] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568819310] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568819352] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568819394] SERVICE ALERT: ieatnetsimv5116-26;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568819440] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568819538] SERVICE ALERT: ieatnetsimv5116-26;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568819748] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568820910] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13997) IS HIGHER THAN 95% OF ENM COMECIM NODES (14057)
[1568823033] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568823227] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568823398] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568824785] SERVICE ALERT: ieatnetsimv5116-95;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568824935] SERVICE ALERT: ieatnetsimv5116-95;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568825352] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.57%
[1568825503] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.76%
[1568825654] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.07%
[1568825804] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.91%
[1568827035] SERVICE ALERT: ieatnetsimv5116-93;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568827072] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568827198] SERVICE ALERT: ieatnetsimv5116-93;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568827254] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568833127] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1568833248] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1568833368] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1568834651] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568834830] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568835027] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568835190] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568842857] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1568843905] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1568844438] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568844604] SERVICE ALERT: ieatnetsimv5116-94;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES