[1563750137] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1563750393] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1563750610] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563750735] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563750901] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (121) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563751658] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.3%
[1563751728] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2033) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1563751811] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.4%
[1563751863] SERVICE ALERT: ieatnetsimv5116-97;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1563751961] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 91.19%
[1563752014] SERVICE ALERT: ieatnetsimv5116-97;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 15.96%
[1563752212] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563752964] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563753547] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.18%
[1563753575] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2033) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563753698] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.9%
[1563753707] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2033) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563753839] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2033) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563754535] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (113) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563754660] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (113) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563754785] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (115) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563754887] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2325) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563755442] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563755592] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;HARD;3;OK- CPU IS USED 74.74%
[1563758430] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (112) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563759066] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563759190] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563759328] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (462) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1563759329] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563760127] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1563761084] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563761984] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563762035] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (109) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563762099] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2157) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563762197] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (109) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563762225] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2157) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563762322] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (109) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563762391] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2157) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563762446] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (109) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563762570] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (109) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563762949] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2033) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563763470] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (1412) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563763478] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (1666) IS LOWER THAN 80% OF ENM CPP NODES (4254)
[1563763595] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (1427) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563763653] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4193) IS HIGHER THAN 95% OF ENM CPP NODES (4104)
[1563763721] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1464) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563763772] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563763904] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563764109] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized nodes com/ecim;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1563764109] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563764943] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2033) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1563765997] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2394) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563766161] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2394) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563766214] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (117) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563766297] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2335) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563766423] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563766570] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563766572] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563766754] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2011) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563766889] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2013) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563767060] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2013) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563768936] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1999) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1563769838] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563769963] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563770087] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563770175] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563770316] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563770440] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical DSC;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1563770440] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563771851] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.59%
[1563772003] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.64%
[1563772153] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.25%
[1563773704] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.31 seconds)
[1563774084] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563774562] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563775482] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563775784] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;HARD;3;OK- CPU IS USED 65.63%
[1563776327] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1563776379] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563777299] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563777330] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (115) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563777689] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563777814] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563777938] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563778039] SERVICE ALERT: ieatnetsimv7035-24;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1563778164] SERVICE ALERT: ieatnetsimv7035-24;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1563778208] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563778293] SERVICE ALERT: ieatnetsimv7035-24;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1563778422] SERVICE ALERT: ieatnetsimv7035-24;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1563778550] SERVICE ALERT: ieatnetsimv7035-24;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1563778638] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1563778956] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563779112] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1563779815] SERVICE ALERT: ieatnetsimv7035-18;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1563779944] SERVICE ALERT: ieatnetsimv7035-18;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1563780074] SERVICE ALERT: ieatnetsimv7035-18;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1563780185] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1563780335] SERVICE ALERT: ieatnetsimv7035-18;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 32.81%
[1563780361] SERVICE ALERT: ieatnetsimv7035-24;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1563780787] SERVICE ALERT: ieatnetsimv7035-24;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.6%
[1563780927] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.59%
[1563781014] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563781077] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.09%
[1563781228] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.88%
[1563781437] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM COMECIM NODES (14035)
[1563781593] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563781736] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1563781744] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563781864] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1563781884] SERVICE ALERT: ieatnetsimv7035-18;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1563782002] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1563782028] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;SOFT;2;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563782357] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1563782382] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563782728] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563782805] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;SOFT;2;OK
[1563783056] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM COMECIM NODES (0)
[1563783513] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (1999) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1563783649] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1563783812] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1563784569] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1563784894] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1563785580] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563785620] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1563785748] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1563785875] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1563785875] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1563786160] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563786334] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563787260] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1563787475] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1999) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1563788034] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (1537) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1563788882] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM COMECIM NODES (14117)
[1563788882] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563789531] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM COMECIM NODES (0)
[1563790060] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1563790071] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563790239] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563790446] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1563791512] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1563791707] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1563791930] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1563792396] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563792561] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;SOFT;2;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563793083] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1563793266] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1563793901] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563794220] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;SOFT;2;OK
[1563795262] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1999) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1563795527] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM COMECIM NODES (2)
[1563795787] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563795860] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1563795986] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (18443) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21429)
[1563796108] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;2;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM COMECIM NODES (14069)
[1563796381] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1563796433] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563796625] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;SOFT;2;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563796780] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563796882] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563796919] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563797029] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1563797118] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1563797200] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (1999) IS HIGHER THAN 95% OF ENM SNMP NODES (750)
[1563797200] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563797248] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1563797378] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1563797567] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1563797802] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1563798610] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6912) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563798735] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563798856] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13425) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1563798924] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1563798954] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (118) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563799034] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2333) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563799078] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (120) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563799112] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1563799132] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2001) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563799160] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563799203] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563799263] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2001) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563799283] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1563799290] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2398) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563799328] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563799411] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2001) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563799453] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563799929] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13251) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563800070] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13251) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563800153] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (1426) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563800211] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13251) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563800279] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (1426) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563800350] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563800404] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563800475] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563800600] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563800680] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (14123) IS HIGHER THAN 95% OF ENM COMECIM NODES (14149)
[1563801163] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563801171] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563801287] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563801296] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563801411] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563801428] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563801545] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563801671] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563801796] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563802423] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.71%
[1563802573] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.53%
[1563803061] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (109) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563804020] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (1368) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563804145] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (1380) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563804271] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1416) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563806121] SERVICE ALERT: ieatnetsimv5116-73;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.31%
[1563806272] SERVICE ALERT: ieatnetsimv5116-73;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.62%
[1563806423] SERVICE ALERT: ieatnetsimv5116-73;check netsim cpu;OK;HARD;3;OK- CPU IS USED 14.35%
[1563806665] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563807435] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14035) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563807876] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1466) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563813870] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (117) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563814325] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1563814364] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1563814446] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1563814485] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1563814567] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1563814605] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1563815368] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1563815378] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1563815484] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563815489] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1563815498] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1563815564] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1563815605] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1563815610] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1563815619] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1563815685] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1563815726] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1563815732] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1563815806] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1563815853] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1563815974] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1563816043] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1563816165] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1563816285] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1563816468] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1563817475] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563817484] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1563818227] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1563818256] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13097) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563818398] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13097) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563818539] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (10748) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563819277] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1563819396] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1563819756] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1563820723] SERVICE FLAPPING ALERT: ieatnetsimv7035-15;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563820906] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563821096] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563821192] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1563821268] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1563821891] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1563822164] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1563822164] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14031) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563823587] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1563824426] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.74%
[1563824512] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2307) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1563824576] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 38.87%
[1563825785] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13013) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563825881] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (1405) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563825927] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13113) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563826007] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (1405) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563826067] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13113) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563826132] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1405) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563826658] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563826843] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1563827028] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1563827028] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1563827291] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1563827419] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1563827550] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1563828224] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1999) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1563829690] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14035) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563829737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RBS;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1563829737] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1563831810] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563831901] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563832599] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (7152) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563835505] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (111) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)