[1574553759] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1574553759] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574553899] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574553912] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12646) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574554025] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574554059] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12646) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574554151] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574554207] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12646) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574554352] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1723) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574555148] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574555175] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1574555211] HOST ALERT: ieatnetsimv9034-78;DOWN;SOFT;1;PING NOK
[1574555234] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (748) IS HIGHER THAN 95% OF ENM  NODES (750)
[1574555262] HOST ALERT: ieatnetsimv9034-78;UP;SOFT;1;PING OK
[1574555621] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (750) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (750)
[1574556350] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574557162] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2333) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574557171] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (134) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574557291] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2522) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574557297] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (148) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574557756] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical DSC;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1574557756] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574557835] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13939) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574558962] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1574558962] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574559228] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (430) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574559355] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (430) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574559482] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574563829] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574563952] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574564014] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574564076] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574564139] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574564262] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574567627] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574567680] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574567753] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574567813] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574567866] SERVICE FLAPPING ALERT: ieatlms7435;check pm collection statistical DSC;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1574567866] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574567877] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574567999] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574568184] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574568219] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2309) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574568224] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (131) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574568348] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2317) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574568350] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (133) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574568477] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2317) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574568478] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (133) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574568562] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574568663] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12544) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574568698] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574568812] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12544) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574568851] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574568960] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12527) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574569767] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574569893] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574570018] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574570021] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574570198] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574570203] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574570328] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574570384] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574570456] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574570588] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574570808] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574570986] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574571165] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574571480] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574571605] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574571731] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574571856] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574571990] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574572246] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574572428] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574572588] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13386) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574572609] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574573239] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (609) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574573365] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574573490] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (632) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574573614] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574573739] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (663) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574573875] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574574004] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (676) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574574132] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (688) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574574258] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (691) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574574320] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  RBS: 1 ERBS: 3 MGW: 1
[1574574449] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574576067] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (577) IS LOWER THAN 80% OF ENM  NODES (750)
[1574576093] SERVICE FLAPPING ALERT: ieatlms5735-1;check system service hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1574576215] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12253) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574576373] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13009) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574576520] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13011) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574576784] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (88) IS FROM 80% TO 95% OF ENM  NODES (98)
[1574576913] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (97) IS HIGHER THAN 95% OF ENM  NODES (98)
[1574577762] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1549) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574577877] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (718) IS HIGHER THAN 95% OF ENM  NODES (750)
[1574577888] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1577) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574578025] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1574579285] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2557) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574579287] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574579663] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574580191] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-BC-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (12)
[1574580317] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-BC-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1574580824] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574580949] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574581075] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574582895] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2369) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574583025] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2369) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574583157] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2369) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574583395] SERVICE ALERT: ieatlms5735-1;check synchronized FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (79) IS LOWER THAN 80% OF ENM  NODES (100)
[1574583405] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574583522] SERVICE ALERT: ieatlms5735-1;check synchronized FRONTHAUL-6080;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (100)!
[1574583542] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574583679] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574583747] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13488) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574584055] HOST ALERT: ieatnetsimv7035-37;DOWN;SOFT;1;PING NOK
[1574584055] HOST ALERT: ieatnetsimv9041-03;DOWN;SOFT;1;PING NOK
[1574584117] HOST ALERT: ieatnetsimv9041-03;UP;SOFT;1;PING OK
[1574584117] HOST ALERT: ieatnetsimv7035-37;UP;SOFT;1;PING OK
[1574584355] HOST ALERT: ieatnetsimv7035-84;DOWN;SOFT;1;PING NOK
[1574584417] HOST ALERT: ieatnetsimv7035-84;UP;SOFT;1;PING OK
[1574584680] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574584695] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574584954] HOST ALERT: ieatnetsimv7035-74;DOWN;SOFT;1;PING NOK
[1574584955] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;1;PING NOK
[1574584955] HOST ALERT: ieatnetsimv9041-24;DOWN;SOFT;1;PING NOK
[1574584955] HOST ALERT: ieatnetsimv7035-41;DOWN;SOFT;1;PING NOK
[1574584996] HOST ALERT: ieatnetsimv9041-01;UP;SOFT;1;PING OK
[1574585016] HOST ALERT: ieatnetsimv7035-74;UP;SOFT;1;PING OK
[1574585017] HOST ALERT: ieatnetsimv9041-24;UP;SOFT;1;PING OK
[1574585017] HOST ALERT: ieatnetsimv7035-41;UP;SOFT;1;PING OK
[1574585488] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574585666] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574586250] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574586494] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (132) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574586620] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (132) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574586761] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (132) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574586878] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574587005] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574587042] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1459) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574587137] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574587167] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1497) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574587294] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1574587377] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12587) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574587525] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12587) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574587673] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12587) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574588007] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574588135] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MTAS NODES !
[1574588286] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574588348] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574588412] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574588472] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574588486] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1681) IS BETWEEN 5% AND 20% OF ENM NODES (13969)
[1574588502] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (424) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574588538] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574588597] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (11) IS HIGHER THAN 20% OF ENM NODES (50)
[1574588620] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1412) IS BETWEEN 5% AND 20% OF ENM NODES (14006)
[1574588627] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (424) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574588748] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1504) IS BETWEEN 5% AND 20% OF ENM NODES (13969)
[1574588753] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (18) IS LOWER THAN 5% OF ENM NODES (1200)
[1574588870] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574589316] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1574589440] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (11) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1574589570] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1574590571] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (1750) IS LOWER THAN 80% OF ENM NODES (3500)
[1574590699] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1750) IS HIGHER THAN 95% OF ENM NODES (1750)
[1574590744] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574592157] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574592202] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1574592359] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (8) IS LOWER THAN 5% OF ENM NODES (13969)
[1574593970] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2557) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574593971] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574594083] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (561) IS LOWER THAN 80% OF ENM  NODES (750)
[1574594209] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (571) IS LOWER THAN 80% OF ENM  NODES (750)
[1574594344] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574594350] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574594476] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574594499] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1345) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574594602] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574594630] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1574596484] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574596607] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574596731] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574597984] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574598108] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1574599409] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (16) IS HIGHER THAN 20% OF ENM NODES (50)
[1574599556] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (23) IS HIGHER THAN 20% OF ENM NODES (50)
[1574599561] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (354) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574599682] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (18) IS HIGHER THAN 20% OF ENM NODES (50)
[1574599693] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (138) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574599829] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (199) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574600334] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574600374] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1574600499] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1574600622] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (80)
[1574600810] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574601179] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2158) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574601189] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (123) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574601310] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2228) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574601317] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (128) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574601438] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2286) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574601444] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (132) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574602964] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574603090] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574603167] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (757) IS BETWEEN 5% AND 20% OF ENM NODES (13969)
[1574603215] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (28) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574603299] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574603300] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (30) IS LOWER THAN 5% OF ENM NODES (13969)
[1574603406] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574603438] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574603438] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574603535] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574603663] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574604315] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1574606771] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574606897] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574606902] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1574607022] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574607273] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (5) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (5)
[1574607425] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13240) IS FROM 80% TO 95% OF ENM  NODES (14006)
[1574607559] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13938) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574609597] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574609723] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574609848] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574610421] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574610515] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574610548] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574610671] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1574610674] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574610678] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (250) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574610797] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574610802] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (255) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574610927] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (1200)
[1574611292] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12872) IS FROM 80% TO 95% OF ENM  NODES (14006)
[1574611426] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12775) IS FROM 80% TO 95% OF ENM  NODES (14006)
[1574611561] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13938) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574612267] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2526) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574612272] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (149) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574612351] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574612428] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574612554] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574612607] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (28) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574612679] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574612802] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574612901] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13760) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574612926] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574613380] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13152) IS FROM 80% TO 95% OF ENM  NODES (14006)
[1574613458] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (24)
[1574613513] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13115) IS FROM 80% TO 95% OF ENM  NODES (14006)
[1574613650] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13938) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574614401] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1574614482] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (2) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574614611] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (2) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574614741] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (2) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574615464] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13182) IS FROM 80% TO 95% OF ENM  NODES (14006)
[1574615598] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13938) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574616375] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574616496] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574616616] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574617018] HOST ALERT: ieatnetsimv9034-96;DOWN;SOFT;1;PING NOK
[1574617080] HOST ALERT: ieatnetsimv9034-96;UP;SOFT;1;PING OK
[1574617880] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574618269] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574618338] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574618350] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RNC;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1574618350] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (4) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574618393] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574618462] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;OK;SOFT;2;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574618516] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (28) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574621486] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574621605] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574621612] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574621733] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574621738] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574621857] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574623082] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2288) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574623084] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (133) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574623212] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2334) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574623214] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574623341] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2334) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574623343] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574623420] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574623546] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574623672] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574623731] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12728) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574623879] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12571) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574624025] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12571) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574624264] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574624389] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574624515] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574625464] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (50)
[1574625566] SERVICE ALERT: ieatlms3901;check pm supervision nodes Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1200)
[1574625574] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (4) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (15)
[1574625692] SERVICE ALERT: ieatlms3901;check pm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1574625702] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (2) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574625721] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574625831] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (2) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574625845] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574625968] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574626949] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2557) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574626951] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574627278] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574628121] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1574628121] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (24)
[1574629440] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (5) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (5)
[1574629573] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574630458] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574630613] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.57%
[1574630763] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 44.31%
[1574631934] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10088) IS LOWER THAN 80% OF ENM  NODES (13969)
[1574632066] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10483) IS LOWER THAN 80% OF ENM  NODES (14006)
[1574632185] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1 MTAS: 1
[1574632204] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (11411) IS FROM 80% TO 95% OF ENM  NODES (14006)
[1574632314] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574632543] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574632600] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574632728] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574633025] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (13) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1574633150] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1574633752] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574633881] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634008] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634020] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13931) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574634158] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634163] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634285] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634291] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634410] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574634411] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634420] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634484] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634532] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634537] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574634610] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634659] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634662] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574634711] HOST ALERT: ieatnetsimv9034-63;DOWN;SOFT;1;PING NOK
[1574634712] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634736] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634737] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '10.210.217.10' (RSA) to the list of known hosts.
[1574634773] HOST ALERT: ieatnetsimv9034-63;UP;SOFT;1;PING OK
[1574634787] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-Indoor;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634842] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634852] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634876] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '10.210.217.10' (RSA) to the list of known hosts.
[1574634970] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574634979] SERVICE ALERT: ieatlms3901;check pm collection statistical CISCO-ASR900;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635015] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '10.210.217.10' (RSA) to the list of known hosts.
[1574635106] SERVICE ALERT: ieatlms3901;check pm collection statistical CISCO-ASR900;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635161] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical DSC;STARTED; Service appears to have started flapping (45.7% change >= 20.0% threshold)
[1574635161] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (46.0% change >= 20.0% threshold)
[1574635161] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (46.1% change >= 20.0% threshold)
[1574635161] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1574635161] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (41.7% change >= 20.0% threshold)
[1574635161] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RNC;STARTED; Service appears to have started flapping (33.9% change >= 20.0% threshold)
[1574635161] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (38.6% change >= 20.0% threshold)
[1574635161] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STARTED; Service appears to have started flapping (48.0% change >= 20.0% threshold)
[1574635233] SERVICE ALERT: ieatlms3901;check pm collection statistical CISCO-ASR900;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635327] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635331] HOST ALERT: ieatnetsimv5116-51;DOWN;SOFT;1;PING NOK
[1574635394] HOST ALERT: ieatnetsimv5116-51;UP;SOFT;1;PING OK
[1574635455] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635581] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635620] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-IS;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635631] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574635746] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-IS;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635766] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574635773] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574635791] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635803] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical DSC;STARTED; Service appears to have started flapping (45.7% change >= 20.0% threshold)
[1574635803] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (46.0% change >= 20.0% threshold)
[1574635803] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (46.1% change >= 20.0% threshold)
[1574635803] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (28.9% change >= 20.0% threshold)
[1574635803] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (41.7% change >= 20.0% threshold)
[1574635803] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RNC;STARTED; Service appears to have started flapping (33.9% change >= 20.0% threshold)
[1574635803] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (38.6% change >= 20.0% threshold)
[1574635803] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STARTED; Service appears to have started flapping (48.0% change >= 20.0% threshold)
[1574635876] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-IS;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635877] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-BSP;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574635891] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574635915] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574635919] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636002] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-BSP;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636016] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574636048] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636128] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-BC-BSP;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636149] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636275] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636400] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636514] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical DSC;STARTED; Service appears to have started flapping (45.7% change >= 20.0% threshold)
[1574636514] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (46.0% change >= 20.0% threshold)
[1574636514] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (46.1% change >= 20.0% threshold)
[1574636514] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (28.9% change >= 20.0% threshold)
[1574636514] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (46.9% change >= 20.0% threshold)
[1574636514] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RNC;STARTED; Service appears to have started flapping (33.9% change >= 20.0% threshold)
[1574636514] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (38.6% change >= 20.0% threshold)
[1574636514] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STARTED; Service appears to have started flapping (46.8% change >= 20.0% threshold)
[1574636586] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '10.210.245.3' (RSA) to the list of known hosts.
[1574636650] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636710] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '10.210.245.3' (RSA) to the list of known hosts.
[1574636777] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574636795] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636822] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636834] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '10.210.245.3' (RSA) to the list of known hosts.
[1574636900] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574636924] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636944] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574636949] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574636999] SERVICE ALERT: ieatlms7435;check pm collection statistical EPG;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '10.210.245.3' (RSA) to the list of known hosts.
[1574637024] SERVICE FLAPPING ALERT: ieatlms7435;check pm collection statistical DSC;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574637024] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574637071] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574637078] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574637123] SERVICE ALERT: ieatlms7435;check pm collection statistical EPG;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '10.210.245.3' (RSA) to the list of known hosts.
[1574637197] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1574637247] SERVICE ALERT: ieatlms7435;check pm collection statistical EPG;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '10.210.245.3' (RSA) to the list of known hosts.
[1574637253] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574637255] SERVICE ALERT: ieatlms7435;check pm collection statistical SBG-IS;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '10.210.245.3' (RSA) to the list of known hosts.
[1574637379] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574637380] SERVICE ALERT: ieatlms7435;check pm collection statistical SBG-IS;CRITICAL;SOFT;2;(No output on stdout) stderr: Warning: Permanently added '10.210.245.3' (RSA) to the list of known hosts.
[1574637417] HOST ALERT: ieatnetsimv9034-88;DOWN;SOFT;1;PING NOK
[1574637479] HOST ALERT: ieatnetsimv9034-88;UP;SOFT;1;PING OK
[1574637504] SERVICE ALERT: ieatlms7435;check pm collection statistical SBG-IS;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '10.210.245.3' (RSA) to the list of known hosts.
[1574637507] SERVICE ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574637548] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (625) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574637607] HOST ALERT: ieatnetsimv7102-58;DOWN;SOFT;1;PING NOK
[1574637669] HOST ALERT: ieatnetsimv7102-58;UP;SOFT;1;PING OK
[1574637691] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574637816] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (651) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574637945] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (664) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574638019] SERVICE ALERT: ieatlms3901;check fm supervision nodes ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1000)
[1574638069] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (680) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574638147] SERVICE ALERT: ieatlms3901;check fm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1574639037] SERVICE ALERT: ieatlms3901;check fm hb failures EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1) IS HIGHER THAN 20% OF ENM NODES (0)
[1574639165] SERVICE ALERT: ieatlms3901;check fm hb failures EPG;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (50)
[1574639641] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  MTAS: 1
[1574639771] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574639876] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!