[1572912757] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1572912936] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1572913189] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1572913878] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572914594] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1572914662] HOST ALERT: ieatnetsimv7035-115;DOWN;SOFT;1;PING NOK
[1572914724] HOST ALERT: ieatnetsimv7035-115;UP;SOFT;1;PING OK
[1572914805] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572914932] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572915057] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572917485] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2337) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572917615] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2436) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572918665] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572924945] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2321) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572925071] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2321) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572925197] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572925871] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572925996] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572926121] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572928804] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572929133] SERVICE ALERT: ieatnetsimv7035-97;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.46%
[1572929283] SERVICE ALERT: ieatnetsimv7035-97;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 17.27%
[1572930430] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1572930557] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1572930686] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1572931167] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1560) IS LOWER THAN 80% OF ENM  NODES (2560)
[1572931242] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.27%
[1572931293] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1560) IS LOWER THAN 80% OF ENM  NODES (2560)
[1572931394] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572931423] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2551) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1572931544] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.55%
[1572931593] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1572932410] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572932537] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572932664] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572933224] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1572933360] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1572933498] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1572934791] HOST ALERT: ieatnetsimv7035-41;DOWN;SOFT;1;PING NOK
[1572934853] HOST ALERT: ieatnetsimv7035-41;UP;SOFT;1;PING OK
[1572934982] SERVICE ALERT: ieatnetsimv5116-06;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1572935117] SERVICE ALERT: ieatnetsimv5116-06;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1572935175] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 51.14%
[1572935259] SERVICE ALERT: ieatnetsimv5116-06;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1572936270] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2506) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572936928] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572938834] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1572939087] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1572940310] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1572941034] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1572942439] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1572944134] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572944261] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572944388] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572945325] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572945371] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572946104] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1572946289] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1572946474] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1572947258] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572947386] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2560) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572947872] SERVICE ALERT: ieatnetsimv5116-05;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1572947875] SERVICE ALERT: ieatnetsimv5116-06;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1572947913] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1572948229] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1572948417] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1572948619] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1572948937] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (9609) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572948980] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (70) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572949081] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12399) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572949106] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572949223] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12608) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572949231] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (76) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572949692] SERVICE ALERT: ieatnetsimv5116-05;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1572949729] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1572949825] SERVICE ALERT: ieatnetsimv5116-05;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1572949860] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1572949957] SERVICE ALERT: ieatnetsimv5116-05;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1572949990] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1572950487] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1572950677] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2321) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572950805] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572950932] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572951661] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1572951841] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1572951939] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1572952022] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1572952124] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1572952208] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1572952310] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1572952389] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1572952570] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1572952848] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572952983] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572953112] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572953197] SERVICE ALERT: ieatnetsimv5116-69;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572953348] SERVICE ALERT: ieatnetsimv5116-69;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572953500] SERVICE ALERT: ieatnetsimv5116-69;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572953997] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954003] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954004] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954006] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954019] SERVICE ALERT: ieatnetsimv5116-54;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954031] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954043] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954047] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954047] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954149] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954154] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954155] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954157] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954170] SERVICE ALERT: ieatnetsimv5116-54;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954175] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1572954182] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.08%
[1572954194] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954198] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954198] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954300] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572954306] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572954306] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572954308] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572954322] SERVICE ALERT: ieatnetsimv5116-54;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572954346] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572954349] SERVICE ALERT: ieatnetsimv5116-67;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.08%
[1572954350] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572954400] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954551] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954691] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954699] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954700] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954703] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572954726] SERVICE ALERT: ieatnetsimv5116-87;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572954843] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.63%
[1572954852] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954853] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572954877] SERVICE ALERT: ieatnetsimv5116-87;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572955003] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572955005] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572955028] SERVICE ALERT: ieatnetsimv5116-87;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572955240] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1572955295] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572955326] SERVICE ALERT: ieatnetsimv5116-88;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572955366] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1572955446] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572955478] SERVICE ALERT: ieatnetsimv5116-88;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572955493] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1572955598] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572955629] SERVICE ALERT: ieatnetsimv5116-88;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572955683] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1572956092] HOST ALERT: ieatnetsimv9041-32;DOWN;SOFT;1;PING NOK
[1572956154] HOST ALERT: ieatnetsimv9041-32;UP;SOFT;1;PING OK
[1572956231] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1572956448] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13695) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572956719] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572957130] SERVICE ALERT: ieatnetsimv5116-69;check netsim cpu;OK;HARD;3;OK- CPU IS USED 58.12%
[1572957595] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572957608] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572957613] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572957618] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572957622] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572957637] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572957644] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572957691] SERVICE ALERT: ieatnetsimv5116-86;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572957747] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572957760] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572957764] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572957769] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572957773] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572957788] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 15.29%
[1572957795] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572957843] SERVICE ALERT: ieatnetsimv5116-86;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572957898] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572957911] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572957916] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572957921] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572957924] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572957937] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 56.74%
[1572957948] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572957952] SERVICE ALERT: ieatnetsimv5116-54;check netsim cpu;OK;HARD;3;OK- CPU IS USED 30.88%
[1572957980] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;HARD;3;OK- CPU IS USED 19.55%
[1572957994] SERVICE ALERT: ieatnetsimv5116-86;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572958141] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572958531] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1572958634] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.85%
[1572958636] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;OK;HARD;3;OK- CPU IS USED 36.31%
[1572958659] SERVICE ALERT: ieatnetsimv5116-87;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.17%
[1572958794] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572958890] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572958976] SERVICE ALERT: ieatnetsimv5116-05;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1572959007] SERVICE FLAPPING ALERT: ieatnetsimv5116-04;check netsim nodes;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1572959007] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1572959040] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.94%
[1572959191] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;HARD;3;OK- CPU IS USED 35.54%
[1572959264] SERVICE ALERT: ieatnetsimv5116-88;check netsim cpu;OK;HARD;3;OK- CPU IS USED 72.54%
[1572959772] SERVICE ALERT: ieatnetsimv5116-43;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.39%
[1572959932] SERVICE ALERT: ieatnetsimv5116-44;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.58%
[1572960057] SERVICE ALERT: ieatnetsimv5116-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.24%
[1572960184] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.69%
[1572960196] SERVICE ALERT: ieatnetsimv5116-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.43%
[1572960335] SERVICE ALERT: ieatnetsimv5116-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 1.68%
[1572960345] SERVICE ALERT: ieatnetsimv5116-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.76%
[1572960354] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.59%
[1572960546] SERVICE ALERT: ieatnetsimv5116-64;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.33%
[1572960555] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.34%
[1572960726] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.32%
[1572960736] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.37%
[1572960796] SERVICE ALERT: ieatnetsimv5116-86;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.38%
[1572961199] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572961200] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12189) IS FROM 80% TO 95% OF ENM  NODES (13982)
[1572961216] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572961341] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12189) IS FROM 80% TO 95% OF ENM  NODES (13982)
[1572961350] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572961368] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572961483] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12189) IS FROM 80% TO 95% OF ENM  NODES (13982)
[1572961501] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572961519] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572961568] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572961569] SERVICE ALERT: ieatlms3901;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1793) IS BETWEEN 5% AND 20% OF ENM NODES (13982)
[1572961707] SERVICE ALERT: ieatlms3901;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1793) IS BETWEEN 5% AND 20% OF ENM NODES (13982)
[1572961719] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572961842] SERVICE ALERT: ieatlms3901;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1793) IS BETWEEN 5% AND 20% OF ENM NODES (13982)
[1572961870] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572963040] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (11763) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13982)
[1572963050] SERVICE ALERT: ieatnetsimv5116-45;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.81%
[1572963057] SERVICE FLAPPING ALERT: ieatnetsimv5116-46;check netsim cpu;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1572963057] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.85%
[1572963067] SERVICE ALERT: ieatnetsimv5116-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 0.75%
[1572963188] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (11763) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13982)
[1572963370] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (11561) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13982)
[1572964006] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.56%
[1572964157] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 49.37%
[1572964284] SERVICE ALERT: ieatnetsimv5116-86;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1572964308] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.8%
[1572964420] SERVICE ALERT: ieatnetsimv5116-86;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1572964426] SERVICE ALERT: ieatnetsimv5116-86;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.62%
[1572964459] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.2%
[1572964552] SERVICE ALERT: ieatnetsimv5116-86;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1572964577] SERVICE ALERT: ieatnetsimv5116-86;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 62.31%
[1572964609] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 81.42%
[1572965350] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572965480] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572965607] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572966002] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572966128] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572966255] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572967276] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12397) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572967419] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12397) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572967524] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572967565] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1572967565] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12397) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572967648] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572967773] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1572967773] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572967959] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572967968] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1572967970] SERVICE ALERT: ieatnetsimv5116-75;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1572968112] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572968120] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1572968122] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1572968176] SERVICE ALERT: ieatnetsimv5116-75;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1572968240] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.7%
[1572968266] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572968272] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1572968391] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 57.12%
[1572968397] SERVICE ALERT: ieatnetsimv5116-75;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1572968553] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1572968853] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1572969216] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572969862] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572971186] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572971379] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572971903] SERVICE FLAPPING ALERT: ieatnetsimv5116-75;check netsim cpu;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1572971903] SERVICE ALERT: ieatnetsimv5116-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.48%
[1572972824] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572972953] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572973082] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572973466] HOST ALERT: ieatnetsimv7035-12;DOWN;SOFT;1;PING NOK
[1572973469] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572973528] HOST ALERT: ieatnetsimv7035-12;UP;SOFT;1;PING OK
[1572973597] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572973726] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1572974993] HOST ALERT: ieatnetsimv7035-33;DOWN;SOFT;1;PING NOK
[1572975055] HOST ALERT: ieatnetsimv7035-33;UP;SOFT;1;PING OK
[1572975496] SERVICE FLAPPING ALERT: ieatnetsimv5116-74;check netsim cpu;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1572975496] SERVICE ALERT: ieatnetsimv5116-74;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.51%
[1572975794] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1572975915] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1572976035] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1572978407] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12397) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572978555] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12397) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572978586] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572978697] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12397) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572978712] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572978838] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572979342] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.05%
[1572979492] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.88%
[1572980291] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572980932] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1572983243] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.02%
[1572983394] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 49.78%
[1572983544] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.23%
[1572983695] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 62.28%
[1572983898] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2407) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572984024] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.4%
[1572984025] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2552) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1572984175] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.35%
[1572985925] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572986023] SERVICE FLAPPING ALERT: ieatnetsimv5116-04;check netsim nodes;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1572986045] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1572989551] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12767) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572989695] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13148) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572989840] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13148) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1572993253] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572993382] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572993509] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1572997061] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)