[1566083232] SERVICE ALERT: ieatwlvm5116;check workload profiles started;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1566083353] SERVICE ALERT: ieatwlvm5116;check workload profiles started;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1566083474] SERVICE ALERT: ieatwlvm5116;check workload profiles started;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE NOT STARTED:
[1566087086] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566087228] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (14032) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566087417] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566087558] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (14028) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566087747] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566087889] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (14028) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566087959] HOST ALERT: ieatnetsimv5116-52;DOWN;SOFT;1;PING NOK
[1566088021] HOST ALERT: ieatnetsimv5116-52;UP;SOFT;1;PING OK
[1566095278] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566095467] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566095610] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14024) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566096094] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566096220] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;SOFT;2;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566096405] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566096590] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566096715] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566099280] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566099422] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (14030) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566099611] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566099754] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (14030) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566110746] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1566116079] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6741) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1566116222] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6725) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1566116413] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6705) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1566117985] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566118174] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566118243] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (20204) IS LOWER THAN 80% OF ENM NODES (25486)
[1566118249] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (4432) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1566118377] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566121999] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14032) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566124434] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1566124615] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1566124796] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1566125094] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566125275] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566125455] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.81%
[1566125456] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566125606] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.19%
[1566125669] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566125757] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.46%
[1566125907] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.03%
[1566125919] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (14032) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566126058] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.08%
[1566126109] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566126208] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.96%
[1566126307] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (14033) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566126359] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.34%
[1566126496] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566126510] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.26%
[1566126661] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.51%
[1566126685] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566126811] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.69%
[1566126916] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566126962] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.67%
[1566127112] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.39%
[1566132780] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566132906] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;SOFT;2;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566133857] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1566134517] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1566135359] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.72%
[1566135428] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.28%
[1566135510] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 61.05%
[1566135578] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.61%
[1566136691] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566136844] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;SOFT;2;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566137029] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566137214] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566137399] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566139644] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.02%
[1566139795] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.63%
[1566139945] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.77%
[1566140096] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.71%
[1566144837] HOST ALERT: ieatnetsimv5116-77;DOWN;SOFT;1;PING NOK
[1566144839] SERVICE ALERT: ieatnetsimv5116-80;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566144839] SERVICE ALERT: ieatnetsimv5116-82;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566144899] HOST ALERT: ieatnetsimv5116-77;UP;SOFT;1;PING OK
[1566144959] SERVICE ALERT: ieatnetsimv5116-80;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566144960] SERVICE ALERT: ieatnetsimv5116-82;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566145080] SERVICE ALERT: ieatnetsimv5116-80;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566145080] SERVICE ALERT: ieatnetsimv5116-82;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566149995] SERVICE ALERT: ieatlms3901;check cm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (20888) IS BETWEEN 80% AND 95% OF ENM NODES (25486)
[1566159663] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.53%
[1566159814] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.72%
[1566163512] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566163559] HOST ALERT: ieatnetsimv5116-112;DOWN;SOFT;1;PING NOK
[1566163622] HOST ALERT: ieatnetsimv5116-112;UP;SOFT;1;PING OK
[1566163699] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566163885] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566164330] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1566164508] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1566164687] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1566167485] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13141) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1566167636] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13732) IS HIGHER THAN 95% OF ENM COMECIM NODES (14149)
[1566168017] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1944) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1566168164] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2072) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566168326] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2160) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566169054] SERVICE ALERT: ieatwlvm5116;check workload profiles;CRITICAL;SOFT;1;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1566169175] SERVICE ALERT: ieatwlvm5116;check workload profiles;OK;SOFT;2;OK- ALL WORKLOAD PROFILES ARE RUNNING