[1574467255] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1574467373] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574467627] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1574468224] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574468424] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1574468509] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1574468530] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1574468661] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (79) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574468747] HOST ALERT: ieatnetsimv7035-75;DOWN;SOFT;1;PING NOK
[1574468802] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2508) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574468809] HOST ALERT: ieatnetsimv7035-75;UP;SOFT;1;PING OK
[1574468951] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (745) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574469048] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (11902) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574469172] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574469196] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12314) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574469343] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12548) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574469788] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (150)
[1574469916] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1574470772] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574471762] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574471895] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (49) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574472558] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (694) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574472687] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (463) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574472814] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (463) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574473349] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1101) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1574473475] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1200)!
[1574473525] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574473656] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574473785] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RNC;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1574473785] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574474376] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574474500] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574475243] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1574475373] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1574475486] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574475501] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1574475606] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1574475606] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574475667] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574475786] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574475849] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574475868] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574475913] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574476001] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574476010] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2356) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574476040] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574476141] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2356) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574476143] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574476270] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2323) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574477309] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1574477395] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (5) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (5)
[1574478136] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574478260] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1574479265] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574479322] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574479443] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574479447] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (16) IS HIGHER THAN 20% OF ENM NODES (50)
[1574479511] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1574479511] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574479571] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574479621] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574481373] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574481554] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574481736] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574481825] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574481955] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574482080] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574482895] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574483036] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574483180] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (50)
[1574483479] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2557) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574483770] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574485398] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574485683] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1574485683] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574486297] HOST ALERT: ieatnetsimv5116-83;DOWN;SOFT;1;PING NOK
[1574486359] HOST ALERT: ieatnetsimv5116-83;UP;SOFT;1;PING OK
[1574487088] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2393) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574487202] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574487216] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2463) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574487326] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574487449] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574487624] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2136) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574487646] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1296) IS FROM 80% TO 95% OF ENM  NODES (1514)
[1574487659] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574487749] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (2560)!
[1574487772] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1514)!
[1574487782] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574487906] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574489299] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574489423] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574489986] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1310) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574490111] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1310) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574490236] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1749) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574490369] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (126) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574490371] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574490448] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574490493] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (126) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574490494] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574490550] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1574490617] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574490619] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1574490859] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574490903] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (664) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574490995] HOST ALERT: ieatnetsimv7035-37;DOWN;SOFT;1;PING NOK
[1574491028] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (673) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574491041] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574491053] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574491058] HOST ALERT: ieatnetsimv7035-37;UP;SOFT;1;PING OK
[1574491164] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (686) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574491223] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574492638] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574492968] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574493147] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574493270] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574493397] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574493621] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574493744] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574494054] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574494156] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574494181] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574494286] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574494308] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574494412] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574494504] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574494545] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2333) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574494599] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12749) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574494674] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2333) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574494689] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574494735] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574494803] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2333) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574494874] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1574494874] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574494883] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12749) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574494885] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574496640] HOST ALERT: ieatnetsimv7035-33;DOWN;SOFT;1;PING NOK
[1574496640] HOST ALERT: ieatnetsimv9041-12;DOWN;SOFT;1;PING NOK
[1574496645] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.18%
[1574496683] HOST ALERT: ieatnetsimv9041-12;UP;SOFT;1;PING OK
[1574496702] HOST ALERT: ieatnetsimv7035-33;UP;SOFT;1;PING OK
[1574496796] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 50.75%
[1574496947] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574497004] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574497098] SERVICE ALERT: ieatnetsimv5116-83;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 53.94%
[1574497721] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574497914] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574498019] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1574498412] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2557) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574498514] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574498539] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574498800] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (900) IS LOWER THAN 80% OF ENM  NODES (1200)
[1574498946] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (900) IS LOWER THAN 80% OF ENM  NODES (1200)
[1574499088] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1193) IS HIGHER THAN 95% OF ENM  NODES (1200)
[1574503603] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574503730] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1574504486] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574504529] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574504610] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574504658] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1574504979] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574505125] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574505157] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574505252] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574505336] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574505379] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (142) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574505741] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13263) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574505804] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574505888] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13263) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574505989] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574506036] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574506104] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12934) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574506175] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574506247] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13064) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574506396] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (13166) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574507819] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574507948] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574508040] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574508385] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574508514] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1574508621] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574508668] SERVICE ALERT: ieatlms3901;check synchronized EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (50)
[1574508745] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574508793] SERVICE ALERT: ieatlms3901;check synchronized EPG;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574508869] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (144) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574508902] HOST ALERT: ieatnetsimv9034-113;DOWN;SOFT;1;PING NOK
[1574508964] HOST ALERT: ieatnetsimv9034-113;UP;SOFT;1;PING OK
[1574509036] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574509502] HOST ALERT: ieatnetsimv9034-11;DOWN;SOFT;1;PING NOK
[1574509565] HOST ALERT: ieatnetsimv9034-11;UP;SOFT;1;PING OK
[1574509802] HOST ALERT: ieatnetsimv9034-112;DOWN;SOFT;1;PING NOK
[1574509864] HOST ALERT: ieatnetsimv9034-112;UP;SOFT;1;PING OK
[1574510014] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13939) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574510046] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1446) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574510178] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1747) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574510973] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (705) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574511002] HOST ALERT: ieatnetsimv9034-07;DOWN;SOFT;1;PING NOK
[1574511064] HOST ALERT: ieatnetsimv9034-07;UP;SOFT;1;PING OK
[1574511098] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (706) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574511223] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (711) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574511985] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574512110] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (50)
[1574512235] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574512359] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574512373] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (328) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574512426] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574512483] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (50)
[1574512508] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574512552] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574512585] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (133) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574512678] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574512712] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574512822] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2317) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574512838] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (133) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574512955] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2317) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574513027] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574513027] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574513084] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2317) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574515502] HOST ALERT: ieatnetsimv9034-114;DOWN;SOFT;1;PING NOK
[1574515564] HOST ALERT: ieatnetsimv9034-114;UP;SOFT;1;PING OK
[1574515627] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574515821] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574515823] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574515947] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (76) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574516013] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574516073] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574516197] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574516285] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574516321] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574516445] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574516694] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2557) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574517270] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574517394] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574517520] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574517995] HOST ALERT: ieatnetsimv9041-32;DOWN;SOFT;1;PING NOK
[1574518057] HOST ALERT: ieatnetsimv9041-32;UP;SOFT;1;PING OK
[1574518432] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (629) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574518560] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (631) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574518595] HOST ALERT: ieatnetsimv7035-41;DOWN;SOFT;1;PING NOK
[1574518657] HOST ALERT: ieatnetsimv7035-41;UP;SOFT;1;PING OK
[1574518688] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (634) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574518883] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574518895] HOST ALERT: ieatnetsimv7035-84;DOWN;SOFT;1;PING NOK
[1574518958] HOST ALERT: ieatnetsimv7035-84;UP;SOFT;1;PING OK
[1574519010] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574519671] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574519770] SERVICE ALERT: ieatlms3901;check cm supervision nodes Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (27) IS LOWER THAN 80% OF ENM NODES (50)
[1574519795] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574519856] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574519895] SERVICE ALERT: ieatlms3901;check cm supervision nodes Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1574519919] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574519980] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574520103] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574520302] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2362) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574520310] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574520431] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2437) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574520434] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;OK;SOFT;2;OK- NUMBER OF PM FILES (28) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574520493] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574520712] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (79) IS FROM 80% TO 95% OF ENM  NODES (98)
[1574520842] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (11987) IS FROM 80% TO 95% OF ENM  NODES (13969)
[1574520860] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (97) IS HIGHER THAN 95% OF ENM  NODES (98)
[1574520976] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13725) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574521112] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1547) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574521119] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12785) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574521140] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (37) IS LOWER THAN 80% OF ENM  NODES (50)
[1574521160] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574521255] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13970) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574521262] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1574521271] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (42) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574521285] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (49) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574521407] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (40) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574521413] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13234) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574521424] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574521554] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13480) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574521594] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574522733] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574522857] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574522983] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574523095] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;1;PING NOK
[1574523157] HOST ALERT: ieatnetsimv9041-01;UP;SOFT;1;PING OK
[1574523286] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574523651] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574523707] SERVICE FLAPPING ALERT: ieatlms7435;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1574523707] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (28) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574523779] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574523907] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574525287] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (11906) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574525325] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574525438] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13890) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574525456] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574525581] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574525705] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574526038] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574526165] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574526290] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574526443] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  MGW: 5
[1574526572] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 144 MGW: 6
[1574526699] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 348
[1574527311] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574527347] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1574527370] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12787) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574527434] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574527503] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12845) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574527558] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574527639] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13949) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574529068] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574529197] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1574529321] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574529458] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1574529733] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574529853] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574529974] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574530187] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574530307] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1574530313] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574530412] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574530446] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574530537] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (42) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574530671] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574531119] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (121) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574531254] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574531264] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12807) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574531380] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (146) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574531417] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (13815) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574531433] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (114) IS LOWER THAN 80% OF ENM  NODES (150)
[1574531561] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (114) IS LOWER THAN 80% OF ENM  NODES (150)
[1574531684] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1574532202] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1301) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574532356] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1574533295] HOST ALERT: ieatnetsimv7035-38;DOWN;SOFT;1;PING NOK
[1574533295] HOST ALERT: ieatnetsimv7035-111;DOWN;SOFT;1;PING NOK
[1574533357] HOST ALERT: ieatnetsimv7035-38;UP;SOFT;1;PING OK
[1574533357] HOST ALERT: ieatnetsimv7035-111;UP;SOFT;1;PING OK
[1574534051] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574534180] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574534265] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574534389] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574534513] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574534899] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (700) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574535028] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (706) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574535153] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (707) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574536959] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574537096] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574538560] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2314) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574538593] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (131) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574538689] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2314) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574538719] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (131) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574538764] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12553) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574538823] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2314) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574538852] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (131) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574538911] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12553) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574538950] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574539083] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574539089] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12553) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574539209] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (28) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574539559] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574539752] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574539945] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574541221] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574541505] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574541629] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574541743] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574541753] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574541770] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574541864] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574541891] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574541985] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574542012] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574542241] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574542359] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1574542362] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574542431] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2461) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574542459] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (145) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574542480] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1574542483] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574542600] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1574542717] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13389) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574542817] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574542856] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1574542942] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574542976] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1574543076] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574543097] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1574543227] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1574543616] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;SOFT;2;OK- nas_healthcheck SUCCESSFULLY PASSED
[1574543619] SERVICE FLAPPING ALERT: ieatlms5735-1;check system service hc;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1574543619] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574543847] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574543875] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1574544301] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574544427] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574544461] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1574544552] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574544890] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1574545014] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1574545033] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (67) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574545095] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574545140] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1574545158] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (68) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574545273] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574545293] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (78) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574545369] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574545451] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574545546] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574545673] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (37) IS LOWER THAN 80% OF ENM  NODES (50)
[1574545676] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574545803] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RNC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574545904] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574545970] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574546043] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2313) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574546065] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (132) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574546181] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574546198] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (132) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574546202] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2313) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574546213] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (23) IS LOWER THAN 80% OF ENM  NODES (30)
[1574546332] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2313) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574546349] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (132) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574546351] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574546356] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12561) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574546461] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574546477] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574546505] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12587) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574546590] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574546652] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12615) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574546755] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1574547117] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574547246] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (78) IS HIGHER THAN 95% OF ENM  NODES (80)
[1574547454] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12589) IS FROM 80% TO 95% OF ENM  NODES (13969)
[1574547509] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574547590] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13581) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574547691] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574547723] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12750) IS FROM 80% TO 95% OF ENM  NODES (13969)
[1574547863] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12761) IS FROM 80% TO 95% OF ENM  NODES (13969)
[1574547872] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574547986] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574547998] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (13193) IS FROM 80% TO 95% OF ENM  NODES (13969)
[1574548283] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574548318] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  RBS: 214 ERBS: 339 MGW: 22
[1574548457] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574548532] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574548665] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574548685] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574548744] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1574548789] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574548823] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574548898] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (1200)
[1574548951] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (49) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574549169] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574549294] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574549358] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574549418] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574549481] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574549605] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574549630] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574549812] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13949) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574550279] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574550288] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574550757] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574550881] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574551534] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574551716] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574551898] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574552290] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1289) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574552419] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1310) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574552547] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1330) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574552700] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1422) IS FROM 80% TO 95% OF ENM  NODES (1514)
[1574552857] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1514)!
[1574553024] SERVICE ALERT: ieatlms7435;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574553165] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (677) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574553209] SERVICE ALERT: ieatlms7435;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574553292] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (662) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574553422] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (703) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574553551] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2557) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574553555] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)