[1576022923] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576022946] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1576023082] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1576023214] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576023218] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;HARD;3;OK
[1576023834] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (14035)
[1576024243] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1576024980] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576026819] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (42) IS HIGHER THAN 20% OF ENM NODES (80)
[1576026944] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (49) IS HIGHER THAN 20% OF ENM NODES (80)
[1576027069] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (50) IS HIGHER THAN 20% OF ENM NODES (80)
[1576027335] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (7) IS HIGHER THAN 20% OF ENM NODES (30)
[1576028585] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8) IS HIGHER THAN 20% OF ENM NODES (30)
[1576028697] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (383) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576028709] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8) IS HIGHER THAN 20% OF ENM NODES (30)
[1576028832] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (39) IS LOWER THAN 5% OF ENM NODES (1200)
[1576028839] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1576028956] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (316) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576029080] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (53) IS LOWER THAN 5% OF ENM NODES (1200)
[1576030108] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (2) IS LOWER THAN 80% OF ENM  NODES (4)
[1576030127] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (21) IS HIGHER THAN 20% OF ENM NODES (100)
[1576030278] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (44) IS HIGHER THAN 20% OF ENM NODES (100)
[1576030282] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (2) IS LOWER THAN 80% OF ENM  NODES (4)
[1576030404] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (19) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576030407] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (4)!
[1576030673] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (10) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576030939] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1576031045] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (6494) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576031173] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (104) IS LOWER THAN 5% OF ENM NODES (14035)
[1576033441] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576034277] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576034544] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (16) IS HIGHER THAN 20% OF ENM NODES (30)
[1576034672] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576034801] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1576034902] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (2810) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576035035] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (810) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576035169] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1510) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576036044] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576036168] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576037046] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1576037607] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576037884] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (12) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576038056] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576038434] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1576038557] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (13) IS HIGHER THAN 20% OF ENM NODES (30)
[1576038682] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1576038781] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (4762) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576039891] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576040015] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576040140] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576041783] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (42) IS HIGHER THAN 20% OF ENM NODES (80)
[1576041911] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (41) IS HIGHER THAN 20% OF ENM NODES (80)
[1576042039] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (23) IS HIGHER THAN 20% OF ENM NODES (80)
[1576042069] SERVICE ALERT: ieatwlvm9034;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1576042285] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576042390] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2144) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576042777] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.83%
[1576042881] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576042928] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.62%
[1576043066] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576043079] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;OK;HARD;3;OK- CPU IS USED 80.57%
[1576043252] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576043750] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MTAS NODES !
[1576044987] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1576045136] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1576045291] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1576045643] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (80)
[1576045891] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576046917] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1576047362] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1576047396] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (98)
[1576047479] SERVICE ALERT: ieatwlvm9034;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1576047784] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (50)
[1576047786] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1576047866] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 285 processes with STATE = RSZDT
[1576047901] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ESC;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576047912] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576047926] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 288 processes with STATE = RSZDT
[1576047927] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576047937] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576047976] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MSC-BC-BSP;CRITICAL;SOFT;1;(Service check timed out after 800.44 seconds)
[1576047986] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 317 processes with STATE = RSZDT
[1576047994] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes SBG-IS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048000] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048014] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6274;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048032] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (6831) IS LOWER THAN 80% OF ENM  NODES (14035)
[1576048034] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;CRITICAL;SOFT;1;(Service check timed out after 800.00 seconds)
[1576048046] SERVICE ALERT: localhost;Total Processes;WARNING;HARD;4;PROCS WARNING: 338 processes with STATE = RSZDT
[1576048076] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes ESC;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048080] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes RBS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048086] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048089] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-BC-IS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048128] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048131] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes RNC;CRITICAL;SOFT;1;(Service check timed out after 800.38 seconds)
[1576048135] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1576048148] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048151] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1576048177] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MTAS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048197] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes DSC;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048218] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (2560)
[1576048219] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes RadioNode;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048222] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (798) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576048226] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1576048232] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1576048237] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes SBG-IS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048244] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048246] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1576048273] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1750) IS HIGHER THAN 95% OF ENM NODES (1750)
[1576048281] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes SGSN-MME;CRITICAL;SOFT;1;(Service check timed out after 800.24 seconds)
[1576048283] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes CISCO-ASR900;CRITICAL;SOFT;2;(Service check timed out after 800.38 seconds)
[1576048288] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-BC-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (13) IS HIGHER THAN 95% OF ENM NODES (13)
[1576048296] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes EPG;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048297] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (13) IS HIGHER THAN 20% OF ENM NODES (30)
[1576048301] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1576048308] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1576048317] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1576048337] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1576048348] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;(Service check timed out after 800.60 seconds)
[1576048348] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;CRITICAL;SOFT;1;(Service check timed out after 800.08 seconds)
[1576048361] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (7666) IS LOWER THAN 80% OF ENM  NODES (14035)
[1576048368] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048369] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (14035) IS HIGHER THAN 95% OF ENM NODES (14035)
[1576048376] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MSC-BC-IS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048398] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1576048407] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes EPG;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048415] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes RNC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (4) IS HIGHER THAN 95% OF ENM NODES (4)
[1576048416] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048421] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes SGSN-MME;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048431] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1576048438] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-Indoor;CRITICAL;SOFT;1;(Service check timed out after 800.99 seconds)
[1576048449] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1576048462] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes CISCO-ASR900;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1576048480] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MSC-BC-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1576048492] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1576048510] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MSC-BC-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (13) IS HIGHER THAN 95% OF ENM NODES (13)
[1576048523] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (25) IS HIGHER THAN 20% OF ENM NODES (100)
[1576048544] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1576048561] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (2502) IS HIGHER THAN 95% OF ENM NODES (2560)
[1576048569] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1576048575] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes DSC;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1576048579] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (97) IS HIGHER THAN 95% OF ENM  NODES (98)
[1576048586] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1576048592] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1576048595] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MSC-BC-BSP;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1576048605] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;(Service check timed out after 800.06 seconds)
[1576048611] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes EPG;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1576048614] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1576048624] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (8523) IS LOWER THAN 80% OF ENM  NODES (14035)
[1576048626] HOST ALERT: ieatnetsimv9041-23;DOWN;SOFT;1;PING NOK
[1576048633] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576048651] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (510) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576048663] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes RBS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1514) IS HIGHER THAN 95% OF ENM NODES (1514)
[1576048664] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MTAS;CRITICAL;SOFT;1;(Service check timed out after 801.00 seconds)
[1576048668] SERVICE ALERT: ieatlms5735-1;check fm hb failures MINI-LINK-Indoor;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MINI-LINK-Indoor NODES !
[1576048670] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1576048678] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1576048688] HOST ALERT: ieatnetsimv9041-23;UP;SOFT;1;PING OK
[1576048689] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ESC;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1576048704] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes EPG;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1576048706] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (46) IS HIGHER THAN 20% OF ENM NODES (100)
[1576048748] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MGW;CRITICAL;SOFT;1;(Service check timed out after 800.24 seconds)
[1576048756] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1576048767] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576048798] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1576048821] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1576048821] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ESC;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1576048832] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1576048836] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes EPG;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1576048854] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1576048865] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1576048875] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes EPG;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1576048883] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1576048888] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1576048894] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (21) IS LOWER THAN 80% OF ENM  NODES (30)
[1576048898] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (40) IS HIGHER THAN 20% OF ENM NODES (100)
[1576048900] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1576048905] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MSC-BC-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1576048925] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1576048930] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MGW;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1576048946] SERVICE ALERT: localhost;Total Processes;OK;HARD;4;PROCS OK: 180 processes with STATE = RSZDT
[1576048957] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ESC;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1576049033] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576049252] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (33) IS HIGHER THAN 20% OF ENM NODES (80)
[1576049384] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (29) IS HIGHER THAN 20% OF ENM NODES (80)
[1576049501] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (11) IS HIGHER THAN 20% OF ENM NODES (30)
[1576049553] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (27) IS HIGHER THAN 20% OF ENM NODES (80)
[1576049605] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (6012) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576049634] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (10) IS HIGHER THAN 20% OF ENM NODES (30)
[1576049759] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1576049869] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576049943] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2290) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1576049994] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576050070] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2556) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1576050132] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1576050196] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2252) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1576050323] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2556) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1576050452] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12997) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576050698] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1576051455] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576051585] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576051612] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576051717] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576051736] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576051752] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  RadioNode: 52
[1576051863] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576051890] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576051936] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576052057] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (78) IS HIGHER THAN 95% OF ENM  NODES (80)
[1576052058] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576052062] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1576052278] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13968) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576052319] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized Router6274;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576053363] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576053669] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1576055321] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1576055321] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (46) IS HIGHER THAN 20% OF ENM NODES (50)
[1576055412] SERVICE ALERT: ieatnetsimv7035-31;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.92%
[1576055562] SERVICE ALERT: ieatnetsimv7035-31;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 18.53%
[1576055663] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (17) IS HIGHER THAN 20% OF ENM NODES (30)
[1576056772] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (13) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576056975] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8) IS HIGHER THAN 20% OF ENM NODES (30)
[1576057591] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (68) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576057653] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2370) IS FROM 80% TO 95% OF ENM  NODES (2562)
[1576057790] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (2382) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1576057849] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576057957] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2498) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1576057971] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (19) IS LOWER THAN 80% OF ENM  NODES (30)
[1576058013] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (29) IS LOWER THAN 80% OF ENM  NODES (80)
[1576058097] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (17) IS LOWER THAN 80% OF ENM  NODES (30)
[1576058142] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576058237] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (18) IS LOWER THAN 80% OF ENM  NODES (30)
[1576058274] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (35) IS LOWER THAN 80% OF ENM  NODES (80)
[1576058409] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576058556] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (37) IS LOWER THAN 80% OF ENM  NODES (80)
[1576058688] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (37) IS LOWER THAN 80% OF ENM  NODES (80)
[1576058823] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576059043] SERVICE ALERT: ieatnetsimv9034-45;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.33%
[1576059073] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (23) IS LOWER THAN 80% OF ENM  NODES (30)
[1576059193] SERVICE ALERT: ieatnetsimv9034-45;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.45%
[1576059224] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576059358] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1576059397] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576059507] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10712) IS LOWER THAN 80% OF ENM  NODES (14035)
[1576059717] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (11184) IS LOWER THAN 80% OF ENM  NODES (14035)
[1576059753] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576059926] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (35) IS LOWER THAN 80% OF ENM  NODES (50)
[1576059948] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1576059948] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (11207) IS LOWER THAN 80% OF ENM  NODES (14035)
[1576060064] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576060627] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (69) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576060761] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576060961] SERVICE ALERT: ieatnetsimv9034-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576061092] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1576061112] SERVICE ALERT: ieatnetsimv9034-65;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 46.39%
[1576061243] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.96%
[1576061393] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576061751] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576061763] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13753) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576061843] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MTAS;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1576061843] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1576061888] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1576061902] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576061916] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.76%
[1576062072] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.99%
[1576062077] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.43%
[1576062114] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.96%
[1576062228] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 46.07%
[1576062264] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.85%
[1576062415] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.23%
[1576063339] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576063381] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (29) IS HIGHER THAN 20% OF ENM NODES (100)
[1576063489] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.91%
[1576063640] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.09%
[1576063653] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576063707] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576063776] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576063831] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1576063910] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1576063979] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1576063997] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (80)
[1576065024] SERVICE ALERT: ieatnetsimv9034-121;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.74%
[1576065426] HOST ALERT: ieatnetsimv7035-95;DOWN;SOFT;1;PING NOK
[1576065488] HOST ALERT: ieatnetsimv7035-95;UP;SOFT;1;PING OK
[1576065704] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;OK;HARD;3;OK- CPU IS USED 35.09%
[1576066045] SERVICE FLAPPING ALERT: ieatnetsimv9034-108;check netsim cpu;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1576066045] SERVICE ALERT: ieatnetsimv9034-108;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.32%
[1576066469] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576066921] SERVICE ALERT: ieatnetsimv7048-08;check netsim disk;OK;HARD;3;OK
[1576066945] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (150)
[1576066985] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (2562)
[1576066997] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576067074] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1576067111] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2559) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1576067199] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (150)
[1576067242] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (2562)
[1576067327] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (150)
[1576067375] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (2560)
[1576067396] SERVICE FLAPPING ALERT: ieatnetsimv9034-105;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576067454] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (150)
[1576067502] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2559) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1576067976] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (130) IS LOWER THAN 80% OF ENM  NODES (1514)
[1576068101] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1513) IS HIGHER THAN 95% OF ENM  NODES (1514)
[1576068542] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (1533) IS BETWEEN 80% AND 95% OF ENM NODES (1750)
[1576068669] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1276) IS LOWER THAN 80% OF ENM NODES (1750)
[1576068795] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1024) IS LOWER THAN 80% OF ENM NODES (1750)
[1576069109] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (696) IS LOWER THAN 80% OF ENM  NODES (1750)
[1576069235] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (441) IS LOWER THAN 80% OF ENM  NODES (1750)
[1576069260] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1576069376] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (140) IS LOWER THAN 80% OF ENM  NODES (1750)
[1576070033] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1025) IS LOWER THAN 80% OF ENM  NODES (1514)
[1576070079] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (10) IS HIGHER THAN 20% OF ENM NODES (30)
[1576070160] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1514)!
[1576070432] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1576070611] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (60) IS HIGHER THAN 20% OF ENM NODES (100)
[1576070736] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (59) IS HIGHER THAN 20% OF ENM NODES (100)
[1576070863] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (48) IS HIGHER THAN 20% OF ENM NODES (100)
[1576071044] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576071164] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576071202] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (80)
[1576071285] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576071577] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (1169) IS LOWER THAN 80% OF ENM NODES (1750)
[1576071705] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (931) IS LOWER THAN 80% OF ENM NODES (1750)
[1576071794] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1576071833] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (720) IS LOWER THAN 80% OF ENM NODES (1750)
[1576071931] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1576072055] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1576073467] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576073587] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576073707] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576073744] SERVICE ALERT: ieatnetsimv9034-45;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1576073747] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576073894] SERVICE ALERT: ieatnetsimv9034-45;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 10.15%
[1576073898] SERVICE ALERT: ieatnetsimv9034-82;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 63.31%
[1576074362] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (1216) IS LOWER THAN 80% OF ENM NODES (1750)
[1576074421] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576074467] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.98%
[1576074469] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576074502] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (1182) IS LOWER THAN 80% OF ENM NODES (1750)
[1576074541] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576074618] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.99%
[1576074629] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (1144) IS LOWER THAN 80% OF ENM NODES (1750)
[1576074662] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576074769] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.97%
[1576074818] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (44) IS HIGHER THAN 20% OF ENM NODES (80)
[1576074866] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (14035)
[1576074907] SERVICE ALERT: ieatnetsimv9034-73;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1576074909] SERVICE ALERT: ieatnetsimv9034-77;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.32%
[1576074916] SERVICE ALERT: ieatnetsimv9034-85;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.7%
[1576074919] SERVICE ALERT: ieatnetsimv9034-89;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576074922] SERVICE ALERT: ieatnetsimv9034-93;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576074923] SERVICE ALERT: ieatnetsimv9034-54;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.79%
[1576074925] SERVICE ALERT: ieatnetsimv9034-86;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576074950] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (41) IS HIGHER THAN 20% OF ENM NODES (80)
[1576074984] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576075060] SERVICE ALERT: ieatnetsimv9034-73;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 56.92%
[1576075062] SERVICE ALERT: ieatnetsimv9034-77;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.95%
[1576075068] SERVICE ALERT: ieatnetsimv9034-85;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576075070] SERVICE ALERT: ieatnetsimv9034-89;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 36.57%
[1576075073] SERVICE ALERT: ieatnetsimv9034-93;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.78%
[1576075074] SERVICE ALERT: ieatnetsimv9034-54;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 44.71%
[1576075076] SERVICE ALERT: ieatnetsimv9034-86;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.83%
[1576075088] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (45) IS HIGHER THAN 20% OF ENM NODES (80)
[1576075219] SERVICE ALERT: ieatnetsimv9034-85;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.18%
[1576075430] SERVICE ALERT: ieatnetsimv9034-118;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576075581] SERVICE ALERT: ieatnetsimv9034-118;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576075732] SERVICE ALERT: ieatnetsimv9034-118;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576076535] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576076686] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576076709] SERVICE ALERT: ieatnetsimv9034-07;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.73%
[1576076836] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576076859] SERVICE ALERT: ieatnetsimv9034-07;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 16.58%
[1576077322] SERVICE FLAPPING ALERT: ieatnetsimv9034-72;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576077656] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1095) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576077809] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1096) IS HIGHER THAN 20% OF ENM NODES (1200)
[1576077939] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (64) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576078071] SERVICE FLAPPING ALERT: ieatnetsimv9034-106;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576078080] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576078214] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576078227] SERVICE FLAPPING ALERT: ieatnetsimv9034-59;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576078327] SERVICE FLAPPING ALERT: ieatnetsimv9034-115;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576078360] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576078399] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;OK;HARD;3;OK- CPU IS USED 21.48%
[1576078477] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9414) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576078483] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576078611] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (14035)
[1576078620] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576078751] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (11) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576078756] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8375) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576078876] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576078893] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (6935) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576078993] SERVICE FLAPPING ALERT: ieatnetsimv9034-100;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576079023] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (14035)
[1576079362] SERVICE ALERT: ieatnetsimv9034-118;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.83%
[1576079382] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1576079496] SERVICE ALERT: ieatlms7435;check fm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1576079619] SERVICE ALERT: ieatlms7435;check fm supervision nodes SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1576079741] SERVICE ALERT: ieatlms7435;check fm supervision nodes SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1576080041] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SBG-IS NODES !
[1576080468] SERVICE ALERT: ieatnetsimv9034-114;check netsim cpu;OK;HARD;3;OK- CPU IS USED 14.81%
[1576080902] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576080907] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576081025] HOST ALERT: ieatnetsimv7035-85;DOWN;SOFT;1;PING NOK
[1576081087] HOST ALERT: ieatnetsimv7035-85;UP;SOFT;1;PING OK
[1576081543] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1576081863] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;OK;HARD;3;OK
[1576082192] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN DSC NODES !
[1576082295] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576082636] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2404) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576082767] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1523) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576082898] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (14035)
[1576084506] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1576085146] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (109) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576085270] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.31%
[1576085273] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (109) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576085405] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (112) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576085421] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.93%
[1576085528] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576085572] SERVICE ALERT: ieatnetsimv9034-106;check netsim cpu;OK;HARD;3;OK- CPU IS USED 53.51%
[1576085632] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576085681] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576085784] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576085795] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576085832] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576085901] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576085923] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1576085937] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1576086010] SERVICE ALERT: ieatnetsimv9034-113;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1576086026] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (80)
[1576086047] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN DSC NODES !
[1576086085] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (70) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576086161] SERVICE ALERT: ieatnetsimv9034-113;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.99%
[1576086213] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576086312] SERVICE ALERT: ieatnetsimv9034-113;check netsim cpu;OK;HARD;3;OK- CPU IS USED 34.69%
[1576086340] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576086802] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH CM SUPERVISION (1432) IS BETWEEN 80% AND 95% OF ENM NODES (1750)
[1576087459] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.86%
[1576087610] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1576087725] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1576087760] SERVICE ALERT: ieatnetsimv9034-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.08%
[1576088383] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1576088568] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1576088753] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1576088775] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (11874) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576088909] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13983) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576089043] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12105) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576089178] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (20) IS LOWER THAN 80% OF ENM NODES (40)
[1576089179] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13979) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576089306] SERVICE ALERT: ieatlms3901;check fm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1576089357] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12646) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576089408] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1576089465] SERVICE ALERT: ieatnetsimv9034-115;check netsim cpu;OK;HARD;3;OK- CPU IS USED 42.03%
[1576089495] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13978) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576089567] SERVICE FLAPPING ALERT: ieatnetsimv9034-116;check netsim cpu;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1576089567] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.93%
[1576089586] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1576089680] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576089754] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (14) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576089765] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1576089804] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576089881] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (80)
[1576089928] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (18) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1576089948] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576090053] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1576090108] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2702) IS BETWEEN 5% AND 20% OF ENM NODES (14035)
[1576090238] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3174) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576090374] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3253) IS HIGHER THAN 20% OF ENM NODES (14035)
[1576090931] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576091057] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576091183] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576091226] HOST ALERT: ieatnetsimv7035-74;DOWN;SOFT;1;PING NOK
[1576091288] HOST ALERT: ieatnetsimv7035-74;UP;SOFT;1;PING OK
[1576091334] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1576091439] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1576091625] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (127) IS FROM 80% TO 95% OF ENM  NODES (150)
[1576091757] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1576091897] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1324) IS LOWER THAN 80% OF ENM  NODES (1750)
[1576092030] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1576092157] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1345) IS LOWER THAN 80% OF ENM  NODES (1750)
[1576092282] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1357) IS LOWER THAN 80% OF ENM  NODES (1750)
[1576092408] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1576092609] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (12) IS LOWER THAN 5% OF ENM NODES (1200)
[1576092664] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576092667] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (139) IS FROM 80% TO 95% OF ENM  NODES (150)
[1576092784] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576092795] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1576092905] SERVICE ALERT: ieatnetsimv7035-28;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576092995] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (23) IS LOWER THAN 80% OF ENM  NODES (30)
[1576093553] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576093678] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576093927] HOST ALERT: ieatnetsimv9034-49;DOWN;SOFT;1;PING NOK
[1576093968] HOST ALERT: ieatnetsimv9034-49;UP;SOFT;1;PING OK
[1576093985] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (117) IS LOWER THAN 5% OF ENM NODES (14035)
[1576094222] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1628) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576094352] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1576094778] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (98)
[1576094808] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1576094990] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes BSC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (98) IS HIGHER THAN 95% OF ENM NODES (98)
[1576095029] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12509) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576095163] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12135) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576095298] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13983) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576095308] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576095364] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576095429] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576095507] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576095549] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576095607] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (62) IS LOWER THAN 80% OF ENM  NODES (80)
[1576095640] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1576095733] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (63) IS LOWER THAN 80% OF ENM  NODES (80)
[1576095759] SERVICE ALERT: ieatnetsimv7035-19;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1576095864] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (66) IS FROM 80% TO 95% OF ENM  NODES (80)
[1576095880] SERVICE ALERT: ieatnetsimv7035-19;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1576096000] SERVICE ALERT: ieatnetsimv7035-19;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1576097116] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12514) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576097250] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12484) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576097384] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13983) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576097448] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576097572] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1576097696] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1576098079] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1618) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1576098207] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1742) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1576099198] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13135) IS FROM 80% TO 95% OF ENM  NODES (14035)
[1576099333] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13354) IS HIGHER THAN 95% OF ENM  NODES (14035)
[1576099472] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1576099472] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1576099600] SERVICE ALERT: ieatnetsimv7035-19;check netsim disk;OK;HARD;3;OK
[1576099813] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (120) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576099866] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1576099944] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (78) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1576099994] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1576100068] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (57) IS LOWER THAN 5% OF ENM NODES (1200)
[1576100122] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1576100805] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576100943] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576101068] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1576101195] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1576101641] SERVICE ALERT: ieatnetsimv9034-91;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.14%
[1576101791] SERVICE ALERT: ieatnetsimv9034-91;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 18.19%
[1576102749] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1576104265] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED