[1574123018] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574123154] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574123291] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;HARD;3;OK
[1574123534] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574123777] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;SOFT;2;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574125133] SERVICE FLAPPING ALERT: ieatnetsimv9034-43;check netsim disk;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1574125133] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;OK;HARD;3;OK
[1574125412] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (11) IS HIGHER THAN 20% OF ENM NODES (50)
[1574125549] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574125672] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574127255] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (4319) IS FROM 80% TO 95% OF ENM  NODES (5000)
[1574127379] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (4570) IS FROM 80% TO 95% OF ENM  NODES (5000)
[1574127504] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (4802) IS HIGHER THAN 95% OF ENM  NODES (5000)
[1574128223] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (563) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574128351] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (623) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574128483] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (623) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574129655] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical FRONTHAUL-6080;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1574134006] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1506) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574134132] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1547) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574134262] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1525) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574136634] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574138374] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.2%
[1574139292] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (591) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574140227] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1000
[1574140357] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1000
[1574140451] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574140486] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1000
[1574140527] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574140579] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574140623] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1560) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574140653] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574140686] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1560) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574140710] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574140750] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1560) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574140782] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574140822] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (2560)!
[1574140842] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574140880] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1560) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574140951] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1560) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574140969] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574141085] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (2356) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574141215] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2520) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1574141467] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1712) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574141688] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574142688] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (2560)!
[1574142777] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574142815] SERVICE ALERT: ieatnetsimv5116-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.46%
[1574142965] SERVICE ALERT: ieatnetsimv5116-33;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 12.29%
[1574144096] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1574144397] SERVICE ALERT: ieatlms3901;check cm supervision nodes Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1200) IS LOWER THAN 80% OF ENM NODES (3600)
[1574144523] SERVICE ALERT: ieatlms3901;check cm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1574145660] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical EPG;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574146499] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (621) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574147499] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574147685] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574147870] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574148860] SERVICE ALERT: ieatnetsimv9034-17;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.96%
[1574149011] SERVICE ALERT: ieatnetsimv9034-17;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 11.41%
[1574149736] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574149924] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1574150054] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1574150161] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (589) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574150183] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1574150333] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574150336] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574150456] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574150524] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574150567] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (836) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1574150576] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574150695] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;OK;SOFT;2;OK- NUMBER OF PM FILES (1000) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (1000)
[1574150711] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574151124] SERVICE ALERT: ieatnetsimv9034-36;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574151275] SERVICE ALERT: ieatnetsimv9034-36;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 3.39%
[1574151645] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures ERBS;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1574151903] SERVICE FLAPPING ALERT: ieatnetsimv7035-51;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574152028] SERVICE ALERT: ieatnetsimv9034-38;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.57%
[1574152178] SERVICE ALERT: ieatnetsimv9034-38;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.19%
[1574152579] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574153524] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574153770] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (743) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574153793] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1574155846] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.23%
[1574155910] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical Router6274;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574155996] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.15%
[1574156754] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13198) IS FROM 80% TO 95% OF ENM  NODES (13957)
[1574156895] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13181) IS FROM 80% TO 95% OF ENM  NODES (13957)
[1574157001] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574157090] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (13256) IS FROM 80% TO 95% OF ENM  NODES (13957)
[1574157187] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574157373] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1574157373] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574157377] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (668) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574157513] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (668) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574157592] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574157643] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (619) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574157717] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574157856] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574158079] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574158210] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574158338] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574158906] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13865) IS HIGHER THAN 95% OF ENM  NODES (13957)
[1574159225] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;OK;HARD;3;OK
[1574159287] SERVICE ALERT: ieatlms5735-1;check synchronized ESC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (0)!
[1574159517] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF PM FILES (10000) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (5000)
[1574159593] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574159657] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (0)!
[1574159890] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574159911] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574159912] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574159966] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574160093] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574160093] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574160188] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574160242] SERVICE ALERT: ieatlms5735-1;check synchronized SBG-IS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (0)!
[1574160257] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-IS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (0)!
[1574160273] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574160304] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574160395] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574160460] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (608) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1574160571] SERVICE ALERT: ieatlms5735-1;check fm hb failures EPG;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN EPG NODES !
[1574160582] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574160651] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-DB-BSP;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (0)!
[1574160670] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (0)!
[1574160834] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (0)!
[1574161262] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-DB-BSP;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-DB-BSP NODES !
[1574161940] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1574162723] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (15) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1574162897] SERVICE ALERT: ieatlms5735-1;check synchronized ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (1000)
[1574163024] SERVICE ALERT: ieatlms5735-1;check synchronized ESC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (1000)
[1574163150] SERVICE ALERT: ieatlms5735-1;check synchronized ESC;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (1000)
[1574163296] SERVICE ALERT: ieatnetsimv9034-61;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574163447] SERVICE ALERT: ieatnetsimv9034-61;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 77.41%
[1574163846] SERVICE ALERT: ieatlms5735-1;check synchronized SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (20)
[1574163861] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574163871] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (13)
[1574163935] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574163974] SERVICE ALERT: ieatlms5735-1;check synchronized SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (20)
[1574163996] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (13)
[1574164057] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1574164067] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (608) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1574164098] SERVICE ALERT: ieatlms5735-1;check synchronized SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (20)
[1574164119] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (13)
[1574164186] SERVICE ALERT: ieatlms5735-1;check fm hb failures EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (50) IS HIGHER THAN 20% OF ENM NODES (50)
[1574164194] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (608) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1574164281] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574164290] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (24)
[1574164305] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574164326] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (608) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1574164352] SERVICE ALERT: ieatlms5735-1;check fm hb failures EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (50) IS HIGHER THAN 20% OF ENM NODES (50)
[1574164414] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (24)
[1574164432] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574164442] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (12)
[1574164480] SERVICE ALERT: ieatlms5735-1;check fm hb failures EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (50) IS HIGHER THAN 20% OF ENM NODES (50)
[1574164551] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (24)
[1574164568] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574164570] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (12)
[1574164696] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (12)
[1574164859] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (594) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574164883] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (10) IS HIGHER THAN 20% OF ENM NODES (24)
[1574165010] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (10) IS HIGHER THAN 20% OF ENM NODES (24)
[1574165143] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (10) IS HIGHER THAN 20% OF ENM NODES (24)
[1574166327] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (15) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574166437] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  RadioNode: 1
[1574166454] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (15) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574166582] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (15) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574166587] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  RadioNode: 1
[1574166737] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  RadioNode: 1
[1574168084] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574168235] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.68%
[1574173673] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574173824] SERVICE ALERT: ieatnetsimv9034-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 35.75%
[1574175438] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574178173] SERVICE ALERT: ieatnetsimv9034-84;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.64%
[1574178267] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.78%
[1574178324] SERVICE ALERT: ieatnetsimv9034-84;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 31.88%
[1574178420] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 59.95%
[1574179266] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (613) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574180715] SERVICE ALERT: ieatlms3901;check fm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (150) IS LOWER THAN 80% OF ENM NODES (450)
[1574180846] SERVICE ALERT: ieatlms3901;check fm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574180875] SERVICE ALERT: ieatnetsimv9034-98;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.78%
[1574181026] SERVICE ALERT: ieatnetsimv9034-98;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 32.67%
[1574182873] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (572) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574183761] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (80) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (240)
[1574183888] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1574183948] SERVICE ALERT: ieatnetsimv9034-119;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574184099] SERVICE ALERT: ieatnetsimv9034-119;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 48.28%
[1574186527] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (7) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574186580] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (4126) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13964)
[1574186623] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (600) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574186656] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (23) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574186688] SERVICE ALERT: ieatlms3901;check fm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1574186689] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 274 processes with STATE = RSZDT
[1574186749] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 288 processes with STATE = RSZDT
[1574186752] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (13358) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1574186781] SERVICE ALERT: ieatlms3901;check synchronized BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (47) IS LOWER THAN 80% OF ENM  NODES (96)
[1574186809] SERVICE ALERT: localhost;Total Processes;OK;SOFT;3;PROCS OK: 235 processes with STATE = RSZDT
[1574186821] SERVICE ALERT: ieatlms3901;check fm hb failures SBG-IS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574186822] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574186823] SERVICE ALERT: ieatlms3901;check fm supervision nodes Router6672;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574186826] SERVICE ALERT: ieatlms3901;check fm supervision nodes RadioNode;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1574186836] SERVICE ALERT: ieatlms3901;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574186843] SERVICE ALERT: ieatlms3901;check cm supervision nodes Router6672;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574186850] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1574186856] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574186862] SERVICE ALERT: ieatlms3901;check pm supervision nodes MGW;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574186870] SERVICE ALERT: ieatlms3901;check synchronized RBS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574186873] SERVICE ALERT: ieatlms3901;check unsync nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1574186876] SERVICE ALERT: ieatlms3901;check fm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1574186879] SERVICE ALERT: ieatlms3901;check cm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574186904] SERVICE ALERT: ieatlms3901;check fm supervision nodes MTAS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574186929] SERVICE ALERT: ieatlms3901;check fm supervision nodes MSC-BC-IS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574186950] SERVICE ALERT: ieatlms3901;check fm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1574187007] SERVICE ALERT: ieatlms3901;check fm supervision nodes RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (13964) IS HIGHER THAN 95% OF ENM NODES (13965)
[1574187011] SERVICE ALERT: ieatlms3901;check synchronized RBS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (3849)!
[1574187015] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574187015] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1574187023] SERVICE ALERT: ieatlms3901;check cm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1574187025] SERVICE ALERT: ieatlms3901;check fm supervision nodes RBS;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1574187028] SERVICE ALERT: ieatlms3901;check cm supervision nodes RNC;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574187031] SERVICE ALERT: ieatlms3901;check synchronized BSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (86) IS FROM 80% TO 95% OF ENM  NODES (96)
[1574187050] SERVICE ALERT: ieatlms3901;check fm supervision nodes ERBS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574187053] SERVICE ALERT: ieatlms3901;check fm hb failures SBG-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SBG-IS NODES !
[1574187066] SERVICE ALERT: ieatlms3901;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1574187077] SERVICE ALERT: ieatlms3901;check fm supervision nodes MSC-BC-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1574187081] SERVICE ALERT: ieatlms3901;check cm supervision nodes MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1750) IS HIGHER THAN 95% OF ENM NODES (1750)
[1574187085] SERVICE ALERT: ieatlms3901;check cm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574187087] SERVICE ALERT: ieatlms3901;check pm supervision nodes MGW;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574187100] SERVICE ALERT: ieatlms3901;check pm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574187118] SERVICE ALERT: ieatlms3901;check fm supervision nodes MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574187161] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (2560)!
[1574187183] SERVICE ALERT: ieatlms3901;check synchronized BSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (89) IS FROM 80% TO 95% OF ENM  NODES (96)
[1574187188] SERVICE ALERT: ieatlms3901;check fm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574187217] SERVICE ALERT: ieatlms3901;check fm supervision nodes RBS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (3849) IS HIGHER THAN 95% OF ENM NODES (3849)
[1574187230] SERVICE ALERT: ieatlms3901;check cm supervision nodes RNC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (6) IS HIGHER THAN 95% OF ENM NODES (6)
[1574187322] SERVICE ALERT: ieatlms3901;check pm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574188733] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574189105] SERVICE ALERT: ieatlms3901;check synchronized BSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (95) IS HIGHER THAN 95% OF ENM  NODES (96)
[1574190239] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (573) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574190368] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (573) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574190499] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (638) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574193504] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1574197376] SERVICE FLAPPING ALERT: ieatnetsimv9034-43;check netsim disk;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1574200219] SERVICE ALERT: ieatlms3901;check pm supervision nodes MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (12) IS LOWER THAN 80% OF ENM NODES (36)
[1574200344] SERVICE ALERT: ieatlms3901;check pm supervision nodes MSC-BC-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1574203208] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (150) IS LOWER THAN 80% OF ENM NODES (300)
[1574203334] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MGW;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574205720] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (24) IS LOWER THAN 80% OF ENM NODES (72)
[1574205845] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)