[1573087878] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1573094160] HOST ALERT: ieatnetsimv017-16;DOWN;SOFT;1;PING NOK
[1573094222] HOST ALERT: ieatnetsimv017-16;UP;SOFT;1;PING OK
[1573095595] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12731) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573095635] HOST ALERT: ieatnetsimv5116-02;DOWN;SOFT;1;PING NOK
[1573095697] HOST ALERT: ieatnetsimv5116-02;UP;SOFT;1;PING OK
[1573095738] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12514) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573095881] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12700) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573096955] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (3461) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573097033] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12369) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13958)
[1573097084] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (3461) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573097135] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573097177] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573097182] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12369) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13958)
[1573097212] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (3461) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573097260] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573097304] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573097328] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12369) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13958)
[1573097374] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2490) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573097384] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573097434] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573097743] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573097869] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573097995] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573098184] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573098316] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;SOFT;2;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573099503] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573100821] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RBS;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1573100821] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (3814) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573100954] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1573100954] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13574) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573100989] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573101043] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (5) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573101064] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573101603] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573102040] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573102165] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573102290] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573102317] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2284) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573102444] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2284) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573102572] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2284) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573103103] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (729) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573103129] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12850) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573103270] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12850) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573103412] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12850) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573104581] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2286) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573104671] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573104710] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (2038) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573104796] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573104837] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (2038) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573104922] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573105895] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573105934] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573105945] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1573106180] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573106214] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1573106539] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573106978] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1573107027] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573107253] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1573107604] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1573108256] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573108379] SERVICE ALERT: ieatnetsimv5116-44;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573108383] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573108447] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2286) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573108484] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573108513] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573108515] SERVICE ALERT: ieatnetsimv5116-44;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1573108810] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573108851] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1573108937] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573109063] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573109391] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573109531] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573109667] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573110973] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573112054] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573112122] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RNC;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1573112122] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (5) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573112137] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573112670] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1573112670] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573113101] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573113225] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573113350] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573114002] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573114233] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573114470] SERVICE ALERT: ieatlms3901;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1573114518] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573114723] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1573114793] SERVICE ALERT: ieatlms3901;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1573114828] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573114835] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1573114935] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573115011] SERVICE ALERT: ieatlms3901;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1573115151] SERVICE ALERT: ieatlms3901;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1573115993] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1573115993] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573116341] HOST ALERT: ieatnetsimv017-27;DOWN;SOFT;1;PING NOK
[1573116404] HOST ALERT: ieatnetsimv017-27;UP;SOFT;1;PING OK
[1573118441] HOST ALERT: ieatnetsimv5116-62;DOWN;SOFT;1;PING NOK
[1573118503] HOST ALERT: ieatnetsimv5116-62;UP;SOFT;1;PING OK
[1573119261] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2286) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573119343] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573119390] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2286) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573119469] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573119517] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573119595] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573122215] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.89%
[1573122365] SERVICE ALERT: ieatnetsimv5116-48;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.88%
[1573123204] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573125054] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13054) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573125202] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13054) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573125345] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573126369] SERVICE ALERT: ieatnetsimv9041-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573126515] SERVICE ALERT: ieatnetsimv9041-10;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573126654] SERVICE ALERT: ieatnetsimv9041-10;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573128485] SERVICE ALERT: ieatnetsimv9041-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130021] SERVICE ALERT: ieatnetsimv7035-59;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_netsim_nodes.sh, ...) failed. errno is 2: No such file or directory
[1573130024] SERVICE ALERT: ieatnetsimv7035-60;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_netsim_nodes.sh, ...) failed. errno is 2: No such file or directory
[1573130025] SERVICE ALERT: ieatnetsimv7035-61;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_netsim_nodes.sh, ...) failed. errno is 2: No such file or directory
[1573130062] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1573130062] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (28.5% change >= 20.0% threshold)
[1573130062] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RNC;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1573130062] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (41.6% change >= 20.0% threshold)
[1573130062] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (36.1% change >= 20.0% threshold)
[1573130062] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (33.5% change >= 20.0% threshold)
[1573130063] SERVICE ALERT: ieatnetsimv9041-16;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130141] SERVICE ALERT: ieatnetsimv7035-59;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1573130144] SERVICE ALERT: ieatnetsimv7035-60;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1573130145] SERVICE ALERT: ieatnetsimv7035-61;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1573130158] SERVICE ALERT: ieatnetsimv9041-34;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130250] SERVICE ALERT: ieatnetsimv9041-15;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130252] SERVICE ALERT: ieatnetsimv9041-21;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130257] SERVICE ALERT: ieatnetsimv9041-27;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130264] SERVICE ALERT: ieatnetsimv9041-24;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130266] SERVICE ALERT: ieatnetsimv9041-18;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130268] SERVICE ALERT: ieatnetsimv9041-30;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130273] SERVICE ALERT: ieatnetsimv9041-33;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130335] SERVICE ALERT: ieatnetsimv9041-31;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573130680] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.17%
[1573130831] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 17.82%
[1573131361] HOST ALERT: ieatnetsimv5116-17;DOWN;SOFT;1;PING NOK
[1573131385] SERVICE ALERT: ieatnetsimv9041-25;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573131388] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2284) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573131423] HOST ALERT: ieatnetsimv5116-17;UP;SOFT;1;PING OK
[1573131448] SERVICE ALERT: ieatnetsimv9041-20;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573131457] SERVICE ALERT: ieatnetsimv9041-17;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573131475] SERVICE ALERT: ieatnetsimv9041-23;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573131478] SERVICE ALERT: ieatnetsimv9041-14;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573131479] SERVICE ALERT: ieatnetsimv9041-32;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573131485] SERVICE ALERT: ieatnetsimv9041-29;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573131488] SERVICE ALERT: ieatnetsimv9041-22;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573131495] SERVICE ALERT: ieatnetsimv9041-19;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573131516] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2284) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573131643] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573132238] HOST ALERT: ieatnetsimv5116-112;DOWN;SOFT;1;PING NOK
[1573132241] HOST ALERT: ieatnetsimv5116-66;DOWN;SOFT;1;PING NOK
[1573132300] HOST ALERT: ieatnetsimv5116-112;UP;SOFT;1;PING OK
[1573132303] HOST ALERT: ieatnetsimv5116-66;UP;SOFT;1;PING OK
[1573134233] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1573134233] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (27.2% change >= 20.0% threshold)
[1573134233] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RNC;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1573134233] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (39.5% change >= 20.0% threshold)
[1573134233] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (34.3% change >= 20.0% threshold)
[1573134233] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (32.6% change >= 20.0% threshold)
[1573134258] SERVICE ALERT: ieatnetsimv017-21;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134306] SERVICE ALERT: ieatnetsimv5116-46;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134307] SERVICE ALERT: ieatnetsimv5116-52;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134317] SERVICE ALERT: ieatnetsimv5116-117;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134329] SERVICE ALERT: ieatnetsimv5116-49;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134373] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134378] SERVICE ALERT: ieatnetsimv017-21;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134426] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134426] SERVICE ALERT: ieatnetsimv5116-46;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134427] SERVICE ALERT: ieatnetsimv5116-52;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134429] SERVICE ALERT: ieatnetsimv017-32;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134429] SERVICE ALERT: ieatnetsimv017-33;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134432] SERVICE ALERT: ieatnetsimv017-34;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134434] SERVICE ALERT: ieatnetsimv017-35;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134437] SERVICE ALERT: ieatnetsimv5116-117;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134437] SERVICE ALERT: ieatnetsimv017-36;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134440] SERVICE ALERT: ieatnetsimv5116-01;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134441] SERVICE ALERT: ieatnetsimv5116-02;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134443] SERVICE ALERT: ieatnetsimv5116-03;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134449] SERVICE ALERT: ieatnetsimv5116-49;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134450] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134452] SERVICE ALERT: ieatnetsimv5116-07;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134455] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134459] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134459] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134469] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134482] SERVICE ALERT: ieatnetsimv017-23;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134484] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134493] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134498] SERVICE ALERT: ieatnetsimv017-21;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134509] SERVICE ALERT: ieatnetsimv017-26;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134512] SERVICE ALERT: ieatnetsimv5116-119;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134513] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134515] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134546] SERVICE ALERT: ieatnetsimv5116-46;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134546] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134547] SERVICE ALERT: ieatnetsimv5116-52;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134549] SERVICE ALERT: ieatnetsimv017-32;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134549] SERVICE ALERT: ieatnetsimv017-33;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134552] SERVICE ALERT: ieatnetsimv017-34;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134554] SERVICE ALERT: ieatnetsimv017-35;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134557] SERVICE ALERT: ieatnetsimv5116-117;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134557] SERVICE ALERT: ieatnetsimv017-36;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134560] SERVICE ALERT: ieatnetsimv5116-01;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134561] SERVICE ALERT: ieatnetsimv5116-02;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134563] SERVICE ALERT: ieatnetsimv5116-03;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134569] SERVICE ALERT: ieatnetsimv5116-49;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134570] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134572] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134572] SERVICE ALERT: ieatnetsimv5116-07;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134575] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134577] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134579] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134579] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134589] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134602] SERVICE ALERT: ieatnetsimv017-23;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134604] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134613] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134617] SERVICE ALERT: ieatnetsimv5116-29;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134629] SERVICE ALERT: ieatnetsimv017-26;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134632] SERVICE ALERT: ieatnetsimv5116-119;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134633] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134635] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134666] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134669] SERVICE ALERT: ieatnetsimv017-32;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134669] SERVICE ALERT: ieatnetsimv017-33;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134672] SERVICE ALERT: ieatnetsimv017-34;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134674] SERVICE ALERT: ieatnetsimv017-35;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134677] SERVICE ALERT: ieatnetsimv017-36;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134680] SERVICE ALERT: ieatnetsimv5116-01;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134681] SERVICE ALERT: ieatnetsimv5116-02;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134683] SERVICE ALERT: ieatnetsimv5116-03;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134690] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134692] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134692] SERVICE ALERT: ieatnetsimv5116-07;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134695] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134697] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134699] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134699] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134709] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134722] SERVICE ALERT: ieatnetsimv017-23;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134724] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134736] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134737] SERVICE ALERT: ieatnetsimv5116-29;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134749] SERVICE ALERT: ieatnetsimv017-26;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134752] SERVICE ALERT: ieatnetsimv5116-119;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134753] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134755] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134812] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134817] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134856] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134857] SERVICE ALERT: ieatnetsimv5116-32;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134857] SERVICE ALERT: ieatnetsimv5116-29;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134873] SERVICE ALERT: ieatnetsimv5116-48;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134894] SERVICE ALERT: ieatnetsimv017-24;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134976] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134977] SERVICE ALERT: ieatnetsimv5116-32;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134983] SERVICE ALERT: ieatnetsimv5116-51;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134991] SERVICE ALERT: ieatnetsimv017-27;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573134993] SERVICE ALERT: ieatnetsimv5116-48;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573135014] SERVICE ALERT: ieatnetsimv017-24;check netsim nodes;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573135032] SERVICE ALERT: ieatnetsimv017-19;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/s4/check_netsim_nodes_wkl.sh, ...) failed. errno is 2: No such file or directory
[1573135048] SERVICE ALERT: ieatnetsimv9041-20;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_netsim_nodes.sh, ...) failed. errno is 2: No such file or directory
[1573135057] SERVICE ALERT: ieatnetsimv9041-17;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_netsim_nodes.sh, ...) failed. errno is 2: No such file or directory
[1573135075] SERVICE ALERT: ieatnetsimv9041-23;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_netsim_nodes.sh, ...) failed. errno is 2: No such file or directory
[1573135078] SERVICE ALERT: ieatnetsimv9041-14;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_netsim_nodes.sh, ...) failed. errno is 2: No such file or directory
[1573135079] SERVICE ALERT: ieatnetsimv9041-32;check netsim nodes;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_netsim_nodes.sh, ...) failed. errno is 2: No such file or directory
[1573135115] SERVICE ALERT: ieatnetsimv9041-29;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135125] SERVICE ALERT: ieatnetsimv9041-22;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135129] SERVICE ALERT: ieatnetsimv9041-19;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135140] SERVICE ALERT: ieatnetsimv5116-32;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573135141] SERVICE ALERT: ieatnetsimv5116-51;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1573135143] SERVICE ALERT: ieatnetsimv017-27;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1573135147] SERVICE ALERT: ieatnetsimv5116-48;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573135163] SERVICE ALERT: ieatnetsimv017-24;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573135197] SERVICE ALERT: ieatnetsimv9041-20;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135197] SERVICE ALERT: ieatnetsimv017-19;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1573135242] SERVICE ALERT: ieatnetsimv9041-32;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135248] SERVICE ALERT: ieatnetsimv9041-17;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135251] SERVICE ALERT: ieatnetsimv9041-14;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135252] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2232) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573135252] SERVICE ALERT: ieatnetsimv9041-23;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135260] SERVICE ALERT: ieatnetsimv9041-29;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135271] SERVICE ALERT: ieatnetsimv9041-19;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135272] SERVICE ALERT: ieatnetsimv9041-22;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135347] SERVICE ALERT: ieatnetsimv9041-20;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135379] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2328) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573135390] SERVICE ALERT: ieatnetsimv9041-32;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135399] SERVICE ALERT: ieatnetsimv9041-17;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135400] SERVICE ALERT: ieatnetsimv9041-14;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135405] SERVICE ALERT: ieatnetsimv9041-29;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135408] SERVICE ALERT: ieatnetsimv9041-23;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135422] SERVICE ALERT: ieatnetsimv9041-19;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135423] SERVICE ALERT: ieatnetsimv9041-22;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135500] SERVICE ALERT: ieatnetsimv9041-16;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135508] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2434) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573135540] SERVICE ALERT: ieatnetsimv017-21;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573135564] SERVICE ALERT: ieatnetsimv017-23;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573135595] SERVICE ALERT: ieatnetsimv9041-34;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135646] SERVICE ALERT: ieatnetsimv9041-16;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135677] SERVICE ALERT: ieatnetsimv9041-21;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135683] SERVICE ALERT: ieatnetsimv9041-27;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135685] SERVICE ALERT: ieatnetsimv9041-15;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135695] SERVICE ALERT: ieatnetsimv9041-24;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135701] SERVICE ALERT: ieatnetsimv9041-30;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135703] SERVICE ALERT: ieatnetsimv9041-33;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135703] SERVICE ALERT: ieatnetsimv9041-18;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135741] SERVICE ALERT: ieatnetsimv9041-34;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135769] SERVICE ALERT: ieatnetsimv9041-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573135796] SERVICE ALERT: ieatnetsimv9041-16;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135818] SERVICE ALERT: ieatnetsimv9041-21;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135828] SERVICE ALERT: ieatnetsimv9041-15;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135829] SERVICE ALERT: ieatnetsimv9041-27;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135840] SERVICE ALERT: ieatnetsimv9041-24;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135844] SERVICE ALERT: ieatnetsimv9041-18;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135852] SERVICE ALERT: ieatnetsimv9041-30;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135853] SERVICE ALERT: ieatnetsimv9041-33;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135882] SERVICE ALERT: ieatnetsimv9041-34;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135913] SERVICE ALERT: ieatnetsimv9041-31;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573135970] SERVICE ALERT: ieatnetsimv9041-21;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135977] SERVICE ALERT: ieatnetsimv9041-15;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135985] SERVICE ALERT: ieatnetsimv9041-27;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135990] SERVICE ALERT: ieatnetsimv9041-18;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573135991] SERVICE ALERT: ieatnetsimv9041-24;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573136001] SERVICE ALERT: ieatnetsimv9041-33;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573136003] SERVICE ALERT: ieatnetsimv9041-30;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573136057] SERVICE ALERT: ieatnetsimv9041-31;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573136382] SERVICE ALERT: ieatnetsimv5116-117;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136384] SERVICE ALERT: ieatnetsimv5116-46;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136385] SERVICE ALERT: ieatnetsimv5116-52;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136396] SERVICE ALERT: ieatnetsimv5116-11;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136407] SERVICE ALERT: ieatnetsimv5116-49;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136452] SERVICE FLAPPING ALERT: ieatnetsimv5116-45;check netsim nodes;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1573136452] SERVICE ALERT: ieatnetsimv5116-45;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136508] SERVICE ALERT: ieatnetsimv5116-01;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136510] SERVICE ALERT: ieatnetsimv5116-02;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136511] SERVICE ALERT: ieatnetsimv5116-03;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136524] SERVICE ALERT: ieatnetsimv017-34;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136525] SERVICE ALERT: ieatnetsimv017-32;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136527] SERVICE ALERT: ieatnetsimv017-33;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136531] SERVICE ALERT: ieatnetsimv5116-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136533] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136534] SERVICE ALERT: ieatnetsimv017-36;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136534] SERVICE ALERT: ieatnetsimv5116-07;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136535] SERVICE ALERT: ieatnetsimv017-35;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136541] SERVICE ALERT: ieatnetsimv5116-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136541] SERVICE ALERT: ieatnetsimv5116-08;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136556] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136575] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136576] SERVICE ALERT: ieatnetsimv017-26;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136589] SERVICE ALERT: ieatnetsimv5116-13;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136590] SERVICE ALERT: ieatnetsimv5116-119;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136596] SERVICE ALERT: ieatnetsimv5116-16;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136653] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136659] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136686] SERVICE ALERT: ieatnetsimv5116-29;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136812] SERVICE ALERT: ieatnetsimv9041-25;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573136816] SERVICE ALERT: ieatnetsimv5116-35;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1573136955] SERVICE ALERT: ieatnetsimv9041-25;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1573137107] SERVICE ALERT: ieatnetsimv9041-25;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1573138261] HOST ALERT: ieatnetsimv5116-26;DOWN;SOFT;1;PING NOK
[1573138323] HOST ALERT: ieatnetsimv5116-26;UP;SOFT;1;PING OK
[1573138898] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573139080] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573139261] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573139749] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13129) IS FROM 80% TO 95% OF ENM  NODES (13956)
[1573139768] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12307) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573139889] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12832) IS FROM 80% TO 95% OF ENM  NODES (13956)
[1573139912] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12433) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573140032] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12821) IS FROM 80% TO 95% OF ENM  NODES (13956)
[1573140055] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12433) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573140089] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573140185] SERVICE ALERT: ieatlms3901;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1573140193] SERVICE ALERT: ieatlms3901;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1573140270] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573140309] SERVICE ALERT: ieatlms3901;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1573140321] SERVICE ALERT: ieatlms3901;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1573140432] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (3456) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573140435] SERVICE ALERT: ieatlms3901;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1573140445] SERVICE ALERT: ieatlms3901;check synchronized MTAS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1573140451] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573140560] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (3456) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573140582] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12550) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573140598] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573140691] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (3464) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573140724] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573140730] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12583) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573140849] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573140876] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12773) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573141475] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (133) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573141536] HOST ALERT: ieatnetsimv5116-70;DOWN;SOFT;1;PING NOK
[1573141540] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573141598] HOST ALERT: ieatnetsimv5116-70;UP;SOFT;1;PING OK
[1573141602] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573141690] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 10.88%
[1573141860] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13649) IS HIGHER THAN 95% OF ENM  NODES (13956)
[1573142239] SERVICE ALERT: ieatlms3901;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1573142248] SERVICE ALERT: ieatlms3901;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1573142746] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.1%
[1573142896] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.97%
[1573143677] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573144722] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1573144811] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573144936] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573145062] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573145912] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573146316] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2290) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573146321] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573146465] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2290) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573146479] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573146593] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2290) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573146605] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573147903] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RBS;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1573147903] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (3846) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573148055] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical DSC;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1573148055] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573148104] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1573148104] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13952) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573148279] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573148477] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573148666] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573148674] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1573148736] HOST ALERT: ieatnetsimv5116-104;DOWN;SOFT;1;PING NOK
[1573148798] HOST ALERT: ieatnetsimv5116-104;UP;SOFT;1;PING OK
[1573150201] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1573150201] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573150211] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1573151513] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (3463) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573151641] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (3463) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573151660] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573151730] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12549) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573151770] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (3463) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573151787] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573151882] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12549) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573151913] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573152003] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573152032] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12549) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573152271] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573152352] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573152396] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573152521] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573152527] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573152653] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (143) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573154158] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.33%
[1573154308] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 23.66%
[1573154498] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12433) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573154641] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12433) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573154785] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12433) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573155379] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (3846) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573155608] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573155658] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13952) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573155735] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573155860] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573156125] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573156278] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1573156399] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1573156519] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1573157409] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2307) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573157421] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573157537] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2307) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573157547] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573157665] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2307) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573157673] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573158408] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573159464] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573162030] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13096) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573162172] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13096) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573162315] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12628) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573164873] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (2046) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573165941] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573166131] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573166190] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (3481) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573166263] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573166320] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573166394] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573166449] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (3481) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573166488] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12648) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573166637] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12648) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573166789] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13952) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13956)
[1573168482] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2270) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573169562] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12416) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573169703] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12515) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573169845] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12515) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1573170002] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (5) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (6)
[1573170058] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (3846) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1573170268] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (136) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573170396] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1573170429] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573170625] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573170687] SERVICE FLAPPING ALERT: ieatnetsimv5116-45;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573170822] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED