[1574208177] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;OK;HARD;3;OK
[1574208527] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1) IS LOWER THAN 80% OF ENM NODES (2560)
[1574208656] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574213460] SERVICE ALERT: ieatnetsimv7102-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.12%
[1574213611] SERVICE ALERT: ieatnetsimv7102-71;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 73.61%
[1574220673] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1630) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574220799] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1620) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574220969] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1652) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574222799] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1664) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574228185] HOST ALERT: ieatnetsimv7035-58;DOWN;SOFT;1;PING NOK
[1574228247] HOST ALERT: ieatnetsimv7035-58;UP;SOFT;1;PING OK
[1574228732] HOST ALERT: ieatnetsimv7035-40;DOWN;SOFT;1;PING NOK
[1574228735] HOST ALERT: ieatnetsimv7035-82;DOWN;SOFT;1;PING NOK
[1574228740] HOST ALERT: ieatnetsimv7035-89;DOWN;SOFT;1;PING NOK
[1574228765] HOST ALERT: ieatnetsimv7035-09;DOWN;SOFT;1;PING NOK
[1574228795] HOST ALERT: ieatnetsimv7035-40;UP;SOFT;1;PING OK
[1574228797] HOST ALERT: ieatnetsimv7035-82;UP;SOFT;1;PING OK
[1574228802] HOST ALERT: ieatnetsimv7035-89;UP;SOFT;1;PING OK
[1574228828] HOST ALERT: ieatnetsimv7035-09;UP;SOFT;1;PING OK
[1574229881] HOST ALERT: ieatnetsimv7048-11;DOWN;SOFT;1;PING NOK
[1574229943] HOST ALERT: ieatnetsimv7048-11;UP;SOFT;1;PING OK
[1574231303] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574231489] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574231675] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574233068] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1574233765] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574233917] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574234002] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (28) IS LOWER THAN 80% OF ENM  NODES (98)
[1574234068] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.74%
[1574234105] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 4
[1574234125] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (97) IS HIGHER THAN 95% OF ENM  NODES (98)
[1574234236] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  RBS: 6 ERBS: 4 MGW: 3
[1574234351] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (3320) IS LOWER THAN 80% OF ENM NODES (5000)
[1574234363] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  RBS: 6 ERBS: 3
[1574234476] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (3640) IS LOWER THAN 80% OF ENM NODES (5000)
[1574234494] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2130) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574234601] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (3840) IS LOWER THAN 80% OF ENM NODES (5000)
[1574234622] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (2383) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574234753] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2559) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1574234780] SERVICE ALERT: ieatnetsimv9041-08;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574234793] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.53%
[1574234945] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.16%
[1574235096] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.06%
[1574235508] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (4820) IS HIGHER THAN 95% OF ENM NODES (5000)
[1574235588] SERVICE ALERT: ieatnetsimv5116-40;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574235591] SERVICE ALERT: ieatnetsimv5116-41;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574235739] SERVICE ALERT: ieatnetsimv5116-40;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 74.2%
[1574235742] SERVICE ALERT: ieatnetsimv5116-41;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574235894] SERVICE ALERT: ieatnetsimv5116-41;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.92%
[1574236454] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574236559] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2400) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574236572] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574236607] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574236684] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (2063) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574236749] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (123) IS FROM 80% TO 95% OF ENM  NODES (150)
[1574236758] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574236810] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (2098) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574236878] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (147) IS HIGHER THAN 95% OF ENM  NODES (150)
[1574237326] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574237423] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574237452] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574237575] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574237578] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574237699] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;HARD;3;OK- CPU IS USED 26.43%
[1574237702] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574237726] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.48%
[1574237826] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1574238123] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13148) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574238257] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13305) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574238355] SERVICE ALERT: ieatnetsimv9041-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.13%
[1574238617] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1438) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574238674] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1085) IS LOWER THAN 80% OF ENM  NODES (1514)
[1574238726] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.71%
[1574238801] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1102) IS LOWER THAN 80% OF ENM  NODES (1514)
[1574238805] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (100) IS LOWER THAN 80% OF ENM  NODES (150)
[1574238862] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574238940] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1155) IS LOWER THAN 80% OF ENM  NODES (1514)
[1574238941] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574238942] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (110) IS LOWER THAN 80% OF ENM  NODES (150)
[1574238989] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (79) IS HIGHER THAN 95% OF ENM  NODES (80)
[1574239068] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (92) IS LOWER THAN 80% OF ENM  NODES (150)
[1574239524] SERVICE ALERT: ieatnetsimv5116-41;check netsim cpu;OK;HARD;3;OK- CPU IS USED 69.22%
[1574240389] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 72.54%
[1574240428] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2519) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1574240577] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574240698] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574240754] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1362) IS FROM 80% TO 95% OF ENM  NODES (1514)
[1574240818] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574240874] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (136) IS FROM 80% TO 95% OF ENM  NODES (150)
[1574241356] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;OK;HARD;3;OK- CPU IS USED 68.22%
[1574242357] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.28%
[1574242508] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.31%
[1574242564] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1510) IS HIGHER THAN 95% OF ENM  NODES (1514)
[1574242646] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574242659] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.65%
[1574242679] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1574242809] SERVICE ALERT: ieatnetsimv5116-37;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.66%
[1574242834] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574243021] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574244036] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1751) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574244162] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1931) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574244289] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized ERBS;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1574244289] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (2222) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574244507] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (723) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574244516] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574244645] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574244770] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (79) IS HIGHER THAN 95% OF ENM  NODES (80)
[1574245033] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574245167] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574245195] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1574245292] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574245416] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574246102] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1929) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574246576] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574246704] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (77) IS HIGHER THAN 95% OF ENM  NODES (80)
[1574247347] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574247472] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574247930] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (2560)!
[1574248460] SERVICE ALERT: ieatlms3901;check pm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1574248488] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574248585] SERVICE ALERT: ieatlms3901;check pm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1574249746] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.33%
[1574249897] SERVICE ALERT: ieatnetsimv5116-46;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 73.61%
[1574250354] SERVICE ALERT: ieatnetsimv5116-41;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.18%
[1574250505] SERVICE ALERT: ieatnetsimv5116-41;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 54.59%
[1574251220] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.98%
[1574251371] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.97%
[1574251521] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.79%
[1574251699] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12980) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13964)
[1574251714] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (662) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574251841] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (662) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574251847] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12980) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13964)
[1574252006] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12980) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13964)
[1574252008] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (662) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574253320] SERVICE ALERT: ieatnetsimv7035-59;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574253479] SERVICE ALERT: ieatnetsimv7035-59;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574253630] SERVICE ALERT: ieatnetsimv7035-59;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574254737] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574254931] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574255077] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574255207] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574255389] SERVICE ALERT: ieatnetsimv7035-60;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574255392] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574255404] SERVICE ALERT: ieatnetsimv7035-49;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574255466] SERVICE ALERT: ieatnetsimv7035-59;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574255532] SERVICE ALERT: ieatnetsimv7035-63;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574255540] SERVICE ALERT: ieatnetsimv7035-60;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574255554] SERVICE ALERT: ieatnetsimv7035-49;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574255577] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574255593] SERVICE ALERT: ieatnetsimv7035-53;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574255637] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13939) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13964)
[1574255685] SERVICE ALERT: ieatnetsimv7035-63;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574255711] SERVICE ALERT: ieatnetsimv7035-49;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574255738] SERVICE ALERT: ieatnetsimv7035-53;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574255832] SERVICE ALERT: ieatnetsimv7035-63;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574255856] SERVICE ALERT: ieatnetsimv7035-49;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574255890] SERVICE ALERT: ieatnetsimv7035-53;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574256030] SERVICE ALERT: ieatnetsimv7035-53;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574256032] SERVICE ALERT: ieatnetsimv7035-74;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574256187] SERVICE ALERT: ieatnetsimv7035-74;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574256236] SERVICE ALERT: ieatnetsimv7035-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574256275] SERVICE ALERT: ieatnetsimv7035-86;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574256340] SERVICE ALERT: ieatnetsimv7035-74;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574256385] SERVICE ALERT: ieatnetsimv7035-68;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574256429] SERVICE ALERT: ieatnetsimv7035-86;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574256466] SERVICE ALERT: ieatnetsimv7035-87;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574256475] SERVICE ALERT: ieatnetsimv7035-78;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574256477] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574256614] SERVICE ALERT: ieatnetsimv7035-87;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574256631] SERVICE ALERT: ieatnetsimv7035-78;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574256632] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574256764] SERVICE ALERT: ieatnetsimv7035-87;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574256784] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574256786] SERVICE ALERT: ieatnetsimv7035-78;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574256797] SERVICE ALERT: ieatnetsimv7035-75;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574256907] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574256914] SERVICE ALERT: ieatnetsimv7035-87;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574256940] SERVICE ALERT: ieatnetsimv7035-75;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574257009] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13157) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13991)
[1574257067] SERVICE ALERT: ieatnetsimv7035-87;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574257098] SERVICE ALERT: ieatnetsimv7035-75;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574257139] SERVICE ALERT: ieatnetsimv7035-80;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574257140] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13991)
[1574257175] SERVICE ALERT: ieatnetsimv7035-90;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574257282] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13751) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13991)
[1574257293] SERVICE ALERT: ieatnetsimv7035-80;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1574257319] SERVICE ALERT: ieatnetsimv7035-90;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574257437] SERVICE ALERT: ieatnetsimv7035-80;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1574257443] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1574257443] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574257548] SERVICE ALERT: ieatnetsimv7035-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1574257636] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574257664] SERVICE ALERT: ieatnetsimv7035-63;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574257699] SERVICE ALERT: ieatnetsimv7035-77;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1574257761] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;SOFT;2;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574258170] SERVICE ALERT: ieatnetsimv7035-74;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574258196] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12005) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574258329] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12250) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574258470] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12393) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574258624] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574258752] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.95%
[1574258900] SERVICE ALERT: ieatnetsimv7035-87;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574258903] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79.72%
[1574258925] SERVICE ALERT: ieatnetsimv7035-75;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574259228] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (750)
[1574259261] SERVICE ALERT: ieatnetsimv7035-80;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574259318] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574259358] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1574259516] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574259712] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574260285] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13901) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574260990] SERVICE ALERT: ieatnetsimv017-01;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574261081] SERVICE ALERT: ieatlms3901;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574261126] SERVICE ALERT: ieatnetsimv017-02;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574261206] SERVICE ALERT: ieatlms3901;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574261220] SERVICE ALERT: ieatnetsimv017-03;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574261339] SERVICE ALERT: ieatlms3901;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574261940] SERVICE ALERT: ieatnetsimv017-04;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262023] SERVICE ALERT: ieatnetsimv017-05;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262034] SERVICE ALERT: ieatnetsimv017-06;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262042] SERVICE ALERT: ieatnetsimv017-07;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262051] SERVICE ALERT: ieatnetsimv017-08;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262061] SERVICE ALERT: ieatnetsimv017-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262182] SERVICE ALERT: ieatnetsimv017-15;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262190] SERVICE ALERT: ieatnetsimv017-16;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262201] SERVICE ALERT: ieatnetsimv5116-100;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262219] SERVICE ALERT: ieatnetsimv5116-101;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262227] SERVICE ALERT: ieatnetsimv5116-102;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262247] SERVICE ALERT: ieatnetsimv5116-103;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262264] SERVICE ALERT: ieatnetsimv5116-104;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262278] SERVICE ALERT: ieatnetsimv5116-105;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262286] SERVICE ALERT: ieatnetsimv017-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262295] SERVICE ALERT: ieatnetsimv017-11;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262304] SERVICE ALERT: ieatnetsimv017-12;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262315] SERVICE ALERT: ieatnetsimv017-13;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262322] SERVICE ALERT: ieatnetsimv017-14;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262374] SERVICE ALERT: ieatnetsimv5116-106;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262382] SERVICE ALERT: ieatnetsimv5116-107;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262389] SERVICE ALERT: ieatnetsimv5116-108;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262397] SERVICE ALERT: ieatnetsimv5116-109;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262404] SERVICE ALERT: ieatnetsimv5116-110;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262412] SERVICE ALERT: ieatnetsimv5116-111;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262419] SERVICE ALERT: ieatnetsimv5116-112;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262428] SERVICE ALERT: ieatnetsimv5116-113;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262435] SERVICE ALERT: ieatnetsimv5116-114;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262442] SERVICE ALERT: ieatnetsimv5116-115;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574262450] SERVICE ALERT: ieatnetsimv5116-116;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574263390] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574263407] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574263409] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (73) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574263549] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (72) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574263550] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574263674] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (1200)
[1574263716] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 2
[1574263848] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574265266] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574265463] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574265659] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574265703] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (11375) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574265845] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (11534) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574265998] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (11569) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574266351] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574266482] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574266616] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574267279] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (221) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574267424] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574267536] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1574267536] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574267562] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (221) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574267692] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (221) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574267821] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574270908] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF PM FILES (2040) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1750)
[1574271188] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 2 MTAS: 1
[1574271354] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574271427] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1574271427] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13502) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574271935] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574272032] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574272129] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574272159] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574272285] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574272322] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574273601] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574273726] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574273821] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574273856] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574273953] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (25) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574273983] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574274102] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574274229] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574274365] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574274497] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574274598] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574274723] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574274851] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574275351] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1574275444] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574275472] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1574275536] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574275564] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574275593] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1574275656] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574275658] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574275685] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574275777] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574275778] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574275854] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1574275882] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574275892] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (47) IS LOWER THAN 80% OF ENM  NODES (80)
[1574275899] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574275907] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (41) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574275975] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1574276003] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574276018] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (43) IS LOWER THAN 80% OF ENM  NODES (80)
[1574276031] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (40) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574276096] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1574276124] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574276152] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574276156] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (42) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574276281] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (36) IS LOWER THAN 80% OF ENM  NODES (50)
[1574276304] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (23) IS LOWER THAN 80% OF ENM  NODES (30)
[1574276406] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574276839] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (8694) IS LOWER THAN 80% OF ENM  NODES (13991)
[1574276973] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13965) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574277268] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12583) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13964)
[1574277416] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12171) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13964)
[1574277425] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574277454] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1574277551] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574277554] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574277568] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13843) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13964)
[1574277639] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574277685] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574277762] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1574277796] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574278107] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574278210] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574279754] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1574279824] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574281563] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (69) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574281689] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (70) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574281820] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574281823] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (79) IS HIGHER THAN 95% OF ENM  NODES (80)
[1574281948] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574282083] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574282145] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized ERBS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574282150] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (1) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1574282228] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (71) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574282345] SERVICE ALERT: ieatlms3901;check pm supervision nodes SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (80)
[1574282353] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (72) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574282477] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (71) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574282507] SERVICE ALERT: ieatlms3901;check fm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1574282513] SERVICE ALERT: ieatlms3901;check pm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1574282637] SERVICE ALERT: ieatlms3901;check fm supervision nodes ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1000)
[1574282741] SERVICE ALERT: ieatlms3901;check pm supervision nodes EPG;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1574282741] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;(Service check timed out after 800.15 seconds)
[1574282745] SERVICE ALERT: ieatlms3901;check cm supervision nodes SGSN-MME;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574282759] SERVICE ALERT: ieatlms3901;check fm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (0)
[1574282762] SERVICE ALERT: ieatlms3901;check cm supervision nodes ESC;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574282805] SERVICE ALERT: ieatlms3901;check pm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574282812] SERVICE ALERT: ieatlms3901;check pm supervision nodes RadioNode;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574282869] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 256 processes with STATE = RSZDT
[1574282873] SERVICE ALERT: ieatlms3901;check cm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1574282891] SERVICE ALERT: ieatlms3901;check pm supervision nodes EPG;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1574282906] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574282929] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 230 processes with STATE = RSZDT
[1574282939] SERVICE ALERT: ieatlms3901;check pm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1574283029] SERVICE ALERT: ieatlms3901;check cm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1574283034] SERVICE ALERT: ieatlms3901;check pm supervision nodes MTAS;CRITICAL;SOFT;1;(Service check timed out after 800.33 seconds)
[1574283087] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1574283133] SERVICE ALERT: ieatlms3901;check unsync nodes;CRITICAL;HARD;3;(Service check timed out after 800.21 seconds)
[1574283242] SERVICE ALERT: ieatlms3901;check pm supervision nodes RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (13964) IS HIGHER THAN 95% OF ENM NODES (13965)
[1574283247] SERVICE ALERT: ieatlms3901;check fm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1574283299] SERVICE ALERT: ieatlms3901;check pm supervision nodes MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574283484] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;SOFT;2;OK- nas_healthcheck SUCCESSFULLY PASSED
[1574283886] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized EPG;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1574283886] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574285001] SERVICE ALERT: ieatlms3901;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574285426] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.1%
[1574285577] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 17.78%
[1574285759] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (1) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1574285889] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1574286019] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (6)
[1574289028] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574289160] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574289706] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13116) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574289843] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13964) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574289976] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13185) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574290110] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13109) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574290275] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13971) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574291094] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574291237] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574291378] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574291453] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.86%
[1574291503] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574291603] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 15.83%
[1574291755] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574291908] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.86%