[1571785894] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571786664] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571787988] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1571808782] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2710) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571809527] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571809582] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571809767] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571809952] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571811535] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571811686] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 58.85%
[1571811831] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3535) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1571811966] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4239) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1571813638] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2734) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571813874] SERVICE ALERT: ieatnetsimv7035-64;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571813904] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4038) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1571813999] SERVICE ALERT: ieatnetsimv7035-64;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571814049] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4250) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1571814279] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.45%
[1571814280] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571814404] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571814429] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 31.68%
[1571814431] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571814583] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571815159] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3999) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571815418] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571815956] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6439) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571816043] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571816194] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571816346] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571816488] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2590) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571817298] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6559) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571818181] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818333] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818480] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818480] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818485] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571818632] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818633] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818687] SERVICE ALERT: ieatnetsimv7035-77;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818692] SERVICE ALERT: ieatnetsimv7035-79;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818696] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818704] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818715] SERVICE ALERT: ieatnetsimv7035-90;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818718] SERVICE ALERT: ieatnetsimv7035-91;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818720] SERVICE ALERT: ieatnetsimv7035-92;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818725] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571818729] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818777] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818784] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571818786] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571818839] SERVICE ALERT: ieatnetsimv7035-77;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818841] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818844] SERVICE ALERT: ieatnetsimv7035-79;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818848] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818856] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818868] SERVICE ALERT: ieatnetsimv7035-90;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818870] SERVICE ALERT: ieatnetsimv7035-91;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818873] SERVICE ALERT: ieatnetsimv7035-92;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818882] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818913] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571818929] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818969] SERVICE ALERT: ieatnetsimv7035-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818981] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571818992] SERVICE ALERT: ieatnetsimv7035-77;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571818995] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571818997] SERVICE ALERT: ieatnetsimv7035-79;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819001] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819008] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819011] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819020] SERVICE ALERT: ieatnetsimv7035-90;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819023] SERVICE ALERT: ieatnetsimv7035-91;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819027] SERVICE ALERT: ieatnetsimv7035-92;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819036] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819040] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819045] SERVICE ALERT: ieatnetsimv7035-88;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819052] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819057] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819059] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819066] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819068] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819069] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819072] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819075] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819076] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819080] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819081] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819100] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571819120] SERVICE ALERT: ieatnetsimv7035-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 44.12%
[1571819133] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819150] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819163] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819192] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819196] SERVICE ALERT: ieatnetsimv7035-88;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819205] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819210] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819212] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819218] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819220] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819221] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819225] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819227] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819228] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819232] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819242] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819280] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819284] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819284] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819285] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819287] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819289] SERVICE ALERT: ieatnetsimv7035-69;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571819314] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819344] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819348] SERVICE ALERT: ieatnetsimv7035-88;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819357] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819362] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819364] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819370] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819373] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819373] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819377] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819380] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819381] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819384] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819394] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819432] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819437] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819437] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571819437] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 38.26%
[1571819440] SERVICE ALERT: ieatnetsimv7035-69;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 51.2%
[1571819477] SERVICE ALERT: ieatnetsimv7035-59;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1571819545] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819582] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.99%
[1571819588] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 90.06%
[1571819589] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571819922] SERVICE ALERT: ieatnetsimv7035-59;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571820388] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1571820509] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1571820629] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1571820968] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571821487] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.98%
[1571821638] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 59.81%
[1571822116] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;OK;HARD;3;OK- CPU IS USED 45.96%
[1571822415] SERVICE FLAPPING ALERT: ieatnetsimv7035-66;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571822415] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.03%
[1571822416] SERVICE FLAPPING ALERT: ieatnetsimv7035-60;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571822416] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 15.62%
[1571822622] SERVICE ALERT: ieatnetsimv7035-77;check netsim cpu;OK;HARD;3;OK- CPU IS USED 16.68%
[1571822627] SERVICE ALERT: ieatnetsimv7035-79;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.33%
[1571822631] SERVICE ALERT: ieatnetsimv7035-81;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.7%
[1571822639] SERVICE ALERT: ieatnetsimv7035-85;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.65%
[1571822650] SERVICE ALERT: ieatnetsimv7035-90;check netsim cpu;OK;HARD;3;OK- CPU IS USED 31.2%
[1571822654] SERVICE ALERT: ieatnetsimv7035-91;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.36%
[1571822658] SERVICE ALERT: ieatnetsimv7035-92;check netsim cpu;OK;HARD;3;OK- CPU IS USED 14.53%
[1571822667] SERVICE ALERT: ieatnetsimv7035-95;check netsim cpu;OK;HARD;3;OK- CPU IS USED 75.54%
[1571822711] SERVICE FLAPPING ALERT: ieatnetsimv7035-57;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571822711] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.9%
[1571822781] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;HARD;3;OK- CPU IS USED 11.72%
[1571822914] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;OK;HARD;3;OK- CPU IS USED 14.53%
[1571822944] SERVICE ALERT: ieatnetsimv7035-67;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.02%
[1571822974] SERVICE FLAPPING ALERT: ieatnetsimv7035-72;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571822974] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.92%
[1571822979] SERVICE ALERT: ieatnetsimv7035-88;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.57%
[1571822987] SERVICE FLAPPING ALERT: ieatnetsimv7035-76;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571822987] SERVICE ALERT: ieatnetsimv7035-76;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.55%
[1571822992] SERVICE FLAPPING ALERT: ieatnetsimv7035-78;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571822992] SERVICE ALERT: ieatnetsimv7035-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.09%
[1571822995] SERVICE FLAPPING ALERT: ieatnetsimv7035-80;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571822995] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.83%
[1571823001] SERVICE FLAPPING ALERT: ieatnetsimv7035-82;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571823001] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.45%
[1571823003] SERVICE FLAPPING ALERT: ieatnetsimv7035-83;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571823003] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;OK;HARD;3;OK- CPU IS USED 16.43%
[1571823004] SERVICE FLAPPING ALERT: ieatnetsimv7035-84;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571823004] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;OK;HARD;3;OK- CPU IS USED 19.24%
[1571823007] SERVICE FLAPPING ALERT: ieatnetsimv7035-86;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571823007] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;OK;HARD;3;OK- CPU IS USED 28.87%
[1571823010] SERVICE FLAPPING ALERT: ieatnetsimv7035-55;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571823010] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;OK;HARD;3;OK- CPU IS USED 41.77%
[1571823011] SERVICE FLAPPING ALERT: ieatnetsimv7035-87;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571823011] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;OK;HARD;3;OK- CPU IS USED 34.14%
[1571823015] SERVICE FLAPPING ALERT: ieatnetsimv7035-89;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571823015] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;OK;HARD;3;OK- CPU IS USED 52.14%
[1571823218] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;OK;HARD;3;OK- CPU IS USED 43.21%
[1571823219] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 50.59%
[1571823577] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;OK;HARD;3;OK- CPU IS USED 29.69%
[1571825413] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 37.95%
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-55;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-57;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-72;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-76;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-78;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-80;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-82;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-83;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-84;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-86;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-87;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826327] SERVICE FLAPPING ALERT: ieatnetsimv7035-89;check netsim cpu;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571826850] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.18%
[1571827001] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.47%
[1571827076] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571827152] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 93.41%
[1571827275] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571827354] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571827471] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571827549] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571827746] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571828656] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3849) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571829349] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571829400] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6588) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571829605] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571830012] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (30) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571830241] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.33%
[1571830246] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.75%
[1571830392] SERVICE ALERT: ieatnetsimv7035-87;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.52%
[1571830398] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.79%
[1571830783] SERVICE FLAPPING ALERT: ieatnetsimv7035-59;check netsim cpu;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1571830783] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.01%
[1571830808] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.18%
[1571830840] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2560) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571830959] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 42.9%
[1571831007] SERVICE FLAPPING ALERT: ieatnetsimv7035-59;check netsim cpu;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1571831027] SERVICE ALERT: ieatlms5735-1;check rpm diff;WARNING;SOFT;1;ATTENTION!!! ISO version is not present
[1571831152] SERVICE ALERT: ieatlms5735-1;check rpm diff;WARNING;SOFT;2;ATTENTION!!! ISO version is not present
[1571831276] SERVICE ALERT: ieatlms5735-1;check rpm diff;WARNING;HARD;3;ATTENTION!!! ISO version is not present
[1571831419] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571831496] SERVICE FLAPPING ALERT: ieatnetsimv7035-59;check netsim cpu;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1571833636] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571833825] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.24%
[1571833825] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571833976] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 19.8%
[1571834021] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571834149] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.22%
[1571834299] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 22.02%
[1571835274] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571835460] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571835645] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571835889] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1571835889] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571836020] HOST ALERT: ieatnetsimv7035-28;DOWN;SOFT;1;PING NOK
[1571836020] HOST ALERT: ieatnetsimv7035-45;DOWN;SOFT;1;PING NOK
[1571836082] HOST ALERT: ieatnetsimv7035-45;UP;SOFT;1;PING OK
[1571836082] HOST ALERT: ieatnetsimv7035-28;UP;SOFT;1;PING OK
[1571837511] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1571837511] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571838071] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (6589) IS HIGHER THAN 95% OF ENM CPP NODES (6589)
[1571838797] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (16584) IS HIGHER THAN 95% OF ENM SNMP NODES (16602)
[1571850881] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH PM SUPERVISION (22408) IS BETWEEN 80% AND 95% OF ENM NODES (27009)
[1571851147] SERVICE ALERT: ieatlms3901;check pm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (26257) IS HIGHER THAN 95% OF ENM NODES (27009)
[1571851976] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.55%
[1571852826] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571853022] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571853219] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571855207] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571855328] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571855449] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1571855690] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571855811] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571855933] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571855959] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1571856080] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1571856122] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1571856201] SERVICE ALERT: ieatlms3901;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1571856243] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1571856293] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1571856364] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1571856414] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1571856535] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1571857173] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1571857294] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1571857415] SERVICE ALERT: ieatlms3901;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1571857439] SERVICE ALERT: ieatlms3901;check san alert healthcheck;CRITICAL;SOFT;1;CRITICAL- san_alert_healthcheck FAILED
[1571857560] SERVICE ALERT: ieatlms3901;check san alert healthcheck;CRITICAL;SOFT;2;CRITICAL- san_alert_healthcheck FAILED
[1571857682] SERVICE ALERT: ieatlms3901;check san alert healthcheck;CRITICAL;HARD;3;CRITICAL- san_alert_healthcheck FAILED
[1571857771] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1571857893] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1571858014] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1571858372] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (877) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1571858400] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1571858700] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1571858700] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571859116] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571859600] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571859860] SERVICE ALERT: ieatlms3901;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1571860025] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1571860979] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571861081] SERVICE ALERT: ieatlms3901;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1571861162] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571861343] SERVICE ALERT: ieatlms3901;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1571861345] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1571861468] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571861655] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571861843] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571861930] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13320) IS FROM 80% TO 95% OF ENM COMECIM NODES (14102)
[1571861980] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1083) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1571862096] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13390) IS FROM 80% TO 95% OF ENM COMECIM NODES (14102)
[1571862275] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13454) IS HIGHER THAN 95% OF ENM COMECIM NODES (14102)
[1571863208] SERVICE FLAPPING ALERT: ieatlms3901;check vcs cluster hc;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1571863208] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571863573] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2560) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571864389] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6587) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571864620] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571864798] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1571865373] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571865425] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES W/O FM HB FAILURES (2542) IS LOWER THAN 80% OF SNMP NODES WITH FM SUPERVISION ENABLED (14851)
[1571865917] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3849) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571866345] SERVICE ALERT: ieatlms3901;check fm hb failures snmp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1571866365] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2710) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571866906] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571871123] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1571871123] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6439) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)