[1567725671] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.37%
[1567725821] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.61%
[1567731908] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567732559] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567736240] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized cpp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567741750] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.36%
[1567741901] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.03%
[1567742059] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 69.45%
[1567743971] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.92%
[1567744122] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.97%
[1567744272] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.45%
[1567744423] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.65%
[1567749571] SERVICE ALERT: ieatnetsimv7035-103;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1567749692] SERVICE ALERT: ieatnetsimv7035-103;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1567749812] SERVICE ALERT: ieatnetsimv7035-103;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1567754719] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567754853] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567754890] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1567754989] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567755374] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567755525] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1567755676] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567755960] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567756012] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13037) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567756145] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567756154] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13037) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567756297] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13037) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567756331] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567756699] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1567756829] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1567756979] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1567757518] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567757669] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567757844] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567757871] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2277) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1567758009] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2352) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1567758142] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2410) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1567758824] SERVICE ALERT: ieatlms3901;check vcs sg;OK;HARD;3;OK
[1567759879] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.33%
[1567759921] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13527) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567759996] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567760029] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.03%
[1567760613] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567762098] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567763689] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.66%
[1567763840] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 28.15%
[1567764235] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1567764366] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1567764496] SERVICE FLAPPING ALERT: ieatlms3901;check vcs sg;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1567764496] SERVICE ALERT: ieatlms3901;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1567765858] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567766126] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567766361] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567766507] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.13%
[1567768343] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567768533] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567768813] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567769027] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567769328] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567769645] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567770010] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567770600] SERVICE ALERT: ieatnetsimv7035-103;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.59%
[1567770607] SERVICE ALERT: ieatnetsimv7035-106;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.12%
[1567770609] SERVICE ALERT: ieatnetsimv7035-107;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.52%
[1567770750] SERVICE ALERT: ieatnetsimv7035-103;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.84%
[1567770758] SERVICE ALERT: ieatnetsimv7035-106;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.85%
[1567770762] SERVICE ALERT: ieatnetsimv7035-107;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.15%
[1567770908] SERVICE ALERT: ieatnetsimv7035-106;check netsim cpu;OK;HARD;3;OK- CPU IS USED 43.12%
[1567771915] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim nodes;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1567771915] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567772605] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (750) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1567772737] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (750) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1567772870] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (750) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1567773848] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.19%
[1567773891] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567773999] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.45%
[1567774106] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567774149] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 78.15%
[1567776317] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567776723] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567776954] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567778745] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM SNMP NODES (0)
[1567779061] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567779356] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567779758] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567780094] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567780573] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567780873] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567781192] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567781345] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567781429] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567782248] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567782265] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;SOFT;2;(Service check timed out after 800.02 seconds)
[1567782369] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567782490] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567782613] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567782679] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567782734] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567782759] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1567782800] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567782855] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1567782880] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1567782921] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567783001] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1567783171] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567783185] SERVICE ALERT: ieatlms3901;check synchronized snmp;CRITICAL;HARD;3;(Service check timed out after 800.48 seconds)
[1567783292] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567783397] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567783413] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567783518] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567783626] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1567783639] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567783747] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1567783829] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1567783868] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1567783950] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1567784071] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1567784073] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1567784194] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1567784315] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1567784818] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (0) IS HIGHER THAN 95% OF ENM CPP NODES (0)
[1567786150] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567786517] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1567786595] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1567786701] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1567787076] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1567787304] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1567787304] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567787418] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567787547] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1567787729] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1567788155] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1567788339] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1567789259] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567791347] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3909) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1567791480] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3909) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1567791613] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3909) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1567792414] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2362) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1567792539] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2362) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1567792666] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2362) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1567795107] SERVICE FLAPPING ALERT: ieatlms3901;check vcs sg;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1567795779] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.03%
[1567795930] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.66%
[1567799780] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.66%
[1567799930] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.85%
[1567800081] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.67%
[1567800232] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.94%
[1567800351] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.78%
[1567800501] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.84%
[1567800652] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.78%
[1567800803] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.69%
[1567800953] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.62%
[1567801104] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 12.03%
[1567801255] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.2%
[1567801405] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 15.79%
[1567801580] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1567810792] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.31%