[1574380830] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574380958] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574381089] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (142) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574381460] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;OK;HARD;3;OK
[1574381679] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13393) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574381908] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574382045] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574382759] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (749) IS BETWEEN 5% AND 20% OF ENM NODES (13991)
[1574382887] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (187) IS LOWER THAN 5% OF ENM NODES (13991)
[1574383020] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (761) IS BETWEEN 5% AND 20% OF ENM NODES (13991)
[1574383149] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (759) IS BETWEEN 5% AND 20% OF ENM NODES (13991)
[1574383278] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (171) IS LOWER THAN 5% OF ENM NODES (13991)
[1574383357] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (14) IS HIGHER THAN 20% OF ENM NODES (50)
[1574383487] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (50)
[1574383495] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13041) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574383600] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574383611] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1574383629] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13293) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574383724] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574383741] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (14) IS HIGHER THAN 20% OF ENM NODES (50)
[1574383764] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13045) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574383869] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574383900] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13039) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574384034] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (13189) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574384086] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (71) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574384211] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574384340] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (70) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574384696] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574384824] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574384950] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (141) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574385238] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (731) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574385847] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13386) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574386147] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (18) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1574386149] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (78) IS HIGHER THAN 95% OF ENM  NODES (80)
[1574386272] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF PM FILES (24) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (24)
[1574387664] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13155) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574387797] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13187) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574387908] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (112) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574387928] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (13197) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574387953] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574388034] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (55) IS LOWER THAN 5% OF ENM NODES (1200)
[1574388078] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574388158] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (115) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574388282] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (118) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574388411] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (123) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574388443] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574388557] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574388850] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (643) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574388938] HOST ALERT: ieatnetsimv7035-34;DOWN;SOFT;1;PING NOK
[1574388978] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (602) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574389000] HOST ALERT: ieatnetsimv7035-34;UP;SOFT;1;PING OK
[1574389106] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (602) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574389745] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13719) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574391049] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized EPG;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1574391129] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MTAS;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1574391564] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13269) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574391699] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13618) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574392023] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (1200)
[1574392713] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (512) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574392879] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1574393139] HOST ALERT: ieatnetsimv7035-31;DOWN;SOFT;1;PING NOK
[1574393201] HOST ALERT: ieatnetsimv7035-31;UP;SOFT;1;PING OK
[1574393637] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13222) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574393787] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13194) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574393921] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (13209) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574394345] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized Router6274;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1574395410] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (67) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574395539] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574395627] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (78) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574395738] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13885) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574395756] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (85) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574395882] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (240) IS LOWER THAN 5% OF ENM NODES (1200)
[1574397167] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1574397327] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1574397551] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13200) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574397684] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (13251) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574397819] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13789) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574398312] HOST ALERT: ieatnetsimv9034-64;DOWN;SOFT;1;PING NOK
[1574398374] HOST ALERT: ieatnetsimv9034-64;UP;SOFT;1;PING OK
[1574399486] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (122) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574399611] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (1200)
[1574399631] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13254) IS FROM 80% TO 95% OF ENM  NODES (13991)
[1574399734] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (94) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574399769] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13871) IS HIGHER THAN 95% OF ENM  NODES (13991)
[1574399859] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (81) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574399982] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (68) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574400050] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (45) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574400174] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574401296] HOST ALERT: ieatnetsimv7035-10;DOWN;SOFT;1;PING NOK
[1574401344] HOST ALERT: ieatnetsimv7035-10;UP;SOFT;1;PING OK
[1574401986] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.87%
[1574402137] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 19.52%
[1574402317] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574402502] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574403060] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574403166] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (4620) IS BETWEEN 80% AND 95% OF ENM NODES (5000)
[1574403181] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574403291] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (5000) IS HIGHER THAN 95% OF ENM NODES (5000)
[1574403301] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574403586] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (1200)
[1574403913] SERVICE ALERT: ieatnetsimv7035-94;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1574405477] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574406243] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574406436] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574406564] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574406630] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574406691] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574406819] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574407313] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574407313] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574407436] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574407438] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574407561] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574407562] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574408347] SERVICE ALERT: ieatnetsimv7102-01;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.45%
[1574408419] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.22%
[1574408420] SERVICE ALERT: ieatnetsimv7102-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.17%
[1574408498] SERVICE ALERT: ieatnetsimv7102-01;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 72.73%
[1574408504] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574408570] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 39.47%
[1574408571] SERVICE ALERT: ieatnetsimv7102-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.08%
[1574408721] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.67%
[1574408872] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574409023] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.17%
[1574409082] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574409209] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574409333] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574409458] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574410219] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1574410219] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (619) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574411165] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574411166] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1574412654] SERVICE ALERT: ieatnetsimv7102-07;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.47%
[1574412731] HOST ALERT: ieatnetsimv017-20;DOWN;SOFT;1;PING NOK
[1574412793] HOST ALERT: ieatnetsimv017-20;UP;SOFT;1;PING OK
[1574412936] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1574412936] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574413897] HOST ALERT: ieatnetsimv7035-105;DOWN;SOFT;1;PING NOK
[1574413959] HOST ALERT: ieatnetsimv7035-105;UP;SOFT;1;PING OK
[1574414079] SERVICE ALERT: ieatlms7435;check fm supervision nodes DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1574414201] SERVICE ALERT: ieatlms7435;check fm supervision nodes DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1574414263] SERVICE ALERT: ieatlms7435;check fm supervision nodes EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (50)
[1574414324] SERVICE ALERT: ieatlms7435;check fm supervision nodes DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1574414386] SERVICE ALERT: ieatlms7435;check fm supervision nodes EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (50)
[1574414508] SERVICE ALERT: ieatlms7435;check fm supervision nodes EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (50)
[1574414695] SERVICE ALERT: ieatlms7435;check fm supervision nodes MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574414697] SERVICE ALERT: ieatlms7435;check fm supervision nodes SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (80)
[1574414733] SERVICE ALERT: ieatlms7435;check fm supervision nodes MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1574414817] SERVICE ALERT: ieatlms7435;check fm supervision nodes MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574414819] SERVICE ALERT: ieatlms7435;check fm supervision nodes SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (80)
[1574414855] SERVICE ALERT: ieatlms7435;check fm supervision nodes MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1574414940] SERVICE ALERT: ieatlms7435;check fm supervision nodes MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1574414943] SERVICE ALERT: ieatlms7435;check fm supervision nodes SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (80)
[1574414980] SERVICE ALERT: ieatlms7435;check fm supervision nodes MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1574415640] HOST ALERT: ieatnetsimv7035-117;DOWN;SOFT;1;PING NOK
[1574415702] HOST ALERT: ieatnetsimv7035-117;UP;SOFT;1;PING OK
[1574415940] HOST ALERT: ieatnetsimv7035-22;DOWN;SOFT;1;PING NOK
[1574415975] HOST ALERT: ieatnetsimv7035-22;UP;SOFT;1;PING OK
[1574415996] HOST ALERT: ieatnetsimv7035-92;DOWN;SOFT;1;PING NOK
[1574415996] HOST ALERT: ieatnetsimv7035-101;DOWN;SOFT;1;PING NOK
[1574415996] HOST ALERT: ieatnetsimv7035-47;DOWN;SOFT;1;PING NOK
[1574416058] HOST ALERT: ieatnetsimv7035-92;UP;SOFT;1;PING OK
[1574416058] HOST ALERT: ieatnetsimv7035-47;UP;SOFT;1;PING OK
[1574416058] HOST ALERT: ieatnetsimv7035-101;UP;SOFT;1;PING OK
[1574416526] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12550) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574416543] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574416616] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2338) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574416629] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574416676] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574416685] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12550) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574416743] SERVICE ALERT: ieatlms7435;check fm supervision nodes MGW;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1574416746] SERVICE ALERT: ieatlms7435;check fm supervision nodes SGSN-MME;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1574416748] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2338) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574416770] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574416782] SERVICE ALERT: ieatlms7435;check fm supervision nodes MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574416783] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574416803] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574416833] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574416876] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574416896] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1574416909] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (50)
[1574417027] SERVICE ALERT: ieatlms7435;check fm supervision nodes DSC;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574417051] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (50)
[1574417211] SERVICE ALERT: ieatlms7435;check fm supervision nodes EPG;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1574417290] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574417996] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (332) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574418122] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (1200)
[1574420411] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (135) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574420510] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (70) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574420516] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12158) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574420538] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574420564] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2271) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574420637] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574420665] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12825) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574420665] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574420692] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2357) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574420763] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574420817] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12830) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574420839] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2357) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574421785] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MSC-DB-BSP;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1574422667] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1574422791] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1574422925] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (10) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1574423050] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (13) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1574423174] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1574423196] HOST ALERT: ieatnetsimv7035-112;DOWN;SOFT;1;PING NOK
[1574423196] HOST ALERT: ieatnetsimv7035-42;DOWN;SOFT;1;PING NOK
[1574423258] HOST ALERT: ieatnetsimv7035-112;UP;SOFT;1;PING OK
[1574423258] HOST ALERT: ieatnetsimv7035-42;UP;SOFT;1;PING OK
[1574423438] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1574423581] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1574424259] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (11) IS HIGHER THAN 20% OF ENM NODES (50)
[1574424395] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (14) IS HIGHER THAN 20% OF ENM NODES (50)
[1574424526] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (50)
[1574424705] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574424843] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574424969] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (28) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574425296] HOST ALERT: ieatnetsimv7035-76;DOWN;SOFT;1;PING NOK
[1574425296] HOST ALERT: ieatnetsimv9041-05;DOWN;SOFT;1;PING NOK
[1574425358] HOST ALERT: ieatnetsimv9041-05;UP;SOFT;1;PING OK
[1574425358] HOST ALERT: ieatnetsimv7035-76;UP;SOFT;1;PING OK
[1574425454] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (272) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574425579] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (272) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574425704] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (1200)
[1574426458] SERVICE FLAPPING ALERT: ieatlms7435;check pm collection statistical DSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574426982] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574427107] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1574427110] SERVICE ALERT: ieatnetsimv7035-93;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574427261] SERVICE ALERT: ieatnetsimv7035-93;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 22.46%
[1574427887] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574428072] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574428130] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1574428258] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574428275] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (50)
[1574428404] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (11) IS HIGHER THAN 20% OF ENM NODES (50)
[1574428530] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1574428654] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1574428707] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574429195] HOST ALERT: ieatnetsimv7035-44;DOWN;SOFT;1;PING NOK
[1574429227] SERVICE ALERT: ieatlms3901;check fm supervision nodes DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1574429258] HOST ALERT: ieatnetsimv7035-44;UP;SOFT;1;PING OK
[1574429320] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (187) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574429354] SERVICE ALERT: ieatlms3901;check fm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1574429444] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (148) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574429495] HOST ALERT: ieatnetsimv7035-45;DOWN;SOFT;1;PING NOK
[1574429557] HOST ALERT: ieatnetsimv7035-45;UP;SOFT;1;PING OK
[1574429568] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (60) IS LOWER THAN 5% OF ENM NODES (1200)
[1574429795] HOST ALERT: ieatnetsimv7035-18;DOWN;SOFT;1;PING NOK
[1574429857] HOST ALERT: ieatnetsimv7035-18;UP;SOFT;1;PING OK
[1574430123] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574431476] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574431571] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1574431646] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574431648] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1574431648] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574432261] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (50)
[1574435087] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574435141] HOST ALERT: ieatnetsimv9041-14;DOWN;SOFT;1;PING NOK
[1574435141] HOST ALERT: ieatnetsimv7035-11;DOWN;SOFT;1;PING NOK
[1574435178] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574435203] HOST ALERT: ieatnetsimv7035-11;UP;SOFT;1;PING OK
[1574435203] HOST ALERT: ieatnetsimv9041-14;UP;SOFT;1;PING OK
[1574435214] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574435274] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12609) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574435274] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2311) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574435305] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574435342] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574435403] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2311) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574435425] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12609) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574435431] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574435532] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2311) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574435587] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12609) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574435864] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574435988] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574436114] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574438031] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574438164] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1574438289] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574438413] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1574438949] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574439038] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1574439141] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574439219] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574439374] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574439501] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574439627] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1574439773] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 2
[1574439916] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574440827] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (4) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574442556] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (134) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574442697] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (134) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574442773] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2338) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574442828] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (134) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574442850] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12659) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574442902] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2338) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574442998] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574443033] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574443232] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1574443318] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (50)
[1574444588] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574444774] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574444959] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574445305] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574445340] HOST ALERT: ieatnetsimv7035-89;DOWN;SOFT;1;PING NOK
[1574445402] HOST ALERT: ieatnetsimv7035-89;UP;SOFT;1;PING OK
[1574445433] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574446003] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574446182] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574446360] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574446450] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574446641] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2303) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574446745] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12554) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574446770] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2304) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574446892] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12570) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574446902] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2304) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574447043] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12570) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574447622] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (1) IS LOWER THAN 80% OF ENM NODES (20)
[1574447747] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1574447973] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1000)
[1574448098] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1574448597] SERVICE ALERT: ieatnetsimv7035-115;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.81%
[1574448747] SERVICE ALERT: ieatnetsimv7035-115;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.21%
[1574450511] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574450671] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574450776] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (633) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574450900] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (646) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574451027] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (660) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574452834] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574453156] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574453337] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574453444] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574453519] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574453588] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574453664] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574453716] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (78) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574453790] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150)
[1574456441] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (703) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574456565] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (709) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574456690] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574456807] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1579) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574456890] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1053) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1574456932] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1605) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574457019] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1120) IS FROM 80% TO 95% OF ENM  NODES (1200)
[1574457071] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1749) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574457143] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1166) IS HIGHER THAN 95% OF ENM  NODES (1200)
[1574457323] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574457450] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (77) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574457483] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574457519] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574457604] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574457647] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574457720] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2429) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574457726] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574457773] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574457849] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2429) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574457898] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13124) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574457981] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2429) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1574458046] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13124) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574458075] SERVICE ALERT: ieatlms3901;check fm supervision nodes RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (13965)
[1574458110] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (589) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574458193] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (13124) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574458217] SERVICE ALERT: ieatlms3901;check fm supervision nodes RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (13964) IS HIGHER THAN 95% OF ENM NODES (13965)
[1574458819] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574458837] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574458976] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574459003] SERVICE ALERT: ieatlms3901;check vcs llt heartbeat hc;OK;SOFT;2;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1574459107] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (3) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (5)
[1574460675] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1345) IS LOWER THAN 80% OF ENM  NODES (1750)
[1574460778] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574460801] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1749) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574460862] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1574460928] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1401) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574460972] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574461052] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1425) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574461165] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574461177] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1456) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574461224] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574461718] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (622) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1574462717] SERVICE ALERT: ieatlms3901;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (5) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (5)
[1574462982] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1749) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574463035] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574463038] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574463180] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574463309] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574463355] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1574463476] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1574463597] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1574463605] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1574463726] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1574463760] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574463847] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1574463943] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;SOFT;2;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1574464064] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574464184] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574464305] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574464320] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574464441] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574464463] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574464498] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;SOFT;1;CRITICAL- san_alert_healthcheck FAILED
[1574464551] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1574464562] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574464583] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574464619] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;SOFT;2;CRITICAL- san_alert_healthcheck FAILED
[1574464655] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1574464672] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1574464704] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574464740] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;HARD;3;CRITICAL- san_alert_healthcheck FAILED
[1574464776] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1574464793] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1574464794] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574464825] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574464839] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574464897] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1574464921] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574464980] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (143) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1574465017] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574465033] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574465049] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1574465420] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965)
[1574466169] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1574466431] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574466582] SERVICE ALERT: ieatnetsimv9041-05;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.18%
[1574466758] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1574466918] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574467002] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574467042] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574467167] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574467188] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED