[1563167705] SERVICE ALERT: ieatnetsimv5116-80;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1563167825] SERVICE ALERT: ieatnetsimv5116-80;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1563167946] SERVICE ALERT: ieatnetsimv5116-80;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1563171682] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1563171941] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563172549] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563172671] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1563172673] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;CRITICAL;HARD;3;(Service check timed out after 801.01 seconds)
[1563172864] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1563173112] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (122) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1563173168] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563173176] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1563173257] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1563173394] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1563173820] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4) IS HIGHER THAN 95% OF ENM CPP NODES (4)
[1563173828] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (4) IS HIGHER THAN 95% OF ENM SNMP NODES (4)
[1563173846] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (32) IS HIGHER THAN 95% OF ENM COMECIM NODES (5)
[1563175545] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563175653] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2193) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563175669] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563175795] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1563175799] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2193) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563175936] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2196) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1563176158] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563176285] SERVICE ALERT: ieatlms5735-1;check synchronized nodes com/ecim;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563176286] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563176286] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1563176410] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1563176736] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563176869] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563176993] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (122) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1563177748] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1563177748] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2489) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1563184583] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563184768] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1563184953] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1563188618] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1563189714] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563189856] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (14035) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1563191134] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1563191315] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1563191496] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1563191620] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563191801] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563191982] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563200326] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (1219) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563201048] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6392) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (5595)
[1563205724] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.33%
[1563205875] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.22%
[1563206025] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.57%
[1563206684] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563206869] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1563206880] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1563207054] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1563207264] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1563207622] SERVICE ALERT: ieatlms3901;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1563208508] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.38%
[1563208665] SERVICE ALERT: ieatnetsimv5116-62;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 68.43%
[1563208919] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1563209106] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1563209292] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1563211276] SERVICE ALERT: ieatnetsimv5116-16;check netsim disk;OK;HARD;3;OK
[1563211965] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1563213256] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.19%
[1563213406] SERVICE ALERT: ieatnetsimv7035-15;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.06%
[1563220117] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563220255] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1563220298] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563220436] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1563220478] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1563220616] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED