[1566342037] SERVICE ALERT: ieatnetsimv5116-44;check netsim disk;OK;HARD;3;OK
[1566343289] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.01%
[1566343439] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.61%
[1566343591] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.96%
[1566343741] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.42%
[1566343893] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.93%
[1566344044] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.82%
[1566344194] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.87%
[1566348675] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.95%
[1566348826] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.86%
[1566348976] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;OK;HARD;3;OK- CPU IS USED 80.76%
[1566349237] SERVICE ALERT: ieatnetsimv5116-44;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566349358] SERVICE ALERT: ieatnetsimv5116-44;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566349479] SERVICE ALERT: ieatnetsimv5116-44;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566356033] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566356184] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 69.33%
[1566356335] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.62%
[1566356485] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 40.96%
[1566356636] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1566356787] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566356938] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.97%
[1566359916] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.2%
[1566360067] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.54%
[1566360569] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.05%
[1566365709] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1566367580] SERVICE ALERT: ieatnetsimv5116-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.88%
[1566367731] SERVICE ALERT: ieatnetsimv5116-11;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566367882] SERVICE ALERT: ieatnetsimv5116-11;check netsim cpu;OK;HARD;3;OK- CPU IS USED 26.74%
[1566369995] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 38.48%
[1566369998] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566370132] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566370268] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566370395] SERVICE ALERT: ieatnetsimv7035-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566370525] SERVICE ALERT: ieatnetsimv7035-31;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566370595] SERVICE ALERT: ieatnetsimv7035-26;check netsim cpu;OK;HARD;3;OK- CPU IS USED 40.74%
[1566370607] SERVICE ALERT: ieatnetsimv7035-31;check netsim cpu;OK;HARD;3;OK- CPU IS USED 31.87%
[1566370655] SERVICE ALERT: ieatnetsimv7035-31;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566370670] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12518) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1566370832] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12555) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1566370848] SERVICE ALERT: ieatnetsimv7035-28;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566370884] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 52.68%
[1566370886] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1566370970] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1566370983] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12820) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1566371009] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1566371024] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1566371191] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566371255] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566371339] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566371366] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2387) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566371452] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (11933) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566371486] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566371492] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2387) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566371550] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566371595] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (11933) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566371619] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2387) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566371739] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (11840) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566371879] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566372384] SERVICE ALERT: ieatnetsimv7035-26;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566372471] SERVICE ALERT: ieatnetsimv7035-31;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566372659] SERVICE ALERT: ieatnetsimv7035-68;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566372807] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13724) IS HIGHER THAN 95% OF ENM COMECIM NODES (14149)
[1566373025] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.5%
[1566373175] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.65%
[1566373339] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566375361] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13727) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566375612] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 800.04 seconds)
[1566375775] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566375853] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1566378115] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.47%
[1566378265] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.57%
[1566380953] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566381486] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.39%
[1566381651] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566381811] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566382200] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.19%
[1566382350] SERVICE ALERT: ieatnetsimv5116-59;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 11.78%
[1566382732] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566382866] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566382891] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.22%
[1566382999] SERVICE ALERT: ieatnetsimv5116-14;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566383027] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566383041] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.23%
[1566383192] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.23%
[1566383249] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566383342] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.12%
[1566383493] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.1%
[1566383519] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566384426] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.66%
[1566384576] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.22%
[1566384727] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.31%
[1566384878] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.22%
[1566385028] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 90.84%
[1566385565] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim nodes;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1566385565] SERVICE ALERT: ieatnetsimv7035-09;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566386025] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2468) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566387670] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.98%
[1566387820] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.37%
[1566387971] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.18%
[1566388122] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.56%
[1566391907] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1566392100] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1566392293] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1566392350] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566392476] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566392601] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566393230] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566393355] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566393370] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566393480] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566393500] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566393629] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566394168] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1566396248] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.5%
[1566396789] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1566396941] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566396968] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1566397078] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566397123] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1566397146] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1566397264] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566397305] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566397450] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1566398526] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.72%
[1566398676] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.04%
[1566398827] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.4%
[1566399390] SERVICE ALERT: ieatnetsimv5116-42;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566399511] SERVICE ALERT: ieatnetsimv5116-42;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566399632] SERVICE ALERT: ieatnetsimv5116-42;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566399879] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.93%
[1566400033] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.22%
[1566400191] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.02%
[1566401128] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2251) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566401260] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2309) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566401391] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2409) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1566402767] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1566402949] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1566403131] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1566403821] SERVICE FLAPPING ALERT: ieatnetsimv5116-66;check netsim cpu;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1566403821] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;HARD;3;OK- CPU IS USED 53.67%
[1566407034] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12938) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1566407179] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13563) IS HIGHER THAN 95% OF ENM COMECIM NODES (14149)
[1566410393] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1566410403] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1874) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1566410535] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1874) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1566410667] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1874) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1566412480] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1566412509] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.59%
[1566416140] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.89%
[1566416291] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.59%
[1566416442] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.85%
[1566416592] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.27%
[1566416743] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.54%
[1566416893] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.88%
[1566417315] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566417892] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2055) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566418024] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2143) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566418123] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12738) IS FROM 80% TO 95% OF ENM COMECIM NODES (14149)
[1566418158] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2215) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566418272] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13445) IS HIGHER THAN 95% OF ENM COMECIM NODES (14149)
[1566418845] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1566419970] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1566419970] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1566419990] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim nodes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566420674] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.33%
[1566420824] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.9%
[1566420974] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 82.67%
[1566424058] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.4%
[1566424208] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.48%
[1566424360] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.41%
[1566424510] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.73%
[1566424606] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.71%
[1566424757] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.18%
[1566424907] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.61%
[1566425058] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.78%
[1566426724] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.25%
[1566428260] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.29%