[1571871778] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1571874071] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571874254] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571874629] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1571875247] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13319) IS FROM 80% TO 95% OF ENM COMECIM NODES (14102)
[1571875523] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1571875556] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13426) IS HIGHER THAN 95% OF ENM COMECIM NODES (14102)
[1571876017] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571876289] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571876649] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.36%
[1571877566] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12830) IS FROM 80% TO 95% OF ENM COMECIM NODES (14102)
[1571877740] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12980) IS FROM 80% TO 95% OF ENM COMECIM NODES (14102)
[1571877904] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13029) IS FROM 80% TO 95% OF ENM COMECIM NODES (14102)
[1571878037] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571878220] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571878404] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1571878887] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5947) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1571879030] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (6569) IS HIGHER THAN 95% OF ENM CPP NODES (6589)
[1571879461] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2300) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571879592] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2300) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571879630] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (126) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571879718] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571879749] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (14083) IS HIGHER THAN 95% OF ENM COMECIM NODES (14102)
[1571879755] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571881829] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1571881950] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1571882067] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571882070] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1571882232] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2203) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571882365] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1571882512] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2328) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571882643] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2381) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1571883853] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.42%
[1571883930] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571884113] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571884296] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571886159] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571886342] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571886525] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1571888188] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5939) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1571888340] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (6586) IS HIGHER THAN 95% OF ENM CPP NODES (6589)
[1571888806] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13119) IS FROM 80% TO 95% OF ENM COMECIM NODES (14102)
[1571888967] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (14095) IS HIGHER THAN 95% OF ENM COMECIM NODES (14102)
[1571889143] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13224) IS FROM 80% TO 95% OF ENM COMECIM NODES (14101)
[1571889310] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13137) IS FROM 80% TO 95% OF ENM COMECIM NODES (14101)
[1571889377] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1571889466] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12914) IS FROM 80% TO 95% OF ENM COMECIM NODES (14102)
[1571890188] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571891304] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized comecim;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1571891304] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (14095) IS HIGHER THAN 95% OF ENM COMECIM NODES (14101)
[1571893852] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571894034] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571894217] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571895376] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2186) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571895519] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2452) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1571896111] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1571899167] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;OK;HARD;3;OK- CPU IS USED 16.16%
[1571899213] SERVICE FLAPPING ALERT: ieatnetsimv7035-59;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1571899283] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (5820) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1571899426] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (6563) IS HIGHER THAN 95% OF ENM CPP NODES (6589)
[1571900431] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1571900431] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized comecim;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1571900431] SERVICE FLAPPING ALERT: ieatlms3901;check vcs cluster hc;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1571901347] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1571901511] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571902088] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571902219] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (25.7% change >= 20.0% threshold)
[1571902219] SERVICE FLAPPING ALERT: ieatlms3901;check vcs cluster hc;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1571902242] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;1;PING NOK
[1571902304] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;2;PING NOK
[1571902366] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;3;PING NOK
[1571902428] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;4;PING NOK
[1571902491] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;5;PING NOK
[1571902552] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;6;PING NOK
[1571902614] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;7;PING NOK
[1571902676] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;8;PING NOK
[1571902738] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;9;PING NOK
[1571902800] HOST ALERT: ieatnetsimv9041-01;DOWN;HARD;10;PING NOK
[1571902853] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2280) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571902989] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2334) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571903152] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2407) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1571903378] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571903565] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571903670] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;OK;HARD;3;OK
[1571903753] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571904924] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2412) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571905054] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2483) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571906007] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.3%
[1571906158] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 35.24%
[1571907919] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1571909033] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571909465] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571910173] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6559) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571910571] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1571910580] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;1;PING NOK
[1571910642] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;2;PING NOK
[1571910705] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;3;PING NOK
[1571910767] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;4;PING NOK
[1571910829] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;5;PING NOK
[1571910891] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;6;PING NOK
[1571910953] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;7;PING NOK
[1571910989] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13118) IS FROM 80% TO 95% OF ENM COMECIM NODES (14101)
[1571911015] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;8;PING NOK
[1571911077] HOST ALERT: ieatnetsimv9041-01;DOWN;SOFT;9;PING NOK
[1571911140] HOST ALERT: ieatnetsimv9041-01;DOWN;HARD;10;PING NOK
[1571911149] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12941) IS FROM 80% TO 95% OF ENM COMECIM NODES (14101)
[1571911329] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized comecim;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1571911329] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12825) IS FROM 80% TO 95% OF ENM COMECIM NODES (14101)
[1571912137] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (1886) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571912139] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (119) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571912161] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (11061) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13978)
[1571912173] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2220) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571912267] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571912267] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2177) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571912313] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12229) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13978)
[1571912342] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2220) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571912393] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (144) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571912397] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2277) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571912430] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6152) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1571912460] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12545) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13978)
[1571912473] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2220) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571912573] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (6349) IS HIGHER THAN 95% OF ENM CPP NODES (6589)
[1571912582] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1571912703] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1571912738] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (920) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1571912743] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (58) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571912760] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2300) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571912823] SERVICE FLAPPING ALERT: ieatnetsimv7035-30;check netsim disk;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1571912823] SERVICE ALERT: ieatnetsimv7035-30;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1571912865] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (920) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1571912870] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (58) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571912888] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2300) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571912986] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1571912986] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized comecim;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1571912986] SERVICE FLAPPING ALERT: ieatnetsimv7035-30;check netsim disk;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1571912996] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (920) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1571912997] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (58) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571913017] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2300) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571913358] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (126) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571913361] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571913484] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571913485] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571913610] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (126) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571913612] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571913838] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571913991] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571914144] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571914586] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13509) IS HIGHER THAN 95% OF ENM COMECIM NODES (14135)
[1571915465] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571916158] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6349) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571916623] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571916865] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571917215] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571917216] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571917429] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571917433] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571917436] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571917583] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571917586] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571917589] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (6169) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571917590] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571917735] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571917738] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571917743] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571917810] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (180) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571919976] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13770) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14016)
[1571920750] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571920973] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1571921190] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1571921976] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (6169) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571924868] SERVICE ALERT: ieatnetsimv7035-51;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1571925953] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3879) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1571926807] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571926886] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13977) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13978)
[1571926888] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1571927404] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;OK;HARD;3;OK- NUMBER OF PM FILES (998) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1571927406] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571927429] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2391) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571927556] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2446) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571927573] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13970) IS HIGHER THAN 95% OF ENM COMECIM NODES (14101)
[1571928733] SERVICE FLAPPING ALERT: ieatnetsimv7035-30;check netsim disk;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1571932040] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571932233] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571932427] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571934303] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571934951] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1571935072] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1571935193] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1571935286] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1571935407] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1571935528] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1571935568] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571935689] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571935774] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.15%
[1571935810] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571936104] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571936153] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1571936165] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571936225] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571936274] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1571936286] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571936346] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571936454] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1571936474] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571937391] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1571937409] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1571937587] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1571937765] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1571938221] SERVICE FLAPPING ALERT: ieatlms5735-1;check system service hc;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1571938221] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571939043] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1571939374] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.06%
[1571940136] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571940318] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571941276] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571941629] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (7038) IS LOWER THAN 80% OF ENM COMECIM NODES (14135)
[1571941771] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;2;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (8099) IS LOWER THAN 80% OF ENM COMECIM NODES (14135)
[1571941912] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (7936) IS LOWER THAN 80% OF ENM COMECIM NODES (14135)
[1571942300] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571942482] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571942727] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.46%
[1571942878] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 54.81%
[1571943141] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571943327] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1571943513] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1571943737] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13702) IS HIGHER THAN 95% OF ENM COMECIM NODES (14135)
[1571944464] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571944646] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571944726] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3669) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1571944828] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1571944864] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4250) IS HIGHER THAN 95% OF ENM CPP NODES (4254)
[1571946021] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571946146] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571946271] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571946565] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 26.39%
[1571946569] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 19.31%
[1571946690] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571947812] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1858) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1571947945] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1858) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1571948077] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1914) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1571948553] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1571948735] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1571948798] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13237) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14016)
[1571948917] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1571948940] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (13237) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (14016)
[1571949081] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13776) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14016)
[1571949876] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1571949876] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571949889] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2467) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1571954379] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1571954379] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1571955301] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2270) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571955437] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2302) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571955569] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2333) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1571956481] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571956607] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2557) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571956733] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571956860] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2557) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571957384] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1571957384] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1864) IS LOWER THAN 80% OF ENM SNMP NODES (2500)