[1568246202] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.72%
[1568246352] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.02%
[1568246507] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.43%
[1568247185] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1568247336] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.11%
[1568247487] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.54%
[1568249483] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.37%
[1568249634] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.54%
[1568249784] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.8%
[1568249935] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.51%
[1568250085] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.99%
[1568250236] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.35%
[1568251118] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.25%
[1568254748] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.96%
[1568254899] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.28%
[1568269164] SERVICE ALERT: ieatnetsimv7035-47;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568269293] SERVICE ALERT: ieatnetsimv7035-47;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568269429] SERVICE ALERT: ieatnetsimv7035-47;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568270773] SERVICE FLAPPING ALERT: ieatnetsimv7035-12;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1568271922] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1568272257] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1568275586] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.51%
[1568275737] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.4%
[1568275888] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.5%
[1568276039] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.41%
[1568276645] SERVICE ALERT: ieatnetsimv7035-47;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568282538] SERVICE FLAPPING ALERT: ieatnetsimv7035-13;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1568282788] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1568282974] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1568283158] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1568283389] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.04%
[1568283540] SERVICE FLAPPING ALERT: ieatnetsimv7035-08;check netsim cpu;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1568283540] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.44%
[1568283695] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.59%
[1568283845] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.89%
[1568283996] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.87%
[1568284005] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM CPP NODES (4224)
[1568284146] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.84%
[1568284297] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.14%
[1568284447] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568284598] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.2%
[1568284759] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (5) IS LOWER THAN 80% OF ENM COMECIM NODES (14049)
[1568284896] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;2;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (5) IS LOWER THAN 80% OF ENM COMECIM NODES (14049)
[1568285033] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (5) IS LOWER THAN 80% OF ENM COMECIM NODES (14049)
[1568285597] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1568285778] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1568285959] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1568286823] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1568286973] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1568287154] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1568287334] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1568287449] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.65%
[1568287600] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.3%
[1568288228] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.98%
[1568288380] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568288531] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.21%
[1568289064] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.09%
[1568289215] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.06%
[1568289827] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2340) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568290077] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1882) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1568290145] SERVICE ALERT: ieatnetsimv7035-108;check netsim cpu;CRITICAL;SOFT;1;(Service check timed out after 800.40 seconds)
[1568290328] SERVICE ALERT: ieatnetsimv7035-108;check netsim cpu;OK;SOFT;2;OK- CPU IS USED %
[1568290455] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12154) IS FROM 80% TO 95% OF ENM COMECIM NODES (14049)
[1568290489] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1568290674] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1568290859] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1568290859] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1568291227] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3423) IS FROM 80% TO 95% OF ENM CPP NODES (4224)
[1568291313] SERVICE ALERT: ieatnetsimv7035-103;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568291466] SERVICE ALERT: ieatnetsimv7035-103;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568291490] SERVICE ALERT: ieatnetsimv7035-108;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568291608] SERVICE ALERT: ieatnetsimv7035-103;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568291634] SERVICE ALERT: ieatnetsimv7035-108;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568291782] SERVICE ALERT: ieatnetsimv7035-108;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568292281] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (14025) IS HIGHER THAN 95% OF ENM COMECIM NODES (14049)
[1568292861] SERVICE ALERT: ieatnetsimv7035-47;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568292914] SERVICE ALERT: ieatnetsimv7035-107;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.7%
[1568292943] SERVICE ALERT: ieatnetsimv7035-100;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568292990] SERVICE ALERT: ieatnetsimv7035-47;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568293065] SERVICE ALERT: ieatnetsimv7035-107;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.95%
[1568293072] SERVICE ALERT: ieatnetsimv7035-100;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568293120] SERVICE ALERT: ieatnetsimv7035-47;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568293409] SERVICE ALERT: ieatnetsimv7035-100;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568293423] SERVICE ALERT: ieatnetsimv7035-103;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568293600] SERVICE ALERT: ieatnetsimv7035-108;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568293655] SERVICE ALERT: ieatnetsimv7035-103;check netsim disk;OK;HARD;3;OK
[1568293684] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2324) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1568294524] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1568294950] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.71%
[1568295101] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.97%
[1568295224] SERVICE ALERT: ieatnetsimv7035-100;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568295252] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.36%
[1568295762] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.78%
[1568295912] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.38%
[1568296063] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.83%
[1568296938] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.84%
[1568297291] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2453) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1568298882] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.92%
[1568299033] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.28%
[1568299184] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.91%
[1568299694] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.11%
[1568300569] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.31%
[1568300719] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.31%
[1568300870] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.28%
[1568306415] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.38%
[1568309099] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1568309429] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1568309677] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (728) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1568309786] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1568309806] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (728) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1568309933] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1568310045] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.79%
[1568310196] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.84%
[1568310347] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.95%
[1568310497] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.76%
[1568310524] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.04%
[1568310648] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.87%
[1568310675] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568310798] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.09%
[1568310826] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 93.69%
[1568310949] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.32%
[1568311035] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (10449) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1568311176] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (10449) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1568311308] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (6968) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1568311698] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (1257) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1568311824] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (1809) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1568311950] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1568311950] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1837) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2530)
[1568315301] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.08%
[1568315451] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.13%
[1568315602] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.38%
[1568322832] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.66%
[1568322983] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.28%
[1568323133] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;HARD;3;OK- CPU IS USED 81.64%
[1568325257] SERVICE FLAPPING ALERT: ieatnetsimv7035-08;check netsim cpu;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1568325257] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.34%
[1568326990] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1568328887] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.49%
[1568328967] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.91%
[1568329038] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.89%
[1568329117] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 40.97%
[1568329189] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.92%