[1574726546] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574726669] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574726793] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (9) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574726914] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (344) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574727038] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (361) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574727112] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized nodes cpp;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1574727112] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1574727161] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574727620] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574727746] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574727871] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574728356] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.3%
[1574728460] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1574728507] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 55.69%
[1574728658] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.59%
[1574728809] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 32.19%
[1574730393] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (743) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574730524] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;OK;SOFT;2;OK- NUMBER OF PM FILES (995) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574730718] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574730768] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (391) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574730855] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574730897] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (417) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574730926] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (698) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574730984] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1574731021] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (425) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574731477] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (28) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574732146] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574732270] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574732409] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574733612] SERVICE ALERT: ieatwlvm5116;check workload profiles;OK;HARD;3;OK- ALL WORKLOAD PROFILES ARE RUNNING
[1574733997] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1574733997] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1574734534] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1574734534] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (446) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574734625] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1574734625] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574734836] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574734964] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574735087] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (111) IS LOWER THAN 80% OF ENM  NODES (150)
[1574735099] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574735212] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1574736580] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1574736580] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574737478] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (852) IS BETWEEN 5% AND 20% OF ENM NODES (13969)
[1574737606] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574737608] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (438) IS LOWER THAN 5% OF ENM NODES (13969)
[1574737740] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (13) IS HIGHER THAN 20% OF ENM NODES (50)
[1574737868] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (13) IS HIGHER THAN 20% OF ENM NODES (50)
[1574738229] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (469) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574738357] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574738373] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1574738676] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1574738705] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574738976] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1574740184] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574740308] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574740434] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574741336] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (918) IS BETWEEN 5% AND 20% OF ENM NODES (13969)
[1574741470] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (887) IS BETWEEN 5% AND 20% OF ENM NODES (13966)
[1574741473] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (50)
[1574741597] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (13969)
[1574741791] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574741918] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574742000] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (144) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574742048] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574742057] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2469) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574742085] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (506) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574742100] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13449) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574742208] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574742332] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (205) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574742456] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (152) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574742579] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574743832] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (16) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1574743958] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (12) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1574744084] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (100)
[1574744350] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574744489] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574744619] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574745186] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.58%
[1574745255] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574745289] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1560) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574745337] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 54.3%
[1574745383] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574745417] SERVICE ALERT: ieatlms3901;check synchronized ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1560) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574745524] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574745526] SERVICE ALERT: ieatlms3901;check cm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574745545] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2475) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1574745551] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12984) IS FROM 80% TO 95% OF ENM  NODES (13969)
[1574745608] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (126) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574745649] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574745658] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1574745665] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2144) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574745684] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13933) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574745727] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (11782) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574745735] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (126) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574745784] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574745795] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2144) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574745865] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (126) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574745875] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (11782) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574745882] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (23) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574745911] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574745923] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2144) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574746011] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (23) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574746023] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (11782) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574746040] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574746142] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (23) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574746168] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574746184] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (634) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574746308] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (634) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574746443] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (626) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574747620] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (12865) IS FROM 80% TO 95% OF ENM  NODES (13950)
[1574747659] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574747760] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12453) IS FROM 80% TO 95% OF ENM  NODES (13949)
[1574747892] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13890) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574748451] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574748575] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574748658] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (76) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574748889] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574748892] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574749013] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574749086] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574749137] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1574749279] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574749704] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13028) IS FROM 80% TO 95% OF ENM  NODES (13942)
[1574749765] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (26) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574749838] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (13516) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574750047] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574750498] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (75) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574750626] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574750751] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574752222] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574752280] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (13) IS HIGHER THAN 20% OF ENM NODES (50)
[1574752351] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574752415] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3705) IS HIGHER THAN 20% OF ENM NODES (13969)
[1574752416] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (12) IS HIGHER THAN 20% OF ENM NODES (50)
[1574752474] SERVICE ALERT: ieatnetsimv7048-13;check netsim disk;OK;HARD;3;OK
[1574752480] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574752546] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (4391) IS HIGHER THAN 20% OF ENM NODES (13969)
[1574752555] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (18) IS HIGHER THAN 20% OF ENM NODES (50)
[1574752580] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (77) IS HIGHER THAN 95% OF ENM  NODES (80)
[1574752679] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (4395) IS HIGHER THAN 20% OF ENM NODES (13969)
[1574752875] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1574753375] SERVICE ALERT: ieatnetsimv7048-11;check netsim disk;OK;HARD;3;OK
[1574753376] SERVICE ALERT: ieatnetsimv7048-15;check netsim disk;OK;HARD;3;OK
[1574753389] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (28) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574753614] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1574753651] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (382) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574753681] SERVICE ALERT: ieatwlvm5116;check workload profiles;OK;HARD;3;OK- ALL WORKLOAD PROFILES ARE RUNNING
[1574753782] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (368) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574753829] SERVICE ALERT: ieatnetsimv7048-10;check netsim disk;OK;HARD;3;OK
[1574753834] SERVICE ALERT: ieatnetsimv7048-14;check netsim disk;OK;HARD;3;OK
[1574753908] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (1200)
[1574754111] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (24)
[1574754185] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1574754249] SERVICE ALERT: ieatlms3901;check cm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1574754433] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574754736] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1574754890] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;SOFT;1;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574754892] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (14) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1574754917] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1574755017] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1574755018] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;SOFT;2;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574755097] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1574755148] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574755722] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RNC;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1574756046] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574756162] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574756227] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574756294] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2001) IS BETWEEN 5% AND 20% OF ENM NODES (13969)
[1574756408] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574756467] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574756553] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574756608] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1574756735] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 4
[1574756861] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 1
[1574756989] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 2
[1574757331] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (46) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574757459] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (47) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574757967] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2550) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574757987] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574758099] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574758134] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574758181] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13940) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574758284] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574758397] HOST ALERT: ieatnetsimv5116-79;DOWN;SOFT;1;PING NOK
[1574758432] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574758459] HOST ALERT: ieatnetsimv5116-79;UP;SOFT;1;PING OK
[1574758470] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574758583] SERVICE ALERT: ieatnetsimv7035-33;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 27.42%
[1574758666] HOST ALERT: ieatnetsimv7035-103;DOWN;SOFT;1;PING NOK
[1574758728] HOST ALERT: ieatnetsimv7035-103;UP;SOFT;1;PING OK
[1574759466] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574759606] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574759697] SERVICE FLAPPING ALERT: ieatlms7435;check pm collection statistical DSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574759703] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (39) IS LOWER THAN 80% OF ENM  NODES (50)
[1574759733] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (70) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574759768] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1574759829] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (36) IS LOWER THAN 80% OF ENM  NODES (50)
[1574759903] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (115) IS LOWER THAN 5% OF ENM NODES (13937)
[1574759955] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (42) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574760593] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1574760596] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574760705] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.15%
[1574760722] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574760848] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574760856] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 24.65%
[1574761026] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1574761414] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1574761521] SERVICE ALERT: ieatnetsimv9041-07;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574761581] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (2026) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574761596] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (118) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574761672] SERVICE ALERT: ieatnetsimv9041-07;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574761710] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (2026) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574761716] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1574761722] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (118) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574761741] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574761760] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574761809] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (11233) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574761823] SERVICE ALERT: ieatnetsimv9041-07;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574761838] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (2038) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574761848] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (119) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574761869] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574761955] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (11233) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574761994] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574762074] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574762104] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (11233) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574762201] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574762329] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574762558] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (640) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574762572] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (11596) IS FROM 80% TO 95% OF ENM  NODES (13969)
[1574762685] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (651) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574762707] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (10960) IS LOWER THAN 80% OF ENM  NODES (13969)
[1574762816] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574762839] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (11172) IS FROM 80% TO 95% OF ENM  NODES (13937)
[1574762986] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574763025] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574763048] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574763111] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574763148] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574763198] SERVICE ALERT: ieatnetsimv7035-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 12.1%
[1574763287] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574763341] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574763417] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574763514] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1138) IS BETWEEN 5% AND 20% OF ENM NODES (13937)
[1574763550] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (41) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574763585] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574763644] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (630) IS LOWER THAN 5% OF ENM NODES (13969)
[1574763714] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (41) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574763845] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574764152] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (702) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574764204] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  ERBS: 9
[1574764292] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (704) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574764333] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  MGW: 25
[1574764421] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (703) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574764467] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:  RBS: 3
[1574764622] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (648) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574764713] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (144) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574764755] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (632) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574764838] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (170) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574764886] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574764962] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (186) IS BETWEEN 5% AND 20% OF ENM NODES (1200)
[1574765184] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (63) IS LOWER THAN 80% OF ENM  NODES (80)
[1574765316] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (79) IS HIGHER THAN 95% OF ENM  NODES (80)
[1574765346] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574765446] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2118) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574765455] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (125) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574765469] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1574765480] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574765554] SERVICE FLAPPING ALERT: ieatlms5735-1;check nas hc;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1574765554] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1574765600] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574765616] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574765698] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized Router6274;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1574765698] SERVICE ALERT: ieatlms5735-1;check synchronized Router6274;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574765819] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (7) IS HIGHER THAN 20% OF ENM NODES (30)
[1574765958] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1574766015] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574766139] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574766227] SERVICE ALERT: ieatlms3901;check synchronized MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (735) IS HIGHER THAN 95% OF ENM  NODES (750)
[1574766453] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13863) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574766947] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574766984] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574767049] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (141) IS FROM 80% TO 95% OF ENM  NODES (150)
[1574767079] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574767109] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (50)
[1574767154] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1574767174] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (141) IS FROM 80% TO 95% OF ENM  NODES (150)
[1574767206] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (78) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574767240] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574767300] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (141) IS FROM 80% TO 95% OF ENM  NODES (150)
[1574767364] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (76) IS HIGHER THAN 95% OF ENM  NODES (80)
[1574767426] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1574767521] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1574767645] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1574767772] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (80)
[1574768491] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (653) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574768510] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574768616] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574768637] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574768762] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574769108] SERVICE ALERT: ieatlms5735-1;check synchronized MGW;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (144) IS HIGHER THAN 95% OF ENM  NODES (150)
[1574769744] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1574769744] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN DSC NODES !
[1574769828] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (220) IS LOWER THAN 80% OF ENM  NODES (1200)
[1574769954] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (265) IS LOWER THAN 80% OF ENM  NODES (1200)
[1574770064] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1320) IS LOWER THAN 80% OF ENM  NODES (13969)
[1574770083] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (402) IS LOWER THAN 80% OF ENM  NODES (1200)
[1574770193] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1146) IS LOWER THAN 80% OF ENM  NODES (13937)
[1574770327] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13847) IS HIGHER THAN 95% OF ENM  NODES (13969)
[1574770473] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (16) IS LOWER THAN 80% OF ENM  NODES (30)
[1574770541] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (673) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574770596] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (26) IS FROM 80% TO 95% OF ENM  NODES (30)
[1574770666] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (674) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574770719] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (11) IS LOWER THAN 80% OF ENM  NODES (30)
[1574770757] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (28) IS LOWER THAN 80% OF ENM  NODES (50)
[1574770790] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (745) IS HIGHER THAN 95% OF ENM  NODES (750)
[1574770812] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574770833] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574770883] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (16) IS LOWER THAN 80% OF ENM  NODES (50)
[1574770939] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574770958] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574771008] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (34) IS LOWER THAN 80% OF ENM  NODES (50)
[1574771030] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (21) IS LOWER THAN 80% OF ENM  NODES (30)
[1574771065] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (75) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574771089] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (4) IS LOWER THAN 80% OF ENM  NODES (80)
[1574771170] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (15) IS LOWER THAN 80% OF ENM  NODES (30)
[1574771212] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574771293] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (17) IS LOWER THAN 80% OF ENM  NODES (30)
[1574771343] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (6) IS LOWER THAN 80% OF ENM  NODES (80)
[1574771473] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (6) IS LOWER THAN 80% OF ENM  NODES (80)
[1574771597] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574771656] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574771897] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1200)!
[1574772168] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (281) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574772382] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (11) IS HIGHER THAN 20% OF ENM NODES (30)
[1574772603] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (545) IS LOWER THAN 80% OF ENM  NODES (750)
[1574772730] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (747) IS HIGHER THAN 95% OF ENM  NODES (750)
[1574772935] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574773116] SERVICE ALERT: ieatnetsimv5116-33;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.17%
[1574773142] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574773266] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574773267] SERVICE ALERT: ieatnetsimv5116-33;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.24%
[1574773349] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574773396] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (100)
[1574773404] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (9) IS LOWER THAN 80% OF ENM  NODES (80)
[1574773417] SERVICE ALERT: ieatnetsimv5116-33;check netsim cpu;OK;HARD;3;OK- CPU IS USED 35.9%
[1574773473] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574773529] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574773598] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1574773669] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574773725] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1516) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574773867] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1548) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574773992] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1575) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574774324] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1574774598] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (2078) IS BETWEEN 5% AND 20% OF ENM NODES (13937)
[1574774614] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized EPG;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1574774614] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574774682] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (6) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574774752] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (2315) IS BETWEEN 5% AND 20% OF ENM NODES (13937)
[1574774814] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574774883] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures RadioNode;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1574774883] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (2708) IS BETWEEN 5% AND 20% OF ENM NODES (13937)
[1574774896] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MTAS;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1574774896] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1574774955] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574774991] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (21) IS HIGHER THAN 20% OF ENM NODES (80)
[1574775114] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (25) IS HIGHER THAN 20% OF ENM NODES (80)
[1574775237] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (80)
[1574775262] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574775388] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574775473] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574775473] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574775513] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574775797] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1749) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574776460] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (673) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574776584] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574776710] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (673) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574776852] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574776977] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (673) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574776999] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574777117] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574777124] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (10) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1574777201] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8) IS HIGHER THAN 20% OF ENM NODES (30)
[1574777248] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (100)
[1574777325] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8) IS HIGHER THAN 20% OF ENM NODES (30)
[1574777458] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574778492] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (39) IS LOWER THAN 5% OF ENM NODES (13937)
[1574779042] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (671) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574779166] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (671) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574779291] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574779394] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (60) IS LOWER THAN 5% OF ENM NODES (1200)
[1574780853] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1574781068] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8) IS HIGHER THAN 20% OF ENM NODES (30)
[1574781730] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574781872] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (76) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574781911] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574782093] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1574782093] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574782219] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1574782401] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1574782461] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574782581] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574782583] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1574782702] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574782974] HOST ALERT: ieatrcxb1289;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1574782999] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (278) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574783064] HOST ALERT: ieatrcxb1289;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1574783124] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574783154] HOST ALERT: ieatrcxb1289;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1574783187] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1574783244] HOST ALERT: ieatrcxb1289;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1574783305] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1574783309] HOST ALERT: ieatrcxb1289;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.210.244.243)
[1574783326] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574783373] HOST ALERT: ieatrcxb1289;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.210.244.243)
[1574783395] SERVICE ALERT: ieatlms7435;check unsync nodes;OK;HARD;3;OK
[1574783438] HOST ALERT: ieatrcxb1289;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.210.244.243)
[1574783502] HOST ALERT: ieatrcxb1289;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.210.244.243)
[1574783521] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574783567] HOST ALERT: ieatrcxb1289;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.210.244.243)
[1574783631] HOST ALERT: ieatrcxb1289;DOWN;HARD;10;CRITICAL - Host Unreachable (10.210.244.243)
[1574783711] SERVICE ALERT: ieatlms7435;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1574783715] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574783839] SERVICE ALERT: ieatlms7435;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1574783966] SERVICE ALERT: ieatlms7435;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1574784673] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574784699] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574784825] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (749) IS HIGHER THAN 95% OF ENM  NODES (750)
[1574784950] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574785074] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574785198] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574785323] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574785454] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574785479] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574785606] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574785734] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574785764] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574785765] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SBG-IS NODES !
[1574786302] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0*1=0)
[1574786332] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0*1=0)
[1574786338] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (0*1=0)
[1574786790] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1574786850] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (278) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574786913] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574786976] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (278) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574787042] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1574787042] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1574787102] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (1200)
[1574787395] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1574787697] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1574787992] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1574788277] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8) IS HIGHER THAN 20% OF ENM NODES (30)
[1574788748] SERVICE ALERT: ieatlms3901;check pm supervision nodes Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1200)
[1574788875] SERVICE ALERT: ieatlms3901;check pm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1574789129] SERVICE ALERT: ieatlms7435;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574789303] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3539) IS HIGHER THAN 20% OF ENM NODES (13969)
[1574789320] SERVICE ALERT: ieatlms7435;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574789369] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574789450] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (5) IS LOWER THAN 5% OF ENM NODES (13937)
[1574789498] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574789512] SERVICE ALERT: ieatlms7435;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574789580] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3539) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574789621] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574789709] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (5) IS LOWER THAN 5% OF ENM NODES (13937)
[1574790520] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (12545) IS LOWER THAN 80% OF ENM NODES (20060)
[1574790535] HOST ALERT: ieatrcxb1289;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1574790547] SERVICE ALERT: ieatlms7435;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574790622] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574790654] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (12545) IS LOWER THAN 80% OF ENM NODES (20060)
[1574790655] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (2035) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574790663] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (118) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574790751] SERVICE ALERT: ieatlms7435;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574790768] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574790788] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (12545) IS LOWER THAN 80% OF ENM NODES (20060)
[1574790868] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574790915] SERVICE ALERT: ieatlms7435;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574790996] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574791174] SERVICE ALERT: ieatlms7435;check vcs sg;OK;HARD;3;OK
[1574791384] SERVICE ALERT: ieatlms7435;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574791658] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (24) IS HIGHER THAN 20% OF ENM NODES (100)
[1574791783] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574791880] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574791908] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1574792789] SERVICE ALERT: ieatlms7435;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574792921] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574792982] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (20)
[1574793045] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574793104] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (20)
[1574793169] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574793225] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574793226] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (20) IS HIGHER THAN 20% OF ENM NODES (20)
[1574793293] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574793352] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1574793438] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3407) IS HIGHER THAN 20% OF ENM NODES (13969)
[1574793477] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574793514] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20060*1=20060)
[1574793536] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574793542] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574793566] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3407) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574793623] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574793660] SERVICE ALERT: ieatlms7435;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574793662] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (20000) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (20060*1=20060)
[1574793666] SERVICE ALERT: ieatlms7435;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574793696] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (4) IS LOWER THAN 5% OF ENM NODES (13937)
[1574793747] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574794265] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2189) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574794271] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (129) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574794733] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574794857] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574794985] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574795086] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574795212] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574795337] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574795484] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1574797305] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3400) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574797436] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3400) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574797567] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3400) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574797879] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (142) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574798945] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574799088] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1574799112] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574799214] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1574799237] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1574799338] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574799461] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (8) IS HIGHER THAN 20% OF ENM NODES (30)
[1574799585] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1574799916] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized Router6274;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574800443] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1574800616] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574800750] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574800874] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574800998] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574801123] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574801178] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (13969)
[1574801302] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574801447] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574801507] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (267) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574801634] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (267) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574801757] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574801954] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1574802928] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574802960] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574803052] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574803087] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1574803177] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574803188] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1574803210] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574803301] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (750)!
[1574803313] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574803334] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN FRONTHAUL-6080 NODES !
[1574803425] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574803436] SERVICE ALERT: ieatlms5735-1;check fm hb failures DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (9) IS HIGHER THAN 20% OF ENM NODES (30)
[1574803554] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574803678] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (670) IS FROM 80% TO 95% OF ENM  NODES (750)
[1574804046] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (17) IS HIGHER THAN 20% OF ENM NODES (80)
[1574804170] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SGSN-MME NODES !
[1574804551] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574804675] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (50)
[1574804787] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3399) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574804799] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574804917] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (13959)
[1574804927] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (50)
[1574805049] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3399) IS HIGHER THAN 20% OF ENM NODES (13959)
[1574805057] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574805086] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574805179] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3399) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574805184] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574805213] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574805312] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (13937)
[1574805320] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574805340] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574805365] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (267) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574805491] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574805615] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (267) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574805741] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (267) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574805864] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (267) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574806091] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574806117] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical DSC;STARTED; Service appears to have started flapping (31.1% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (33.5% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (45.7% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (29.9% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STARTED; Service appears to have started flapping (49.1% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures MTAS;STARTED; Service appears to have started flapping (27.2% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures RadioNode;STARTED; Service appears to have started flapping (33.6% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6672;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms5735-1;check nas hc;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1574806117] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized nodes cpp;STARTED; Service appears to have started flapping (21.9% change >= 20.0% threshold)
[1574806220] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574806351] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574806414] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574806476] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (1560) IS LOWER THAN 80% OF ENM NODES (2560)
[1574806543] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574806573] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574806575] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574806577] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574806579] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574806602] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;HARD;3;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1574806669] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574806698] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574806699] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574806702] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574806704] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574806746] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1560) IS LOWER THAN 80% OF ENM  NODES (2560)
[1574806823] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574806825] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574806827] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574806829] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574806871] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2559) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1574807023] SERVICE ALERT: ieatlms5735-1;check pm collection statistical EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574807024] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574807026] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1750*4=7000)
[1574807027] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20*1=20)
[1574807058] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574807153] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574807154] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20*1=20)
[1574807157] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1750*4=7000)
[1574807158] SERVICE ALERT: ieatlms5735-1;check pm collection statistical EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574807182] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574807280] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574807281] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20*1=20)
[1574807283] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1750*4=7000)
[1574807286] SERVICE ALERT: ieatlms5735-1;check pm collection statistical EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574807306] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574807341] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (76) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574807472] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574807473] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12*1=12)
[1574807475] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574807477] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574807597] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574807598] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12*1=12)
[1574807601] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574807602] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574807722] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574807723] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (12*1=12)
[1574807726] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574807729] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574807749] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574807874] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574807922] SERVICE ALERT: ieatlms5735-1;check pm collection statistical BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (98*1=98)
[1574807924] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574807925] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574807926] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13*1=13)
[1574807999] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574808047] SERVICE ALERT: ieatlms5735-1;check pm collection statistical BSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (98*1=98)
[1574808049] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574808050] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13*1=13)
[1574808060] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574808171] SERVICE ALERT: ieatlms5735-1;check pm collection statistical BSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (98*1=98)
[1574808174] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574808175] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13*1=13)
[1574808184] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574808652] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (6) IS LOWER THAN 5% OF ENM NODES (30)
[1574808676] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2441) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574808748] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13941*1=13941)
[1574808828] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574808878] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13941*1=13941)
[1574808921] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3399) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574808924] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574808932] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1574808932] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13301) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574808951] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (144) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574808953] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574809008] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13937*1=13937)
[1574809049] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (13942)
[1574809051] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574809060] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514*1=1514)
[1574809079] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574809180] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (2) IS LOWER THAN 5% OF ENM NODES (50)
[1574809186] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514*1=1514)
[1574809312] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514*1=1514)
[1574810958] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574811091] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (70) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574811219] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (70) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574811755] HOST ALERT: ieatnetsimv9041-33;DOWN;SOFT;1;PING NOK
[1574811817] HOST ALERT: ieatnetsimv9041-33;UP;SOFT;1;PING OK
[1574812255] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574812285] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2291) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574812380] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (5) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574812415] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2291) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574812511] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1574812553] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2291) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574812558] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (134) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574812560] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12595) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574812685] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (134) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574812709] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12595) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)