[1575158430] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.01%
[1575158581] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 42.8%
[1575158903] SERVICE ALERT: ieatnetsimv7102-16;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.68%
[1575159053] SERVICE ALERT: ieatnetsimv7102-16;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.05%
[1575159063] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (3) IS BETWEEN 5% AND 20% OF ENM NODES (30)
[1575159187] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (30)
[1575161668] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1575162502] SERVICE ALERT: ieatnetsimv7102-01;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.46%
[1575162652] SERVICE ALERT: ieatnetsimv7102-01;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 30.94%
[1575170855] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1612) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1575170981] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1605) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1575171114] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1612) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1575172918] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1741) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1575179128] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized DSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575181011] HOST ALERT: ieatnetsimv7048-06;DOWN;SOFT;1;PING NOK
[1575181074] HOST ALERT: ieatnetsimv7048-06;UP;SOFT;1;PING OK
[1575181130] HOST ALERT: ieatnetsimv9041-17;DOWN;SOFT;1;PING NOK
[1575181192] HOST ALERT: ieatnetsimv9041-17;UP;SOFT;1;PING OK
[1575181878] HOST ALERT: ieatnetsimv7035-44;DOWN;SOFT;1;PING NOK
[1575181941] HOST ALERT: ieatnetsimv7035-44;UP;SOFT;1;PING OK
[1575183763] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575183920] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575184076] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575185759] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized EPG;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575188293] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1575189052] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized RadioNode;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575189346] SERVICE ALERT: ieatnetsimv7102-41;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1575189496] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MTAS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575189497] SERVICE ALERT: ieatnetsimv7102-41;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.4%
[1575189512] SERVICE ALERT: ieatlms7555;check node fs;OK;HARD;3;OK
[1575189648] SERVICE ALERT: ieatnetsimv7102-41;check netsim cpu;OK;HARD;3;OK- CPU IS USED 78.24%
[1575202148] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575202304] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575202461] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575206296] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (73) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575206422] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (72) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575206567] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575208858] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (13056) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575209000] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (12912) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575209134] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (12647) IS FROM 80% TO 95% OF ENM  NODES (14007)
[1575209177] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575209298] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575209320] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575209339] SERVICE ALERT: ieatlms3901;check cm supervision nodes RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (5) IS LOWER THAN 80% OF ENM NODES (15)
[1575209425] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575209454] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1575209468] SERVICE ALERT: ieatlms3901;check cm supervision nodes RNC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (5) IS HIGHER THAN 95% OF ENM NODES (5)
[1575209554] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (27) IS FROM 80% TO 95% OF ENM  NODES (30)
[1575210171] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (76) IS HIGHER THAN 95% OF ENM  NODES (80)
[1575210499] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures Router6274;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1575211259] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1575211357] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (29) IS HIGHER THAN 95% OF ENM  NODES (30)
[1575211498] SERVICE FLAPPING ALERT: ieatlms7555;check node fs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1575211498] SERVICE ALERT: ieatlms7555;check node fs;OK;HARD;3;OK
[1575212749] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13699) IS HIGHER THAN 95% OF ENM  NODES (14007)
[1575215669] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1575215789] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1575215910] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1575217378] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (69) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575217512] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (71) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575217642] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (71) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575219451] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized SGSN-MME;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1575219451] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1575224855] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1575225013] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (77) IS HIGHER THAN 95% OF ENM  NODES (80)
[1575227734] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575227891] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575228047] SERVICE ALERT: ieatlms7555;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1575231111] HOST ALERT: ieatnetsimv7048-08;DOWN;SOFT;1;PING NOK
[1575231173] HOST ALERT: ieatnetsimv7048-08;UP;SOFT;1;PING OK
[1575233483] SERVICE ALERT: ieatlms7555;check node fs;OK;HARD;3;OK