[1568847769] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;OK;HARD;3;OK
[1568853766] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (16351) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (20362)
[1568859091] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2349) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568859224] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2445) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1568872947] SERVICE ALERT: ieatnetsimv5116-32;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568873576] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568874126] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1568874252] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1568875437] SERVICE ALERT: ieatnetsimv5116-33;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568875533] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1568875727] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1568875919] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1568876007] SERVICE ALERT: ieatnetsimv5116-36;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568876234] SERVICE ALERT: ieatnetsimv5116-34;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568876671] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568876697] SERVICE ALERT: ieatnetsimv5116-41;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568877054] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568877517] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568877792] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1568880434] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1568880618] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1568880803] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1568882138] SERVICE ALERT: ieatnetsimv5116-58;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568882463] SERVICE ALERT: ieatnetsimv5116-50;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568882603] SERVICE ALERT: ieatnetsimv5116-44;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568882734] SERVICE ALERT: ieatnetsimv5116-48;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568882740] SERVICE ALERT: ieatnetsimv5116-55;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568882831] SERVICE ALERT: ieatnetsimv5116-90;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568882873] SERVICE ALERT: ieatnetsimv5116-64;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568882926] SERVICE ALERT: ieatnetsimv5116-66;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568883235] SERVICE ALERT: ieatnetsimv5116-60;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568883255] SERVICE ALERT: ieatnetsimv5116-62;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568883518] SERVICE ALERT: ieatnetsimv5116-73;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568883838] SERVICE ALERT: ieatnetsimv5116-80;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568883921] SERVICE ALERT: ieatnetsimv5116-83;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568884098] SERVICE ALERT: ieatnetsimv5116-74;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568884129] SERVICE ALERT: ieatnetsimv5116-75;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568884348] SERVICE ALERT: ieatnetsimv5116-92;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1568885922] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568886069] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1568886713] SERVICE ALERT: ieatlms3901;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2350) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568886856] SERVICE ALERT: ieatlms3901;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2412) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1568887369] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1568887490] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1568887610] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1568889633] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1568889811] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1568889970] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1568889990] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1568890152] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1568890294] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568890334] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1568890444] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 47.27%
[1568891757] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (1984) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1568891889] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2010) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568892022] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2038) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568893835] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2434) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1568894355] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.89%
[1568894505] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 34.8%
[1568895648] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2143) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568895781] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2165) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568895796] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1568895913] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2194) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1568895958] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1568896776] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1568897490] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1568897643] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1568897734] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1568897734] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2500) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1568897793] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.46%
[1568899563] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (91) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1568899688] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (91) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1568899813] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (91) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1568900668] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1568900895] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1568902404] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568902423] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568902533] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1568902559] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568902575] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568902705] SERVICE ALERT: ieatnetsimv5116-30;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568902718] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1568902730] SERVICE ALERT: ieatnetsimv5116-31;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568902903] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1568903090] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3932) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1568903233] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3932) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1568903388] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3932) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1568903858] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (3931) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1568904052] SERVICE ALERT: ieatnetsimv5116-29;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1568904211] SERVICE ALERT: ieatnetsimv5116-29;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1568904256] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2417) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1568904358] SERVICE ALERT: ieatnetsimv5116-29;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1568904574] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (1544) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1568904768] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1568904768] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1568907904] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (15304) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (19614)
[1568912083] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (3931) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1568919145] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.05%
[1568919295] SERVICE ALERT: ieatnetsimv5116-39;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 39.55%
[1568926564] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1568930527] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1568930527] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (3931) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1568931946] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)