[1574812814] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (133) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574812858] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12565) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574813068] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6672 NODES !
[1574813296] HOST ALERT: ieatnetsimv5116-68;DOWN;SOFT;1;PING NOK
[1574813358] HOST ALERT: ieatnetsimv5116-68;UP;SOFT;1;PING OK
[1574813702] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;OK;HARD;3;OK
[1574814825] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (76) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574815094] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (17) IS HIGHER THAN 20% OF ENM NODES (80)
[1574815102] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (860) IS LOWER THAN 80% OF ENM  NODES (1200)
[1574815102] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;OK;HARD;3;OK
[1574815220] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (13) IS BETWEEN 5% AND 20% OF ENM NODES (80)
[1574815226] SERVICE ALERT: ieatlms5735-1;check synchronized Router6672;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (1200)!
[1574815340] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (11655) IS FROM 80% TO 95% OF ENM  NODES (13937)
[1574815345] SERVICE ALERT: ieatlms5735-1;check fm hb failures SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (17) IS HIGHER THAN 20% OF ENM NODES (80)
[1574815475] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (11388) IS FROM 80% TO 95% OF ENM  NODES (13937)
[1574815607] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (11917) IS FROM 80% TO 95% OF ENM  NODES (13937)
[1574815611] HOST ALERT: ieatnetsimv7035-42;DOWN;SOFT;1;PING NOK
[1574815673] HOST ALERT: ieatnetsimv7035-42;UP;SOFT;1;PING OK
[1574816025] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574816116] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (7) IS HIGHER THAN 20% OF ENM NODES (30)
[1574816153] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (41) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574816162] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2483) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574816288] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (43) IS FROM 80% TO 95% OF ENM  NODES (50)
[1574816379] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (4910) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574816384] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (7) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574816509] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3395) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574816516] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (13) IS HIGHER THAN 20% OF ENM NODES (50)
[1574816638] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (3395) IS HIGHER THAN 20% OF ENM NODES (13937)
[1574816643] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (18) IS HIGHER THAN 20% OF ENM NODES (50)
[1574816671] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (266) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574816795] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (266) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574816920] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (349) IS HIGHER THAN 20% OF ENM NODES (1200)
[1574817422] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13917) IS HIGHER THAN 95% OF ENM  NODES (13937)
[1574818095] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1574818110] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1574818110] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (3) IS LOWER THAN 5% OF ENM NODES (100)
[1574818431] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574818558] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574818684] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574819771] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2272) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574819899] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2272) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574820028] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2284) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574820250] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (8) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1574820802] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1457) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574820903] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574820931] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1745) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574821023] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574821054] HOST ALERT: ieatnetsimv7035-03;DOWN;SOFT;1;PING NOK
[1574821055] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1534) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1574821116] HOST ALERT: ieatnetsimv7035-03;UP;SOFT;1;PING OK
[1574821144] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574821179] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (1745) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1574821713] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574821841] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574821964] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (26) IS HIGHER THAN 20% OF ENM NODES (100)
[1574823789] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574830054] HOST ALERT: ieatnetsimv7035-16;DOWN;SOFT;1;PING NOK
[1574830116] HOST ALERT: ieatnetsimv7035-16;UP;SOFT;1;PING OK
[1574830992] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574831116] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574831240] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574832874] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized nodes cpp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574834006] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MTAS;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574834297] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (72) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574834353] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1574834542] HOST ALERT: ieatnetsimv017-32;DOWN;SOFT;1;PING NOK
[1574834542] HOST ALERT: ieatnetsimv017-33;DOWN;SOFT;1;PING NOK
[1574834605] HOST ALERT: ieatnetsimv017-33;UP;SOFT;1;PING OK
[1574834605] HOST ALERT: ieatnetsimv017-32;UP;SOFT;1;PING OK
[1574835476] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574835629] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (22) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574835632] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;OK;HARD;3;OK- NUMBER OF PM FILES (1199) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574835634] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (651) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574836121] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574836214] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574836396] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574836505] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1574836529] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (72) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (12*1=12)
[1574836529] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;OK;HARD;3;OK- NUMBER OF PM FILES (1000) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574836533] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (3) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574836535] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574836576] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574836979] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-BC-IS;OK;HARD;3;OK- NUMBER OF PM FILES (72) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13*1=13)
[1574836981] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (100) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574836990] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;OK;HARD;3;OK- NUMBER OF PM FILES (50) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574837829] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13698) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935*1=13935)
[1574837904] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574839245] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (579) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750*1=750)
[1574840116] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.59%
[1574840267] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574840282] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1574840284] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574840418] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.92%
[1574840457] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574840551] SERVICE ALERT: ieatnetsimv7035-103;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.67%
[1574840560] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.76%
[1574840644] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1574840701] SERVICE ALERT: ieatnetsimv7035-103;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.66%
[1574840711] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.96%
[1574840862] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.98%
[1574841004] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1574841012] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.14%
[1574841450] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (10281) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935*1=13935)
[1574841509] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (63) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574841591] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (10281) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13937*1=13937)
[1574841623] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (120) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574841635] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574841639] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (2039) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574841722] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (6844) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13937*1=13937)
[1574841760] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1574841760] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574842684] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (1307) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574842811] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (1307) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574842835] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (16) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574842838] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (660) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574842938] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1337) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574842963] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (16) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574842965] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (900) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574843089] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (19) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574843092] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (900) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574843326] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (22) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574843452] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (15) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574843577] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (15) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574843737] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (247) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574843739] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (2) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574843741] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (74) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574843867] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (2) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574843872] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (88) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574843875] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (247) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574843919] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical DSC;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1574843993] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574844003] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (113) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574844026] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (247) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574844049] SERVICE ALERT: ieatnetsimv5116-31;check netsim cpu;OK;HARD;3;OK- CPU IS USED 51.59%
[1574844152] SERVICE ALERT: ieatnetsimv7035-104;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.11%
[1574844187] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (75) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574844195] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (38) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574844303] SERVICE ALERT: ieatnetsimv7035-104;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.59%
[1574844313] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (50) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574844323] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (24) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574844441] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (50) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574844450] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (24) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574845086] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (78) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574845230] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (134) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574845247] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2216) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574845378] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574845505] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (140) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574846693] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (18) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574846800] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574846817] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (18) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574846961] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (18) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574846985] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574847179] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574847743] SERVICE ALERT: ieatnetsimv7035-100;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.55%
[1574847893] SERVICE ALERT: ieatnetsimv7035-100;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 27.09%
[1574848044] SERVICE ALERT: ieatnetsimv7035-100;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.3%
[1574848195] SERVICE ALERT: ieatnetsimv7035-100;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.04%
[1574849053] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574849111] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574851074] SERVICE ALERT: ieatwlvm9034;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1574851684] SERVICE ALERT: ieatwlvm9034;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1574852292] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574852349] SERVICE ALERT: ieatwlvm5116;check workload profiles;OK;HARD;3;OK- ALL WORKLOAD PROFILES ARE RUNNING
[1574852423] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (79) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574852458] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2551) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574852488] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13907) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574853096] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.48%
[1574853182] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (74) IS FROM 80% TO 95% OF ENM  NODES (80)
[1574853246] SERVICE ALERT: ieatnetsimv9034-116;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 10.75%
[1574853308] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1574854239] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (91) IS FROM 80% TO 95% OF ENM  NODES (98)
[1574854364] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (95) IS HIGHER THAN 95% OF ENM  NODES (98)
[1574854529] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.03%
[1574854682] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 58.01%
[1574854727] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574854878] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 29.32%
[1574855625] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2401) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1574855643] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1574855765] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2552) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1574856145] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1574856145] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13701) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935*1=13935)
[1574856166] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574857229] SERVICE ALERT: ieatlms5735-1;check pm collection statistical CISCO-ASR900;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574857345] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574857345] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574857502] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6672;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574857502] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6672;OK;HARD;3;OK- NUMBER OF PM FILES (1199) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1200*1=1200)
[1574857807] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1574857807] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (22) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24*1=24)
[1574857981] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574857981] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30*1=30)
[1574858415] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574858415] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574858431] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ESC;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574858431] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ESC;OK;HARD;3;OK- NUMBER OF PM FILES (1000) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (1000*1=1000)
[1574858847] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574858847] SERVICE ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF PM FILES (100) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (100*1=100)
[1574858857] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6274;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574858857] SERVICE ALERT: ieatlms5735-1;check pm collection statistical Router6274;OK;HARD;3;OK- NUMBER OF PM FILES (50) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (50*1=50)
[1574859075] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574859225] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.37%
[1574859918] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574860103] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574860289] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574861199] SERVICE ALERT: ieatnetsimv9034-118;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574861349] SERVICE ALERT: ieatnetsimv9034-118;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2.9%
[1574861596] SERVICE ALERT: ieatnetsimv9041-11;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1574861747] SERVICE ALERT: ieatnetsimv9041-11;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1574861837] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1574861843] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1574861897] SERVICE ALERT: ieatnetsimv9041-11;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1574862002] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RNC;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1574862002] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RNC;OK;HARD;3;OK- NUMBER OF PM FILES (3) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (4*1=4)
[1574862154] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1574862154] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574862204] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH PM SUPERVISION (750) IS BETWEEN 80% AND 95% OF ENM NODES (811)
[1574862330] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH PM SUPERVISION (750) IS BETWEEN 80% AND 95% OF ENM NODES (810)
[1574862385] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (750) IS BETWEEN 80% AND 95% OF ENM NODES (810)
[1574862485] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (750) IS BETWEEN 80% AND 95% OF ENM NODES (811)
[1574862511] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (750) IS BETWEEN 80% AND 95% OF ENM NODES (811)
[1574862637] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (750) IS BETWEEN 80% AND 95% OF ENM NODES (811)
[1574863266] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2312) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574863315] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12564) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574863358] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (70) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574863395] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2312) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574863463] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12669) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574863485] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574863518] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574863524] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2335) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560*1=2560)
[1574863611] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12710) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13965*1=13965)
[1574863612] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80*1=80)
[1574863644] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574863645] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574863776] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (139) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150*1=150)
[1574864563] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574864656] SERVICE FLAPPING ALERT: ieatlms5735-1;check fm hb failures DSC;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1574864683] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574864804] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574865454] SERVICE ALERT: ieatnetsimv7035-97;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.31%
[1574865528] SERVICE ALERT: ieatnetsimv9041-11;check netsim cpu;OK;HARD;3;OK- CPU IS USED 33.33%
[1574865605] SERVICE ALERT: ieatnetsimv7035-97;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 4.25%
[1574865737] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (44.8% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (43.5% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ESC;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RNC;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6274;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6672;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1574865737] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1574865782] SERVICE ALERT: ieatlms5735-1;check pm collection statistical BSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (98*=)
[1574866008] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (44.8% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (43.5% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ESC;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RNC;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6274;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6672;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1574866008] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574866042] HOST ALERT: ieatnetsimv5116-02;DOWN;SOFT;1;PING NOK
[1574866102] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (558) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750*=)
[1574866104] HOST ALERT: ieatnetsimv5116-02;UP;SOFT;1;PING OK
[1574866520] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (44.8% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (43.5% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ESC;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RNC;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6274;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6672;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1574866520] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574866544] SERVICE ALERT: ieatlms3901;check number nodes ieatnetsimv017 DSC;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574866555] SERVICE ALERT: ieatlms3901;check number nodes ieatnetsimv5116 DSC;CRITICAL;SOFT;1;(No output on stdout) stderr: Warning: Permanently added '141.137.208.23' (RSA) to the list of known hosts.
[1574866627] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (44.8% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (43.5% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ESC;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RNC;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6274;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6672;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1574866627] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574866661] SERVICE ALERT: ieatlms3901;check number nodes ieatnetsimv5116 DSC;OK;SOFT;2;Files /usr/local/nagios/libexec/s4/list_check_number_nodes/check_ieatnetsimv5116.netsim and /usr/local/nagios/libexec/s4/list_check_number_nodes/host_ieatnetsimv5116.netsim differ
[1574866666] SERVICE ALERT: ieatlms3901;check number nodes ieatnetsimv017 DSC;OK;SOFT;2;Files /usr/local/nagios/libexec/s4/list_check_number_nodes/check_ieatnetsimv017.netsim and /usr/local/nagios/libexec/s4/list_check_number_nodes/host_ieatnetsimv017.netsim differ
[1574866737] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (44.8% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (43.5% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ESC;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RNC;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6274;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6672;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1574866737] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574866842] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (3849*=)
[1574866856] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (44.8% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (43.5% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ESC;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical FRONTHAUL-6080;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MTAS;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RNC;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6274;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical Router6672;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1574866856] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1574866917] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (608) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514*=)
[1574867133] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2317) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560*=)
[1574867142] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF PM FILES (2040) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1750*=)
[1574867212] SERVICE ALERT: ieatlms3901;check pm collection statistical EPG;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (50*=)
[1574867230] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (70) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80*=)
[1574867239] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (12628) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965*=)
[1574867242] SERVICE ALERT: ieatlms3901;check pm collection statistical BSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (96*=)
[1574867250] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574867378] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574867392] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (138) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150*=)
[1574867415] SERVICE ALERT: ieatlms7435;check synchronized RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (18516) IS FROM 80% TO 95% OF ENM  NODES (20060)
[1574867549] SERVICE ALERT: ieatlms7435;check synchronized RadioNode;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (19241) IS HIGHER THAN 95% OF ENM  NODES (20060)
[1574867553] HOST ALERT: ieatwlvm7002;DOWN;SOFT;1;PING NOK
[1574867615] HOST ALERT: ieatwlvm7002;UP;SOFT;1;PING OK
[1574868050] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1574868050] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-6352;OK;HARD;3;OK- NUMBER OF PM FILES (463) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (750*=)
[1574868142] HOST ALERT: ieatnetsimv017-35;DOWN;SOFT;1;PING NOK
[1574868204] HOST ALERT: ieatnetsimv017-35;UP;SOFT;1;PING OK
[1574868404] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1574868485] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SBG-IS;OK;HARD;3;TOTAL NUMBER OF PM FILES: 0
[1574868490] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MINI-LINK-Indoor;OK;HARD;3;TOTAL NUMBER OF PM FILES: 0
[1574868494] SERVICE ALERT: ieatlms5735-1;check pm collection statistical EPG;OK;HARD;3;TOTAL NUMBER OF PM FILES: 0
[1574869035] SERVICE ALERT: ieatlms7555;check pm collection statistical MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF PM FILES (19997) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (5000*=)
[1574869123] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30*=)
[1574869218] SERVICE ALERT: ieatlms5735-1;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (15) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30*=)
[1574869366] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (5000) IS FROM 80% TO 95% OF ENM  NODES (5850)
[1574869469] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574869491] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (5000) IS FROM 80% TO 95% OF ENM  NODES (6020)
[1574869535] SERVICE ALERT: ieatlms7555;check pm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH PM SUPERVISION (5000) IS BETWEEN 80% AND 95% OF ENM NODES (6020)
[1574869606] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574869617] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (4999) IS FROM 80% TO 95% OF ENM  NODES (6020)
[1574869652] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (5000) IS BETWEEN 80% AND 95% OF ENM NODES (6020)
[1574869659] SERVICE ALERT: ieatlms7555;check pm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH PM SUPERVISION (5000) IS BETWEEN 80% AND 95% OF ENM NODES (6020)
[1574869692] SERVICE ALERT: ieatlms7555;check cm supervision nodes MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH CM SUPERVISION (5000) IS BETWEEN 80% AND 95% OF ENM NODES (6020)
[1574869744] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1574869778] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (5000) IS LOWER THAN 80% OF ENM NODES (6380)
[1574869785] SERVICE ALERT: ieatlms7555;check pm supervision nodes MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (5000) IS LOWER THAN 80% OF ENM NODES (6440)
[1574869818] SERVICE ALERT: ieatlms7555;check cm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (5000) IS LOWER THAN 80% OF ENM NODES (6676)
[1574869904] SERVICE ALERT: ieatlms7555;check fm supervision nodes MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (5000) IS LOWER THAN 80% OF ENM NODES (7160)
[1574869943] SERVICE ALERT: ieatlms7555;check cm supervision nodes MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (5000) IS LOWER THAN 80% OF ENM NODES (7160)
[1574870183] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1574871103] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1574871346] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1574871423] SERVICE ALERT: ieatlms7555;check synchronized MINI-LINK-Indoor;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (4999) IS LOWER THAN 80% OF ENM  NODES (7160)
[1574871639] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1574872024] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;CRITICAL;HARD;3;(Service check timed out after 800.96 seconds)
[1574872266] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1574872559] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;CRITICAL;SOFT;2;(Service check timed out after 800.01 seconds)
[1574872696] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13965*=)
[1574872703] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574872823] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574872943] SERVICE ALERT: ieatnetsimv7048-12;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574873186] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1574874489] SERVICE ALERT: ieatlms7435;check fm supervision nodes RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (0) IS HIGHER THAN 95% OF ENM NODES (0)
[1574874631] SERVICE ALERT: ieatlms7435;check unsync nodes;OK;HARD;3;OK
[1574874798] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574874987] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574874992] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1574875177] SERVICE ALERT: ieatlms7435;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1574875181] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1574875254] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574875371] SERVICE ALERT: ieatlms7435;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1574875448] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574875642] SERVICE ALERT: ieatlms7435;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574876027] SERVICE ALERT: ieatlms7435;check fm hb failures SBG-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN SBG-IS NODES !
[1574877713] HOST ALERT: ieatnetsimv7035-61;DOWN;SOFT;1;PING NOK
[1574877775] HOST ALERT: ieatnetsimv7035-61;UP;SOFT;1;PING OK
[1574878301] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574878427] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (1) IS BETWEEN 5% AND 20% OF ENM NODES (13)
[1574878550] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-IS;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-IS NODES !
[1574879245] SERVICE ALERT: ieatlms7435;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (20000) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0*=)
[1574880407] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13734) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935*=)
[1574882627] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1574882821] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1574883016] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1574885156] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1574885244] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1574885276] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1574885364] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1574885397] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1574885422] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1574885485] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1574885543] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1574885555] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574885664] SERVICE FLAPPING ALERT: ieatlms5735-1;check nas hc;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1574885664] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1574885675] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574885796] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574886110] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1574886117] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1574886179] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574886218] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1574886231] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1574886238] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1574886300] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574886338] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1574886352] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1574886359] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1574886405] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1574886421] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1574886459] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1574886525] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1574886646] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1574886690] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1574887347] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1574887662] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574888283] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1574889055] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1574889453] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1574889454] HOST ALERT: ieatnetsimv7035-91;DOWN;SOFT;1;PING NOK
[1574889454] HOST ALERT: ieatnetsimv7035-97;DOWN;SOFT;1;PING NOK
[1574889454] HOST ALERT: ieatnetsimv7035-08;DOWN;SOFT;1;PING NOK
[1574889484] HOST ALERT: ieatnetsimv7035-91;UP;SOFT;1;PING OK
[1574889516] HOST ALERT: ieatnetsimv7035-97;UP;SOFT;1;PING OK
[1574889516] HOST ALERT: ieatnetsimv7035-08;UP;SOFT;1;PING OK
[1574889527] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1574889712] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1574889898] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1574890020] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1574890038] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1574890246] SERVICE FLAPPING ALERT: ieatnetsimv7035-21;check netsim disk;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1574890246] SERVICE ALERT: ieatnetsimv7035-21;check netsim disk;OK;HARD;3;OK
[1574892575] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (24)
[1574892626] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (98)
[1574892778] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1574892876] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (97) IS HIGHER THAN 95% OF ENM  NODES (98)
[1574893564] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1574897599] SERVICE ALERT: ieatlms3901;check pm supervision nodes MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (750) IS LOWER THAN 80% OF ENM NODES (958)
[1574897743] SERVICE ALERT: ieatlms3901;check fm supervision nodes MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (750) IS LOWER THAN 80% OF ENM NODES (961)