[1565565280] SERVICE ALERT: ieatnetsimv7035-13;check netsim disk;OK;HARD;3;OK
[1565565930] HOST ALERT: ieatnetsimv5116-31;DOWN;SOFT;1;PING NOK
[1565565992] HOST ALERT: ieatnetsimv5116-31;UP;SOFT;1;PING OK
[1565568886] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565569183] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565570407] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs sg;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1565572439] HOST ALERT: ieatnetsimv5116-25;DOWN;SOFT;1;PING NOK
[1565572501] HOST ALERT: ieatnetsimv5116-25;UP;SOFT;1;PING OK
[1565573674] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized comecim;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1565577079] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2359) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565577220] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2330) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565577358] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2343) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565578514] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565579083] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565579169] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1565579169] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2400) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1565579509] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565583575] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1565588228] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565588409] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565588478] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1565588590] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565588660] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1565588841] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1565590155] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1565590215] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1565590343] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1565590470] SERVICE ALERT: ieatlms5735-1;check vcs sg;OK;HARD;3;OK
[1565590595] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1565590839] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1565591102] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1565591378] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1565591590] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1565591614] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1565591786] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2368) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565591917] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;2;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2369) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1565592048] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2386) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1565592251] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1565592371] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1565592504] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1565592527] SERVICE ALERT: ieatnetsimv7035-60;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565592656] SERVICE ALERT: ieatnetsimv7035-60;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1565592857] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565593101] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565593186] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565593394] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565593572] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1565594217] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2274) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1565594409] SERVICE ALERT: ieatnetsimv7035-83;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565595308] SERVICE ALERT: ieatnetsimv7035-11;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565596161] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565596336] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565596510] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565597639] SERVICE ALERT: ieatnetsimv5116-52;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565597760] SERVICE ALERT: ieatnetsimv5116-52;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565597823] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565597841] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565597880] SERVICE ALERT: ieatnetsimv5116-52;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565597926] SERVICE ALERT: ieatnetsimv7035-61;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565598055] SERVICE ALERT: ieatnetsimv7035-61;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565598184] SERVICE ALERT: ieatnetsimv7035-61;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565599554] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14033) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565599861] SERVICE ALERT: ieatnetsimv7035-105;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565600169] SERVICE ALERT: ieatnetsimv5116-52;check netsim disk;OK;HARD;3;OK
[1565600615] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (7011) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1565601993] SERVICE ALERT: ieatnetsimv5116-77;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1565606824] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565607013] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565607202] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565608688] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565608874] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565609059] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565612692] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1565612692] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2515) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565613659] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1565613792] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1565613929] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1565615669] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.36%
[1565615819] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 61.43%
[1565618022] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14033) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565619571] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.42%
[1565619723] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.51%
[1565619874] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.92%
[1565621710] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565621899] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565622088] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1565622088] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565623557] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565623684] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2515) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1565624075] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1565624255] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1565624436] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1565625154] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1565625334] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1565625512] SERVICE ALERT: ieatlms3901;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1565625529] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565625709] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565625890] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1565626013] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1565626200] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1565626388] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1565626729] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.41%
[1565626880] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.59%
[1565627030] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 91.42%
[1565627298] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1565627481] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1565627664] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1565629310] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14033) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1565629460] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6504) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1565629526] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1565629603] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (7128) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1565629709] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1565629892] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1565631940] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized snmp;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1565632653] SERVICE ALERT: ieatlms3901;check pm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH PM SUPERVISION (23796) IS HIGHER THAN 95% OF ENM NODES (24755)
[1565633820] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23796) IS BETWEEN 80% AND 95% OF ENM NODES (25505)
[1565634126] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23796) IS BETWEEN 80% AND 95% OF ENM NODES (25505)
[1565634352] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.7%
[1565634355] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23796) IS BETWEEN 80% AND 95% OF ENM NODES (25505)
[1565634502] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.52%
[1565634653] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.7%
[1565634803] SERVICE ALERT: ieatnetsimv7035-10;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.73%
[1565637155] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1565637340] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1565637525] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1565637860] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.66%
[1565639022] SERVICE ALERT: ieatlms3901;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11311) IS FROM 80% TO 95% OF ENM COMECIM NODES (12128)
[1565639078] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;1;CRITICAL- THERE ARE FAULTED SG!
[1565639182] SERVICE ALERT: ieatlms3901;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (11923) IS HIGHER THAN 95% OF ENM COMECIM NODES (12128)
[1565639206] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;SOFT;2;CRITICAL- THERE ARE FAULTED SG!
[1565639334] SERVICE ALERT: ieatlms5735-1;check vcs sg;CRITICAL;HARD;3;CRITICAL- THERE ARE FAULTED SG!
[1565639504] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1565639689] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1565639877] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1565641174] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1565641491] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.86%
[1565641531] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1565641641] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.44%
[1565641834] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1565643072] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1565643321] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1565643574] SERVICE ALERT: ieatnetsimv5116-84;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1565644789] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1565645969] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (5322) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1565646134] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (7142) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1565648039] SERVICE ALERT: ieatnetsimv5116-54;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1565648087] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1565648159] SERVICE ALERT: ieatnetsimv5116-54;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1565648280] SERVICE ALERT: ieatnetsimv5116-54;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1565648992] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.37%
[1565649142] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.38%
[1565649293] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.7%