[1571182396] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571182478] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1571182524] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571182602] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (27) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1571182627] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571182652] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2320) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571182727] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1571182753] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571182882] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571182942] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12569) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13950)
[1571183089] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12569) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13950)
[1571183236] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12558) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13950)
[1571186261] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571186489] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571186863] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13795) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13950)
[1571187750] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12037) IS FROM 80% TO 95% OF ENM COMECIM NODES (14153)
[1571187892] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (11950) IS FROM 80% TO 95% OF ENM COMECIM NODES (14153)
[1571188036] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13915) IS HIGHER THAN 95% OF ENM COMECIM NODES (14153)
[1571188946] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571189072] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571189199] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571189869] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2339) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571189998] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571190095] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571190128] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2339) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571190221] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (137) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571190348] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (144) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571191764] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1571192829] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571195259] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13280) IS FROM 80% TO 95% OF ENM COMECIM NODES (14153)
[1571195402] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13918) IS HIGHER THAN 95% OF ENM COMECIM NODES (14151)
[1571196294] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1571197340] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571201154] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (133) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571201281] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571201290] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13147) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571201408] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (146) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571201438] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (13394) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571202564] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1571202749] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1571205447] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571205597] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.67%
[1571206303] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.04%
[1571206387] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.37%
[1571206454] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.99%
[1571206538] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.48%
[1571206605] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571206689] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 93.53%
[1571207273] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571207397] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571207403] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;SOFT;2;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571207407] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3849) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571207527] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF PM FILES (24) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571208165] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571208376] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571209126] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571209946] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized comecim;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1571212624] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571212775] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571212926] SERVICE ALERT: ieatnetsimv5116-30;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.28%
[1571213920] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;HARD;3;OK- CPU IS USED 78.57%
[1571214067] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571214078] SERVICE ALERT: ieatnetsimv7035-17;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571214220] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571214229] SERVICE ALERT: ieatnetsimv7035-17;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 4.84%
[1571214271] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571214372] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571214414] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571214422] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571214423] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571214565] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571214570] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571214574] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571214574] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571214717] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571214722] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571214726] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571214874] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571215113] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215157] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215170] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215224] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215264] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571215307] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 8.82%
[1571215314] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215322] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571215375] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 9.67%
[1571215405] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215415] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571215417] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215465] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571215473] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215473] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571215498] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215556] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571215564] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571215569] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571215605] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215617] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571215625] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571215650] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571215696] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1571215707] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.42%
[1571215707] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571215720] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571215757] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571215776] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571215801] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;OK;HARD;3;OK- CPU IS USED 80.56%
[1571215825] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1571215858] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571215908] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571216009] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571216548] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571216699] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571216850] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571216984] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571216995] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571217136] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571217146] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571217283] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571217287] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571217297] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571217435] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571217586] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571218205] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.41%
[1571218504] SERVICE ALERT: ieatnetsimv7035-83;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.7%
[1571219045] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.83%
[1571219104] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.26%
[1571219351] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.53%
[1571219369] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571219407] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.68%
[1571219520] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571219671] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571219674] SERVICE ALERT: ieatnetsimv7035-03;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571219729] SERVICE ALERT: ieatnetsimv7035-01;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571219738] SERVICE ALERT: ieatnetsimv7035-04;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1571219825] SERVICE ALERT: ieatnetsimv7035-03;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571219880] SERVICE ALERT: ieatnetsimv7035-01;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571219889] SERVICE ALERT: ieatnetsimv7035-04;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571219976] SERVICE ALERT: ieatnetsimv7035-03;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571220031] SERVICE ALERT: ieatnetsimv7035-01;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571220044] SERVICE ALERT: ieatnetsimv7035-04;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1571220297] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571220417] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571220422] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571220546] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571221035] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.61%
[1571221074] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6439) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571221930] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571222052] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571222056] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571222178] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571222183] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571222303] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571222548] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2327) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571222676] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2327) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571222765] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (72) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571222805] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2552) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571222889] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (72) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571223014] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (72) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571223662] SERVICE ALERT: ieatnetsimv7035-01;check netsim cpu;OK;HARD;3;OK- CPU IS USED 14.23%
[1571223676] SERVICE ALERT: ieatnetsimv7035-04;check netsim cpu;OK;HARD;3;OK- CPU IS USED 40.8%
[1571223914] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1571224096] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1571224151] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571224277] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1571224666] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.54%
[1571224816] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.28%
[1571224861] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571226512] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (36.2% change >= 20.0% threshold)
[1571226721] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1571226846] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (36.2% change >= 20.0% threshold)
[1571226937] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1571227118] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1571227151] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571227162] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1571227162] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571227255] SERVICE ALERT: ieatnetsimv5116-04;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1571227504] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571227731] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (29) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1571227753] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571227878] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571227980] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3999) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571228002] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571228127] SERVICE ALERT: ieatnetsimv7035-03;check netsim cpu;OK;HARD;3;OK- CPU IS USED 13.01%
[1571228221] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.09%
[1571228223] SERVICE ALERT: ieatnetsimv7035-48;check netsim cpu;OK;HARD;3;OK- CPU IS USED 8.91%
[1571228231] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 53.47%
[1571228234] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.23%
[1571228345] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (3471) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571228690] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571229060] SERVICE ALERT: ieatnetsimv5116-06;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1571229061] SERVICE ALERT: ieatnetsimv5116-05;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1571231336] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1571231336] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571231725] SERVICE ALERT: ieatnetsimv7035-02;check netsim cpu;OK;HARD;3;OK- CPU IS USED 10.11%
[1571231858] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;OK;HARD;3;OK- CPU IS USED 7.75%
[1571231860] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.22%
[1571231878] SERVICE ALERT: ieatnetsimv7035-74;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.26%
[1571231881] SERVICE ALERT: ieatnetsimv7035-75;check netsim cpu;OK;HARD;3;OK- CPU IS USED 15.86%
[1571231955] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (3846) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571232579] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1571233180] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1571234329] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (134) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571234336] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2407) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571234369] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (13027) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571234457] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;SOFT;2;OK- NUMBER OF PM FILES (145) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571234468] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571234522] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;SOFT;2;OK- NUMBER OF PM FILES (13807) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571235424] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 5.27%
[1571235428] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 12.84%
[1571236204] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1571236385] SERVICE ALERT: ieatlms5735-1;check node fs;OK;SOFT;2;OK
[1571236604] SERVICE ALERT: ieatnetsimv7035-47;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571236608] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571236616] SERVICE ALERT: ieatnetsimv7035-50;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571236622] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571236632] SERVICE ALERT: ieatnetsimv7035-56;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571236675] SERVICE ALERT: ieatnetsimv7035-57;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571236694] SERVICE ALERT: ieatnetsimv7035-70;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1571236732] SERVICE ALERT: ieatnetsimv7035-47;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571236736] SERVICE ALERT: ieatnetsimv7035-48;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571236744] SERVICE ALERT: ieatnetsimv7035-50;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571236750] SERVICE ALERT: ieatnetsimv7035-55;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571236760] SERVICE ALERT: ieatnetsimv7035-56;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571236803] SERVICE ALERT: ieatnetsimv7035-57;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571236822] SERVICE ALERT: ieatnetsimv7035-70;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1571238774] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.34%
[1571238925] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 69.6%
[1571239076] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.44%
[1571239226] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 57.02%
[1571239381] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.38%
[1571239532] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 69.19%
[1571239701] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (150) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571240197] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2590) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571240945] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6559) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571241615] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571241786] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (130) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571241804] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2336) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571241813] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571241871] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12722) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571241912] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (133) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571241933] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2365) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571242009] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571242018] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12722) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571242038] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (133) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571242061] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2365) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571242147] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571242170] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12722) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571242273] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571242398] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (73) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571244810] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2590) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571245490] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6409) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571245646] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1571245646] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571245669] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571245686] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571245686] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (150) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571245753] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571245798] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1571245798] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13933) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571245881] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (30) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571246005] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571246009] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1571246009] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571246108] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (30) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571246368] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (3517) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571246498] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (3519) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571246627] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (3519) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571247563] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571247759] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571247955] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571249252] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571249278] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2382) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571249379] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571249406] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2382) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571249435] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12926) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571249505] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (138) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571249534] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2382) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571249583] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12926) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571249611] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571249743] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (74) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571249751] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12947) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571249818] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.78%
[1571249871] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (76) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571249970] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1571250121] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;OK;HARD;3;OK- CPU IS USED 44.29%
[1571250236] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RBS;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1571250236] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (3807) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571251633] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1571251633] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571253479] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571253606] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571253732] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571253751] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.68%
[1571253846] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (3471) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571253902] SERVICE ALERT: ieatnetsimv7035-34;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 13.47%
[1571253974] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (3471) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571254103] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (3471) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571255972] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571256097] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF PM FILES (24) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1571256711] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (150) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571256744] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571256980] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13510) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571257110] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1571257306] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1571257338] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (78) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571257503] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571257503] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1571257712] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (3799) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571260317] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (130) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571260353] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2329) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571260443] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (130) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571260482] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2329) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571260569] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (130) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1571260608] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (12508) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571260616] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2327) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1571260779] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (12508) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571260926] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12508) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1571260955] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571261081] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571261179] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1571261207] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (71) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1571261322] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (3471) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571261436] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571261459] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (3471) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571261594] SERVICE ALERT: ieatlms3901;check pm collection statistical RBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (3471) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (3849)
[1571262143] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6409) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571262399] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (150) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571263069] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6559) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571263612] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3879) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571263821] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (30) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)
[1571264213] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.48%
[1571264364] SERVICE ALERT: ieatnetsimv7035-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.55%
[1571264487] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6409) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6588)