[1573776299] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (12)
[1573776422] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-BC-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1573776677] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (20) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573778084] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;OK;HARD;3;OK
[1573779317] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (27) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573779411] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (2046) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13974)
[1573779414] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573779548] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573779675] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573781470] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (11202) IS LOWER THAN 80% OF ENM  NODES (14024)
[1573781789] SERVICE FLAPPING ALERT: ieatlms3901;check synchronized DSC;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1573782427] SERVICE ALERT: ieatlms3901;check synchronized EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (38) IS LOWER THAN 80% OF ENM  NODES (50)
[1573782552] SERVICE ALERT: ieatlms3901;check synchronized EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (38) IS LOWER THAN 80% OF ENM  NODES (50)
[1573782677] SERVICE ALERT: ieatlms3901;check synchronized EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (38) IS LOWER THAN 80% OF ENM  NODES (50)
[1573783894] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (17) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573785993] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573786524] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1573786524] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (134) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573786638] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1573786638] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (12221) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (13974)
[1573786886] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;OK;HARD;3;OK- NUMBER OF PM FILES (4500) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573786967] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1573787716] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1573789449] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.03%
[1573789599] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (18) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573789600] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 79.28%
[1573791112] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (22) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573791648] SERVICE ALERT: ieatlms3901;check pm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1573791775] SERVICE ALERT: ieatlms3901;check pm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1573792623] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1573792623] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2103) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573794718] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573794859] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (21) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573794984] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (24)
[1573795205] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1537) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1573795329] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1549) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1573795454] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1558) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1573796591] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573797330] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (118) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573797465] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (11112) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13974)
[1573797467] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (916) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1573797707] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1573797707] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (27.4% change >= 20.0% threshold)
[1573797707] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (42.2% change >= 20.0% threshold)
[1573797707] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (27.4% change >= 20.0% threshold)
[1573797707] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (40.1% change >= 20.0% threshold)
[1573798589] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (18) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573798716] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (18) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573798846] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (18) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573799831] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1996) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573801082] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (157) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1573803284] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1573803405] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1573803525] SERVICE ALERT: ieatnetsimv9034-43;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1573804227] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573804414] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1573804602] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1573805079] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573805272] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573805465] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1573806745] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573806866] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573806876] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1573806987] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573806996] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1573807117] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1573807392] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573807512] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1573807553] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1573807633] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1573807674] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1573807794] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1573807811] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573807932] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573808053] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573808269] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573808493] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573808619] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573808687] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;SOFT;1;CRITICAL- san_alert_healthcheck FAILED
[1573808720] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1573808746] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573808807] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;SOFT;2;CRITICAL- san_alert_healthcheck FAILED
[1573808840] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1573808928] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;HARD;3;CRITICAL- san_alert_healthcheck FAILED
[1573808961] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1573809056] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1573809056] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1573809056] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (45.1% change >= 20.0% threshold)
[1573809056] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1573809056] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (43.2% change >= 20.0% threshold)
[1573809140] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1573809140] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1573809140] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (45.1% change >= 20.0% threshold)
[1573809140] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1573809140] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (43.2% change >= 20.0% threshold)
[1573809407] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1573809407] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1573809407] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (45.1% change >= 20.0% threshold)
[1573809407] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1573809407] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (43.2% change >= 20.0% threshold)
[1573809579] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1573809579] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1573809579] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (45.1% change >= 20.0% threshold)
[1573809579] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1573809579] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (43.2% change >= 20.0% threshold)
[1573809604] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1573809628] HOST ALERT: ieatrcxb7567;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.210.217.9)
[1573809630] HOST ALERT: ieatrcxb7654;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.210.217.97)
[1573809631] HOST ALERT: ieatrcxb7655;DOWN;SOFT;1;CRITICAL - Host Unreachable (10.210.217.98)
[1573809651] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (19) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573809693] HOST ALERT: ieatrcxb7567;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.210.217.9)
[1573809695] HOST ALERT: ieatrcxb7654;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.210.217.97)
[1573809696] HOST ALERT: ieatrcxb7655;DOWN;SOFT;2;CRITICAL - Host Unreachable (10.210.217.98)
[1573809725] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1573809757] HOST ALERT: ieatrcxb7567;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.210.217.9)
[1573809759] HOST ALERT: ieatrcxb7654;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.210.217.97)
[1573809760] HOST ALERT: ieatrcxb7655;DOWN;SOFT;3;CRITICAL - Host Unreachable (10.210.217.98)
[1573809822] HOST ALERT: ieatrcxb7567;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.210.217.9)
[1573809824] HOST ALERT: ieatrcxb7654;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.210.217.97)
[1573809825] HOST ALERT: ieatrcxb7655;DOWN;SOFT;4;CRITICAL - Host Unreachable (10.210.217.98)
[1573809846] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1573809858] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1573809886] HOST ALERT: ieatrcxb7567;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.210.217.9)
[1573809888] HOST ALERT: ieatrcxb7654;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.210.217.97)
[1573809889] HOST ALERT: ieatrcxb7655;DOWN;SOFT;5;CRITICAL - Host Unreachable (10.210.217.98)
[1573809951] HOST ALERT: ieatrcxb7567;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.210.217.9)
[1573809953] HOST ALERT: ieatrcxb7654;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.210.217.97)
[1573809954] HOST ALERT: ieatrcxb7655;DOWN;SOFT;6;CRITICAL - Host Unreachable (10.210.217.98)
[1573809979] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1573810015] HOST ALERT: ieatrcxb7567;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.210.217.9)
[1573810017] HOST ALERT: ieatrcxb7654;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.210.217.97)
[1573810018] HOST ALERT: ieatrcxb7655;DOWN;SOFT;7;CRITICAL - Host Unreachable (10.210.217.98)
[1573810080] HOST ALERT: ieatrcxb7567;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.210.217.9)
[1573810082] HOST ALERT: ieatrcxb7654;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.210.217.97)
[1573810083] HOST ALERT: ieatrcxb7655;DOWN;SOFT;8;CRITICAL - Host Unreachable (10.210.217.98)
[1573810100] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1573810144] HOST ALERT: ieatrcxb7567;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.210.217.9)
[1573810146] HOST ALERT: ieatrcxb7654;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.210.217.97)
[1573810147] HOST ALERT: ieatrcxb7655;DOWN;SOFT;9;CRITICAL - Host Unreachable (10.210.217.98)
[1573810209] HOST ALERT: ieatrcxb7567;DOWN;HARD;10;CRITICAL - Host Unreachable (10.210.217.9)
[1573810211] HOST ALERT: ieatrcxb7654;DOWN;HARD;10;CRITICAL - Host Unreachable (10.210.217.97)
[1573810212] HOST ALERT: ieatrcxb7655;DOWN;HARD;10;CRITICAL - Host Unreachable (10.210.217.98)
[1573810400] SERVICE ALERT: ieatlms7555;check rpm diff;WARNING;SOFT;1;WARNING - RPM ON SERVER ARE DIFFERENT FROM ISO (1.83.51)
[1573810522] SERVICE ALERT: ieatlms7555;check rpm diff;WARNING;SOFT;2;WARNING - RPM ON SERVER ARE DIFFERENT FROM ISO (1.83.51)
[1573810645] SERVICE ALERT: ieatlms7555;check rpm diff;WARNING;HARD;3;WARNING - RPM ON SERVER ARE DIFFERENT FROM ISO (1.83.51)
[1573811199] SERVICE ALERT: ieatnetsimv7102-03;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.71%
[1573811350] SERVICE ALERT: ieatnetsimv7102-03;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 2%
[1573811480] HOST ALERT: ieatnetsimv7035-06;DOWN;SOFT;1;PING NOK
[1573811480] HOST ALERT: ieatnetsimv7035-71;DOWN;SOFT;1;PING NOK
[1573811542] HOST ALERT: ieatnetsimv7035-71;UP;SOFT;1;PING OK
[1573811542] HOST ALERT: ieatnetsimv7035-06;UP;SOFT;1;PING OK
[1573811667] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1681) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1573812352] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;OK;HARD;3;OK- NUMBER OF PM FILES (5400) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573812448] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573812579] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573812592] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1573812634] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1573812738] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573813099] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573813259] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (16) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573813387] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (16) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573813516] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1573813523] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (16) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (24)
[1573813683] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1573813802] SERVICE ALERT: ieatlms5735-1;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1573814025] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1573814026] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2395) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573814151] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (2395) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573814161] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1573814247] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573814276] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (2395) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573814298] SERVICE ALERT: ieatlms7555;check unsync nodes;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1573814429] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573814610] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573814767] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2400) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573814893] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2400) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573815021] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2400) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573815053] SERVICE ALERT: ieatlms5735-1;check mdt hc;OK;HARD;3;OK- mdt_healthcheck SUCCESSFULLY PASSED
[1573815277] SERVICE ALERT: ieatnetsimv7035-61;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1573815312] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573815425] SERVICE ALERT: ieatnetsimv7035-61;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1573815494] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573815675] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1573815958] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573816087] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573816114] SERVICE ALERT: ieatlms7555;check unsync nodes;OK;HARD;3;OK
[1573816234] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573816471] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573818628] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573819850] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;OK;HARD;3;OK- NUMBER OF PM FILES (10800) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573821384] SERVICE ALERT: ieatlms3901;check cm supervision nodes RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (14024)
[1573821526] SERVICE ALERT: ieatlms3901;check cm supervision nodes RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (14022) IS HIGHER THAN 95% OF ENM NODES (14024)
[1573822937] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573823118] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573823299] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573823350] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (98)
[1573823473] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573823476] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes BSC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (98) IS HIGHER THAN 95% OF ENM NODES (98)
[1573823599] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573823732] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573823736] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573823913] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573824095] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573824191] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.44%
[1573824341] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 19.97%
[1573825411] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573825605] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573825798] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1573826527] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (43.8% change >= 20.0% threshold)
[1573826527] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (52.2% change >= 20.0% threshold)
[1573826742] SERVICE ALERT: ieatlms3901;check cm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (150) IS LOWER THAN 80% OF ENM NODES (300)
[1573826849] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573826871] SERVICE ALERT: ieatlms3901;check cm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1573826894] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (43.8% change >= 20.0% threshold)
[1573826894] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (52.2% change >= 20.0% threshold)
[1573826961] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1573826970] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573826980] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1573826993] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;SOFT;1;CRITICAL- san_alert_healthcheck FAILED
[1573827035] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1573827090] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573827100] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1573827113] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;SOFT;2;CRITICAL- san_alert_healthcheck FAILED
[1573827149] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (43.8% change >= 20.0% threshold)
[1573827149] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (52.2% change >= 20.0% threshold)
[1573827156] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1573827221] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1573827234] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;CRITICAL;HARD;3;CRITICAL- san_alert_healthcheck FAILED
[1573827277] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1573827343] SERVICE ALERT: ieatlms3901;check pm collection statistical SBG-IS;OK;HARD;3;OK- NUMBER OF PM FILES (4500) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573827500] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573827621] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1573827742] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1573827742] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1573827750] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (43.8% change >= 20.0% threshold)
[1573827750] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (57.2% change >= 20.0% threshold)
[1573827750] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1573827841] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (43.8% change >= 20.0% threshold)
[1573827841] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (57.2% change >= 20.0% threshold)
[1573827841] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1573827917] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573827960] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;STARTED; Service appears to have started flapping (38.9% change >= 20.0% threshold)
[1573827960] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical SBG-IS;STARTED; Service appears to have started flapping (57.2% change >= 20.0% threshold)
[1573827960] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1573828038] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573828084] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1573828159] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573828203] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1573828205] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1573828324] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1573828325] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1573828445] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1573828823] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573829004] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573829086] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs cluster hc;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1573829086] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573829185] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1573829471] SERVICE FLAPPING ALERT: ieatlms5735-1;check system service hc;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1573829471] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573830023] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs llt heartbeat hc;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1573830023] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;HARD;3;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1573830739] SERVICE FLAPPING ALERT: ieatlms3901;check node fs hc;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1573830739] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573830750] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1573830750] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573830851] SERVICE ALERT: ieatlms7435;check pm collection statistical SBG-IS;OK;HARD;3;OK- NUMBER OF PM FILES (6000) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (20)
[1573830896] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (67) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1573830898] SERVICE FLAPPING ALERT: ieatlms5735-1;check san alert healthcheck;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1573830898] SERVICE ALERT: ieatlms5735-1;check san alert healthcheck;OK;HARD;3;OK- san_alert_healthcheck SUCCESSFULLY PASSED
[1573830948] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1573830949] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6672;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH FM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1200)
[1573830950] SERVICE FLAPPING ALERT: ieatlms5735-1;check storagepool hc;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1573830950] SERVICE ALERT: ieatlms5735-1;check storagepool hc;OK;HARD;3;OK- storagepool_healthcheck SUCCESSFULLY PASSED
[1573831046] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573831076] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1573831130] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1573831227] SERVICE ALERT: ieatlms3901;check node fs;OK;SOFT;2;OK
[1573831311] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1573831409] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573831590] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573831771] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1573831885] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1573831994] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.78%
[1573831995] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573832015] SERVICE ALERT: ieatlms3901;check synchronized MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (104) IS LOWER THAN 80% OF ENM  NODES (150)
[1573832068] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;OK;SOFT;2;OK- vcs_llt_heartbeat_healthcheck SUCCESSFULLY PASSED
[1573832106] SERVICE FLAPPING ALERT: ieatlms5735-1;check nas hc;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1573832106] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1573832144] SERVICE ALERT: ieatnetsimv5116-65;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 48.87%
[1573832146] SERVICE ALERT: ieatlms3901;check synchronized MGW;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (131) IS FROM 80% TO 95% OF ENM  NODES (150)
[1573832146] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573832240] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1573832240] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2461) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573832272] SERVICE ALERT: ieatlms3901;check synchronized MGW;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (128) IS FROM 80% TO 95% OF ENM  NODES (150)
[1573832298] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1573832600] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573832782] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573832809] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573832815] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.34%
[1573832819] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (25) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (30)
[1573832928] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 287 processes with STATE = RSZDT
[1573832962] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1573832963] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1573832966] SERVICE ALERT: ieatnetsimv5116-38;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 50.91%
[1573832988] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 295 processes with STATE = RSZDT
[1573833033] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1573833049] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 329 processes with STATE = RSZDT
[1573833077] SERVICE ALERT: ieatlms3901;check synchronized RNC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (5) IS HIGHER THAN 95% OF ENM  NODES (6)
[1573833109] SERVICE ALERT: localhost;Total Processes;WARNING;HARD;4;PROCS WARNING: 310 processes with STATE = RSZDT
[1573833113] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.61%
[1573833159] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573833183] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.62%
[1573833269] SERVICE ALERT: ieatlms5735-1;check fm hb failures ERBS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833269] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833271] SERVICE ALERT: ieatlms5735-1;check fm hb failures MGW;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833273] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-BSP;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1573833278] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MSC-BC-IS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833278] SERVICE ALERT: ieatlms5735-1;check fm hb failures SBG-IS;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1573833285] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes SGSN-MME;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833297] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes BSC;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1573833309] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes SGSN-MME;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833322] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833334] SERVICE ALERT: ieatnetsimv5116-94;check netsim cpu;OK;HARD;3;OK- CPU IS USED 64.45%
[1573833352] SERVICE ALERT: ieatlms5735-1;check number nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833354] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573833356] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833360] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical MGW;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1573833360] SERVICE ALERT: ieatlms3901;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573833391] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833394] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes EPG;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833396] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes BSC;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833396] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MTAS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833407] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MGW;CRITICAL;SOFT;1;(Service check timed out after 800.52 seconds)
[1573833423] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes ESC;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1573833427] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833444] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833462] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MSC-BC-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (13) IS HIGHER THAN 95% OF ENM NODES (13)
[1573833472] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833474] SERVICE ALERT: ieatlms5735-1;check fm hb failures MGW;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MGW NODES !
[1573833481] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes DSC;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1573833482] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-DB-BSP;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833488] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes BSC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (98) IS HIGHER THAN 95% OF ENM NODES (98)
[1573833490] SERVICE ALERT: ieatlms3901;check pm collection statistical ESC;OK;HARD;3;OK- NUMBER OF PM FILES (965) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1000)
[1573833490] SERVICE ALERT: ieatlms5735-1;check fm hb failures ERBS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN ERBS NODES !
[1573833494] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical RadioNode;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1573833494] SERVICE ALERT: ieatlms3901;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13812) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13974)
[1573833497] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1573833501] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ERBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (2560) IS HIGHER THAN 95% OF ENM NODES (2560)
[1573833508] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (0)!
[1573833515] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1573833527] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes EPG;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1573833528] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2526) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1573833530] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1573833531] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1573833545] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1573833547] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573833557] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (1750) IS HIGHER THAN 95% OF ENM NODES (1750)
[1573833560] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1573833570] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1573833577] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-BC-IS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833578] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MTAS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833588] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes Router6274;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1573833608] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes DSC;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833622] SERVICE ALERT: ieatlms3901;check pm collection statistical MTAS;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573833623] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MGW;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1573833656] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes BSC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (98) IS HIGHER THAN 95% OF ENM NODES (98)
[1573833661] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ESC;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1573833666] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1573833676] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MTAS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833678] SERVICE ALERT: ieatlms5735-1;check fm hb failures SBG-IS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN SBG-IS NODES !
[1573833691] SERVICE ALERT: ieatlms5735-1;check synchronized RNC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (3) IS HIGHER THAN 95% OF ENM  NODES (4)
[1573833704] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-BC-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (13) IS HIGHER THAN 95% OF ENM NODES (13)
[1573833709] SERVICE ALERT: localhost;Total Processes;OK;HARD;4;PROCS OK: 110 processes with STATE = RSZDT
[1573833726] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1573833738] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6274;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573833742] SERVICE ALERT: ieatlms5735-1;check fm hb failures MSC-BC-BSP;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MSC-BC-BSP NODES !
[1573833746] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1573833747] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-DB-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (24) IS HIGHER THAN 95% OF ENM NODES (24)
[1573833749] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1573833788] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes ESC;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1000) IS HIGHER THAN 95% OF ENM NODES (1000)
[1573833792] SERVICE ALERT: ieatlms5735-1;check number nodes;OK;SOFT;2;OK- NUMBER OF NODES CONFIGURED IN ENM (22115) IS HIGHER THAN 95% OF NETSIM NODES (0)
[1573833800] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes MTAS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1573833861] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes Router6274;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1573833925] SERVICE ALERT: ieatlms3901;check pm supervision nodes DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (30)
[1573834083] SERVICE ALERT: ieatlms3901;check synchronized MGW;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1573834209] SERVICE ALERT: ieatlms3901;check pm supervision nodes DSC;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (30) IS HIGHER THAN 95% OF ENM NODES (30)
[1573834396] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes FRONTHAUL-6080;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (100)
[1573834431] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (1514)
[1573834528] SERVICE ALERT: ieatlms3901;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (80)
[1573834689] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (0)!
[1573834909] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 299 processes with STATE = RSZDT
[1573834935] SERVICE ALERT: ieatlms3901;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (28) IS HIGHER THAN 95% OF ENM  NODES (30)
[1573834942] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-BC-BSP;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573834945] SERVICE ALERT: ieatlms5735-1;check synchronized FRONTHAUL-6080;CRITICAL;SOFT;1;(Service check timed out after 800.53 seconds)
[1573834968] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MINI-LINK-6352;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573834969] SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 290 processes with STATE = RSZDT
[1573834970] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MINI-LINK-Indoor;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573834976] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes FRONTHAUL-6080;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (100) IS HIGHER THAN 95% OF ENM NODES (100)
[1573834980] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes RadioNode;CRITICAL;SOFT;1;(Service check timed out after 800.95 seconds)
[1573835028] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes Router6672;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573835029] SERVICE ALERT: localhost;Total Processes;OK;SOFT;3;PROCS OK: 230 processes with STATE = RSZDT
[1573835030] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes EPG;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573835034] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes SBG-IS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573835035] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes RBS;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1573835059] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1573835081] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (91) IS FROM 80% TO 95% OF ENM  NODES (98)
[1573835087] SERVICE ALERT: ieatlms5735-1;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.66 seconds)
[1573835102] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MINI-LINK-Indoor;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1750) IS HIGHER THAN 95% OF ENM NODES (1750)
[1573835132] SERVICE ALERT: ieatlms3901;check pm collection statistical MSC-DB-BSP;OK;HARD;3;OK- NUMBER OF PM FILES (24) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (24)
[1573835166] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes RBS;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (1514) IS HIGHER THAN 95% OF ENM NODES (1514)
[1573835205] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes MINI-LINK-6352;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (750) IS HIGHER THAN 95% OF ENM NODES (750)
[1573835207] SERVICE ALERT: ieatlms5735-1;check synchronized FRONTHAUL-6080;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (100)!
[1573835209] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes MSC-BC-BSP;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (12) IS HIGHER THAN 95% OF ENM NODES (12)
[1573835226] SERVICE ALERT: ieatlms5735-1;check fm hb failures MTAS;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN MTAS NODES !
[1573835236] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes RadioNode;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (13935) IS HIGHER THAN 95% OF ENM NODES (13935)
[1573835241] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes RBS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1514) IS HIGHER THAN 95% OF ENM NODES (1514)
[1573835252] SERVICE ALERT: ieatlms3901;check synchronized BSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (93) IS HIGHER THAN 95% OF ENM  NODES (96)
[1573835253] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes Router6672;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (1200) IS HIGHER THAN 95% OF ENM NODES (1200)
[1573835260] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1573835274] SERVICE ALERT: ieatlms5735-1;check fm supervision nodes EPG;OK;SOFT;2;OK- NUMBER OF NODES WITH FM SUPERVISION (50) IS HIGHER THAN 95% OF ENM NODES (50)
[1573835277] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (94) IS HIGHER THAN 95% OF ENM  NODES (98)
[1573835285] SERVICE ALERT: ieatlms5735-1;check fm hb failures cpp;OK;SOFT;2;OK- NO FM HB FAILURES ARE PRESENT IN RBS ERBS MTAS MGW NODES !
[1573835358] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1573835492] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13831) IS HIGHER THAN 95% OF ENM  NODES (14024)
[1573835649] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6672;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (1056) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (1200)
[1573835772] SERVICE ALERT: ieatlms3901;check pm collection statistical CISCO-ASR900;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1573835897] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes SBG-IS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (20)
[1573835928] SERVICE ALERT: ieatnetsimv5116-80;check netsim cpu;OK;HARD;3;OK- CPU IS USED 9.13%
[1573835971] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH PM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (80)
[1573836065] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1573836076] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes SBG-IS;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (20) IS HIGHER THAN 95% OF ENM NODES (20)
[1573836110] SERVICE ALERT: ieatlms5735-1;check pm supervision nodes SGSN-MME;OK;SOFT;2;OK- NUMBER OF NODES WITH PM SUPERVISION (80) IS HIGHER THAN 95% OF ENM NODES (80)
[1573836380] HOST ALERT: ieatnetsimv7035-114;DOWN;SOFT;1;PING NOK
[1573836425] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical DSC;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1573836425] SERVICE ALERT: ieatlms3901;check pm collection statistical DSC;OK;HARD;3;OK- NUMBER OF PM FILES (30) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (30)
[1573836442] HOST ALERT: ieatnetsimv7035-114;UP;SOFT;1;PING OK
[1573836513] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (12)
[1573836594] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.03%
[1573836650] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573836686] SERVICE ALERT: ieatlms3901;check synchronized RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (2) IS LOWER THAN 80% OF ENM  NODES (6)
[1573836710] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (12)
[1573836744] SERVICE ALERT: ieatnetsimv5116-95;check netsim cpu;OK;HARD;3;OK- CPU IS USED 75.47%
[1573836744] SERVICE ALERT: ieatnetsimv5116-34;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.07%
[1573836752] SERVICE ALERT: ieatlms3901;check synchronized MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (13) IS LOWER THAN 80% OF ENM  NODES (30)
[1573836772] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573836798] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573836807] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1573836812] SERVICE ALERT: ieatlms3901;check synchronized RNC;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (4) IS HIGHER THAN 95% OF ENM  NODES (6)
[1573836833] SERVICE ALERT: ieatlms5735-1;check synchronized MSC-BC-BSP;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM  NODES (12)
[1573836887] SERVICE ALERT: ieatlms3901;check synchronized MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (22) IS LOWER THAN 80% OF ENM  NODES (30)
[1573836925] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573837011] SERVICE ALERT: ieatlms3901;check synchronized MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (18) IS LOWER THAN 80% OF ENM  NODES (30)
[1573837138] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2416) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573837266] SERVICE ALERT: ieatlms3901;check synchronized ERBS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (2413) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573837308] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (2802) IS LOWER THAN 80% OF ENM  NODES (14024)
[1573837402] SERVICE ALERT: ieatlms3901;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2556) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1573837443] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (2802) IS LOWER THAN 80% OF ENM  NODES (14024)
[1573837586] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13878) IS HIGHER THAN 95% OF ENM  NODES (14024)
[1573837684] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (1998) IS LOWER THAN 80% OF ENM  NODES (2560)
[1573838092] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (22) IS LOWER THAN 80% OF ENM  NODES (30)
[1573838117] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (611) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573838149] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (8228) IS LOWER THAN 80% OF ENM  NODES (13935)
[1573838151] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (19) IS LOWER THAN 80% OF ENM  NODES (30)
[1573838172] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (44) IS LOWER THAN 80% OF ENM  NODES (80)
[1573838217] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (18) IS LOWER THAN 80% OF ENM  NODES (30)
[1573838280] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (8179) IS LOWER THAN 80% OF ENM  NODES (13935)
[1573838292] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (16) IS LOWER THAN 80% OF ENM  NODES (30)
[1573838301] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (46) IS LOWER THAN 80% OF ENM  NODES (80)
[1573838341] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (15) IS LOWER THAN 80% OF ENM  NODES (30)
[1573838421] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (9118) IS LOWER THAN 80% OF ENM  NODES (13935)
[1573838425] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573838438] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (43) IS LOWER THAN 80% OF ENM  NODES (80)
[1573838451] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (13) IS LOWER THAN 80% OF ENM  NODES (30)
[1573838561] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1229) IS FROM 80% TO 95% OF ENM  NODES (1514)
[1573838685] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1235) IS FROM 80% TO 95% OF ENM  NODES (1514)
[1573838810] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1253) IS FROM 80% TO 95% OF ENM  NODES (1514)
[1573838964] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (50) IS LOWER THAN 80% OF ENM  NODES (80)
[1573839001] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (35) IS LOWER THAN 80% OF ENM  NODES (98)
[1573839020] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (27) IS LOWER THAN 80% OF ENM  NODES (50)
[1573839032] SERVICE FLAPPING ALERT: ieatlms3901;check node fs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1573839032] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1573839089] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (50) IS LOWER THAN 80% OF ENM  NODES (80)
[1573839125] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (35) IS LOWER THAN 80% OF ENM  NODES (98)
[1573839144] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (20) IS LOWER THAN 80% OF ENM  NODES (50)
[1573839215] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1573839249] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (3) IS LOWER THAN 80% OF ENM  NODES (98)
[1573839268] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (14) IS LOWER THAN 80% OF ENM  NODES (50)
[1573839408] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (8118) IS LOWER THAN 80% OF ENM  NODES (14024)
[1573839442] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;OK;HARD;3;OK- NUMBER OF PM FILES (40) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (50)
[1573839491] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (2536) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1573839547] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (9165) IS LOWER THAN 80% OF ENM  NODES (14024)
[1573839669] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1573839700] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (9165) IS LOWER THAN 80% OF ENM  NODES (14024)
[1573839793] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1573839930] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH HB FAILURE (4) IS BETWEEN 5% AND 20% OF ENM NODES (50)
[1573840235] SERVICE ALERT: ieatlms5735-1;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13765) IS HIGHER THAN 95% OF ENM  NODES (13935)
[1573840246] SERVICE ALERT: ieatlms5735-1;check synchronized SGSN-MME;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1573840434] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1573840472] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1573840531] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1573840531] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2560) IS EQUAL TO NUMBER OF NODES WITH PM SUPERVISION ENABLED (2560)
[1573840614] SERVICE ALERT: ieatlms5735-1;check synchronized RBS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1514)!
[1573840616] SERVICE ALERT: ieatlms3901;check synchronized MTAS;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (24) IS HIGHER THAN 95% OF ENM  NODES (30)
[1573840616] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1573840661] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1573840793] HOST ALERT: ieatrcxb6218-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573840809] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1573840852] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573840866] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1573840883] HOST ALERT: ieatrcxb6218-1;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573840973] HOST ALERT: ieatrcxb6218-1;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1573841021] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (52) IS LOWER THAN 80% OF ENM  NODES (80)
[1573841037] HOST ALERT: ieatrcxb6218-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1573841052] SERVICE ALERT: ieatlms5735-1;check synchronized BSC;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (95) IS HIGHER THAN 95% OF ENM  NODES (98)
[1573841061] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573841073] SERVICE ALERT: ieatlms5735-1;check synchronized EPG;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (48) IS HIGHER THAN 95% OF ENM  NODES (50)
[1573841148] SERVICE ALERT: ieatlms3901;check synchronized SGSN-MME;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (80)!
[1573841255] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573841523] SERVICE ALERT: ieatlms3901;check synchronized RadioNode;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (13646) IS HIGHER THAN 95% OF ENM  NODES (14024)
[1573841724] SERVICE ALERT: ieatlms3901;check pm collection statistical MINI-LINK-6352;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (581) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (750)
[1573841944] SERVICE ALERT: ieatlms5735-1;check synchronized DSC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1573842076] SERVICE ALERT: ieatlms5735-1;check synchronized MTAS;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1573842083] SERVICE ALERT: ieatlms3901;check synchronized EPG;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (44) IS FROM 80% TO 95% OF ENM  NODES (50)
[1573842572] HOST ALERT: ieatrcxb6298-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573842662] HOST ALERT: ieatrcxb6298-1;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573842752] HOST ALERT: ieatrcxb6298-1;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1573842816] HOST ALERT: ieatrcxb6298-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1573842990] SERVICE ALERT: ieatlms3901;check synchronized DSC;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (30)!
[1573843049] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (39) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1573843129] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573843181] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1573843248] HOST ALERT: ieatrcxb6294-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573843313] SERVICE ALERT: ieatlms3901;check pm collection statistical Router6274;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (39) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (50)
[1573843338] HOST ALERT: ieatrcxb6294-1;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573843339] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573843402] HOST ALERT: ieatrcxb6294-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1573843533] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1573843888] SERVICE ALERT: ieatlms3901;check synchronized EPG;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (50)!
[1573844394] HOST ALERT: ieatrcxb6219-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573844458] HOST ALERT: ieatrcxb6219-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1573845562] HOST ALERT: ieatrcxb6215-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573845626] HOST ALERT: ieatrcxb6215-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1573845874] HOST ALERT: ieatrcxb6299-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573845964] HOST ALERT: ieatrcxb6299-1;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573846054] HOST ALERT: ieatrcxb6299-1;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1573846118] HOST ALERT: ieatrcxb6299-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1573846192] SERVICE ALERT: ieatlms3901;check synchronized CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF  NODES SYNCHRONIZED (150) IS LOWER THAN 80% OF ENM  NODES (300)
[1573846317] SERVICE ALERT: ieatlms3901;check synchronized CISCO-ASR900;OK;SOFT;2;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1573847207] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573847351] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.12%
[1573847502] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 54.25%
[1573847672] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1643) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1573847797] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1537) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1573847883] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes CISCO-ASR900;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH CM SUPERVISION (0) IS LOWER THAN 80% OF ENM NODES (150)
[1573847921] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1560) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1573847972] HOST ALERT: ieatrcxb6295-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573848008] SERVICE ALERT: ieatlms5735-1;check cm supervision nodes CISCO-ASR900;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (150) IS HIGHER THAN 95% OF ENM NODES (150)
[1573848062] HOST ALERT: ieatrcxb6295-1;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573848126] HOST ALERT: ieatrcxb6295-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1573849726] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- NUMBER OF  NODES SYNCHRONIZED (1749) IS HIGHER THAN 95% OF ENM  NODES (1750)
[1573849880] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;CRITICAL;SOFT;1;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1) IS HIGHER THAN 20% OF ENM NODES (4)
[1573850004] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;CRITICAL;SOFT;2;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1) IS HIGHER THAN 20% OF ENM NODES (4)
[1573850131] SERVICE ALERT: ieatlms5735-1;check fm hb failures RNC;CRITICAL;HARD;3;CRITICAL- NUMBER OF NODES WITH HB FAILURE (1) IS HIGHER THAN 20% OF ENM NODES (4)
[1573850372] HOST ALERT: ieatrcxb6216-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573850462] HOST ALERT: ieatrcxb6216-1;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573850526] HOST ALERT: ieatrcxb6216-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.19 ms
[1573851671] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573852096] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (2394) IS FROM 80% TO 95% OF ENM  NODES (2560)
[1573852223] SERVICE ALERT: ieatlms5735-1;check synchronized ERBS;OK;SOFT;2;OK- NUMBER OF  NODES SYNCHRONIZED (2555) IS HIGHER THAN 95% OF ENM  NODES (2560)
[1573852446] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6672;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (10) IS LOWER THAN 5% OF ENM NODES (1200)
[1573852489] SERVICE ALERT: ieatlms5735-1;check fm hb failures RadioNode;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (139) IS LOWER THAN 5% OF ENM NODES (13935)
[1573852492] SERVICE ALERT: ieatlms5735-1;check fm hb failures FRONTHAUL-6080;OK;HARD;3;OK- NUMBER OF NODES WITH HB FAILURE (1) IS LOWER THAN 5% OF ENM NODES (100)
[1573852574] HOST ALERT: ieatrcxb6296-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573852664] HOST ALERT: ieatrcxb6296-1;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1573852695] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573852728] HOST ALERT: ieatrcxb6296-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1573852889] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573854334] SERVICE ALERT: ieatlms5735-1;check fm hb failures Router6274;OK;HARD;3;OK- NO FM HB FAILURES ARE PRESENT IN Router6274 NODES !
[1573854872] HOST ALERT: ieatrcxb6220-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573854890] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573854962] HOST ALERT: ieatrcxb6220-1;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573855026] HOST ALERT: ieatrcxb6220-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1573855085] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573855280] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573855333] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1573855515] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573856975] HOST ALERT: ieatrcxb6301-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573857065] HOST ALERT: ieatrcxb6301-1;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1573857129] HOST ALERT: ieatrcxb6301-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.83 ms
[1573857153] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573857346] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573857496] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1573857678] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;SOFT;2;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573857696] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (130) IS FROM 80% TO 95% OF ENM  NODES (150)
[1573857820] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (130) IS FROM 80% TO 95% OF ENM  NODES (150)
[1573857944] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (130) IS FROM 80% TO 95% OF ENM  NODES (150)
[1573858733] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;1;WARNING- NUMBER OF  NODES SYNCHRONIZED (1589) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1573858861] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;SOFT;2;WARNING- NUMBER OF  NODES SYNCHRONIZED (1612) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1573858987] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;WARNING;HARD;3;WARNING- NUMBER OF  NODES SYNCHRONIZED (1634) IS FROM 80% TO 95% OF ENM  NODES (1750)
[1573859354] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573859365] HOST ALERT: ieatrcxb6204;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573859455] HOST ALERT: ieatrcxb6204;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573859519] HOST ALERT: ieatrcxb6204;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1573859547] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573859660] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1573859741] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573859748] SERVICE ALERT: ieatlms5735-1;check synchronized CISCO-ASR900;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (150)!
[1573859841] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1573860024] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1573860086] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1573860267] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1573860448] SERVICE ALERT: ieatlms3901;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1573860802] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1573860802] SERVICE ALERT: ieatlms5735-1;check synchronized MINI-LINK-Indoor;OK;HARD;3;OK- ALL  NODES ARE SYNCHRONIZED (1750)!
[1573861496] HOST ALERT: ieatrcxb6268-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1573861586] HOST ALERT: ieatrcxb6268-1;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1573861637] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1573861676] HOST ALERT: ieatrcxb6268-1;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1573861741] HOST ALERT: ieatrcxb6268-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1573861832] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1573861886] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1573862025] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1573862068] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1573862249] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED