[1569107937] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;OK;HARD;3;OK
[1569115138] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1569115258] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1569115379] SERVICE ALERT: ieatnetsimv5116-96;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1569123172] HOST ALERT: ieatnetsimv5116-54;DOWN;SOFT;1;PING NOK
[1569123234] HOST ALERT: ieatnetsimv5116-54;UP;SOFT;1;PING OK
[1569132413] SERVICE ALERT: ieatnetsimv7035-108;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1569132414] SERVICE ALERT: ieatnetsimv7035-107;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1569132540] SERVICE ALERT: ieatnetsimv7035-108;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1569132542] SERVICE ALERT: ieatnetsimv7035-107;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1569141471] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (1694) IS LOWER THAN 80% OF ENM CPP NODES (4254)
[1569143329] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3964) IS FROM 80% TO 95% OF ENM CPP NODES (4254)
[1569145279] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (1694) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)
[1569149257] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES W/O FM HB FAILURES (3795) IS FROM 80% TO 95% OF CPP NODES WITH FM SUPERVISION ENABLED (4254)