[1566255690] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1566257563] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1566258452] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;OK;HARD;3;OK- CPU IS USED 70.47%
[1566258494] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.48%
[1566258646] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.2%
[1566258796] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.61%
[1566258947] SERVICE ALERT: ieatnetsimv5116-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 30.94%
[1566258947] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.81%
[1566259098] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.72%
[1566259249] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.6%
[1566259401] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.14%
[1566259551] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.95%
[1566262083] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.8%
[1566262234] SERVICE ALERT: ieatnetsimv5116-55;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 73.03%
[1566263182] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.88%
[1566263458] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.23%
[1566263609] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.38%
[1566263759] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.68%
[1566263910] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.42%
[1566264066] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.28%
[1566264216] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.64%
[1566266813] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.48%
[1566266963] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 90.12%
[1566267114] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim cpu;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1566267114] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.5%
[1566267347] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;SOFT;1;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2356) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566267494] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;SOFT;2;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2406) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1566269962] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.53%
[1566270113] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.54%
[1566270274] SERVICE ALERT: ieatnetsimv5116-71;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566270745] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.62%
[1566276612] SERVICE ALERT: ieatnetsimv7035-108;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.04%
[1566276763] SERVICE ALERT: ieatnetsimv7035-108;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 97.66%
[1566276914] SERVICE ALERT: ieatnetsimv7035-108;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.68%
[1566278461] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.18%
[1566278612] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.68%
[1566280385] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.08%
[1566280535] SERVICE ALERT: ieatnetsimv5116-66;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 19.05%
[1566280544] SERVICE ALERT: ieatnetsimv7035-108;check netsim cpu;OK;HARD;3;OK- CPU IS USED 60.14%
[1566281437] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1566281628] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1566281813] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1566284123] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.19%
[1566284147] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.99%
[1566284274] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.81%
[1566284298] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.91%
[1566284424] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.43%
[1566284449] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.56%
[1566284575] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.58%
[1566284599] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.68%
[1566284636] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1566284726] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.07%
[1566284755] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.53%
[1566284830] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1566284906] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.62%
[1566285024] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1566285176] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.83%
[1566285326] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 96.63%
[1566285477] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 94.6%
[1566285966] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.02%
[1566286117] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.24%
[1566286267] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.02%
[1566286418] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.79%
[1566286898] SERVICE FLAPPING ALERT: ieatlms5735-1;check system service hc;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1566286898] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1566288357] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.6%
[1566291039] SERVICE ALERT: ieatnetsimv7035-79;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566291169] SERVICE ALERT: ieatnetsimv7035-79;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566291249] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566291298] SERVICE ALERT: ieatnetsimv7035-79;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566291320] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.64%
[1566291392] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566291471] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.5%
[1566291539] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566291621] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.14%
[1566291772] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.59%
[1566291852] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566291877] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566291922] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 92.42%
[1566291995] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566292069] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566292152] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566292208] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566292284] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566292372] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1566292418] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566292548] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566292566] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1566292614] SERVICE ALERT: ieatnetsimv7035-70;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566292727] SERVICE ALERT: ieatnetsimv7035-53;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566292749] SERVICE ALERT: ieatnetsimv7035-70;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566292760] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1566292856] SERVICE ALERT: ieatnetsimv7035-53;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1566292879] SERVICE ALERT: ieatnetsimv7035-70;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566293112] SERVICE ALERT: ieatnetsimv7035-79;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566293181] SERVICE ALERT: ieatnetsimv7035-68;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566293311] SERVICE ALERT: ieatnetsimv7035-68;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566293330] SERVICE ALERT: ieatnetsimv7035-63;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566293334] SERVICE ALERT: ieatnetsimv7035-64;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566293370] SERVICE ALERT: ieatnetsimv7035-65;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566293445] SERVICE ALERT: ieatnetsimv7035-68;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566293460] SERVICE ALERT: ieatnetsimv7035-63;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1566293465] SERVICE ALERT: ieatnetsimv7035-64;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1566293501] SERVICE ALERT: ieatnetsimv7035-65;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1566293601] SERVICE ALERT: ieatnetsimv7035-64;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1566293635] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3902) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1566293777] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES SYNCHRONIZED (4081) IS HIGHER THAN 95% OF ENM CPP NODES (4255)
[1566294364] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566294633] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1566294694] SERVICE ALERT: ieatnetsimv7035-70;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566294734] SERVICE ALERT: ieatnetsimv7035-68;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566294884] SERVICE ALERT: ieatnetsimv7035-68;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 13.9%
[1566295145] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1566295418] SERVICE ALERT: ieatnetsimv7035-64;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566295553] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.15%
[1566295759] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2534) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2561)
[1566295830] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (14033) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (14035)
[1566296308] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.11%
[1566299456] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.52%
[1566299607] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.6%
[1566299758] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.08%
[1566299908] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.58%
[1566300362] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.53%
[1566300513] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.02%
[1566300664] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 97.21%
[1566300814] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.02%
[1566301118] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3917) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1566301253] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4003) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1566301387] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (4003) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1566302785] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.4%
[1566302936] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.52%
[1566305496] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.18%
[1566305499] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.68%
[1566305501] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.86%
[1566305507] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566305519] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566305522] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566305523] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.6%
[1566305527] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566305529] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.88%
[1566305540] SERVICE ALERT: ieatnetsimv7035-70;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566305542] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566305547] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566305562] SERVICE ALERT: ieatnetsimv7035-79;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1566305597] SERVICE ALERT: ieatnetsimv7035-70;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1566305646] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566305649] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 62.3%
[1566305652] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.95%
[1566305658] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566305670] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566305673] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566305674] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.13%
[1566305678] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.97%
[1566305679] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.94%
[1566305691] SERVICE ALERT: ieatnetsimv7035-70;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566305693] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.98%
[1566305698] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566305713] SERVICE ALERT: ieatnetsimv7035-79;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1566305766] SERVICE ALERT: ieatnetsimv7035-70;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1566305797] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566305802] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;OK;HARD;3;OK- CPU IS USED 77.24%
[1566305809] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566305821] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.71%
[1566305824] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566305824] SERVICE ALERT: ieatnetsimv7035-64;check netsim cpu;OK;HARD;3;OK- CPU IS USED 55.54%
[1566305828] SERVICE ALERT: ieatnetsimv7035-65;check netsim cpu;OK;HARD;3;OK- CPU IS USED 60.36%
[1566305830] SERVICE ALERT: ieatnetsimv7035-66;check netsim cpu;OK;HARD;3;OK- CPU IS USED 55.46%
[1566305842] SERVICE ALERT: ieatnetsimv7035-70;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1566305843] SERVICE ALERT: ieatnetsimv7035-71;check netsim cpu;OK;HARD;3;OK- CPU IS USED 55.69%
[1566305849] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.97%
[1566305864] SERVICE ALERT: ieatnetsimv7035-79;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.93%
[1566306315] SERVICE FLAPPING ALERT: ieatnetsimv5116-71;check netsim cpu;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1566307145] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.97%
[1566307295] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.36%
[1566309026] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1566309157] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1566309288] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;CRITICAL;HARD;3;CRITICAL- NUMBER OF SNMP NODES SYNCHRONIZED (0) IS LOWER THAN 80% OF ENM SNMP NODES (2500)
[1566309348] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566309385] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1566309428] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4%
[1566309440] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.69%
[1566309451] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;OK;HARD;3;OK- CPU IS USED 4.35%
[1566309454] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.46%
[1566309472] SERVICE ALERT: ieatnetsimv7035-70;check netsim cpu;OK;HARD;3;OK- CPU IS USED 6.49%
[1566309479] SERVICE ALERT: ieatnetsimv7035-73;check netsim cpu;OK;HARD;3;OK- CPU IS USED 2.79%
[1566309494] SERVICE ALERT: ieatnetsimv7035-79;check netsim cpu;OK;HARD;3;OK- CPU IS USED 3.72%
[1566309529] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566309566] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1566309710] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566309747] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1566311765] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.29%
[1566311915] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.86%
[1566312902] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;WARNING;HARD;3;WARNING- NUMBER OF SNMP NODES SYNCHRONIZED (2127) IS FROM 80% TO 95% OF ENM SNMP NODES (2500)
[1566312984] SERVICE ALERT: ieatnetsimv5116-39;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566313104] SERVICE ALERT: ieatnetsimv5116-39;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1566313225] SERVICE ALERT: ieatnetsimv5116-39;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1566314459] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.29%
[1566314609] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.54%
[1566314714] SERVICE ALERT: ieatlms5735-1;check synchronized snmp;OK;HARD;3;OK- NUMBER OF SNMP NODES SYNCHRONIZED (2482) IS HIGHER THAN 95% OF ENM SNMP NODES (2500)
[1566316067] SERVICE ALERT: ieatnetsimv5116-61;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.63%
[1566316088] SERVICE ALERT: ieatnetsimv5116-69;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.52%
[1566316218] SERVICE ALERT: ieatnetsimv5116-61;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 28.18%
[1566316239] SERVICE ALERT: ieatnetsimv5116-69;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 18.03%
[1566318360] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.02%
[1566318510] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.77%
[1566318661] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.89%
[1566318811] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.78%
[1566321846] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.43%
[1566321996] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.34%
[1566322147] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.82%
[1566324580] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.55%
[1566324734] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.78%
[1566324887] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.73%
[1566325037] SERVICE ALERT: ieatnetsimv5116-70;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 73.13%
[1566325778] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.02%
[1566325929] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.42%
[1566326079] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 93.18%
[1566327988] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.14%
[1566328103] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;1;(Service check timed out after 801.01 seconds)
[1566328138] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.93%
[1566328907] SERVICE FLAPPING ALERT: ieatlms5735-1;check system service hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1566329024] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;CRITICAL;SOFT;2;(Service check timed out after 801.01 seconds)
[1566329061] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566329197] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1566329242] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566329379] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1566329424] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1566329561] SERVICE ALERT: ieatlms3901;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1566329769] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.27%
[1566329812] SERVICE ALERT: ieatnetsimv5116-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1566329919] SERVICE ALERT: ieatnetsimv7035-14;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 87.41%
[1566330510] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.18%
[1566330661] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 5.36%
[1566331887] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.87%
[1566332037] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.88%
[1566332188] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.69%
[1566332338] SERVICE ALERT: ieatnetsimv7035-11;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.6%
[1566333362] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.79%
[1566333512] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 88.37%
[1566333663] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.04%
[1566333814] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 92.84%
[1566333965] SERVICE ALERT: ieatnetsimv7035-08;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 96.93%
[1566339089] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.64%
[1566339239] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.59%
[1566339390] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93%
[1566339540] SERVICE ALERT: ieatnetsimv7035-12;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.62%
[1566341614] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.54%
[1566341765] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 11.88%
[1566341796] SERVICE ALERT: ieatnetsimv5116-44;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1566341917] SERVICE ALERT: ieatnetsimv5116-44;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%