[1570662817] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570663011] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1570663109] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570663204] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1570663290] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570663472] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1570663472] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570664354] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570664535] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570664716] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1570664716] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1570665082] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570667217] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (110) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570667342] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (111) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570667468] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1570667468] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (73) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570667526] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1570667650] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (1278) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570667782] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (1278) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570667827] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1570667874] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12776) IS FROM 80% TO 95% OF ENM COMECIM NODES (14101)
[1570667909] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1570667909] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570668041] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12760) IS FROM 80% TO 95% OF ENM COMECIM NODES (14101)
[1570668156] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1570668199] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;HARD;3;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (12980) IS FROM 80% TO 95% OF ENM COMECIM NODES (14101)
[1570668452] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570668545] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570668578] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570668678] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1570668706] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1570668706] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (45) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570671831] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13487) IS HIGHER THAN 95% OF ENM COMECIM NODES (14101)
[1570675029] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570677099] SERVICE ALERT: ieatlms3901;check hw resources hc;OK;HARD;3;OK- hw_resources_healthcheck SUCCESSFULLY PASSED
[1570678090] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570678271] SERVICE ALERT: ieatlms3901;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570678453] SERVICE ALERT: ieatlms3901;check node fs;OK;HARD;3;OK
[1570682944] SERVICE FLAPPING ALERT: ieatlms3901;check system service hc;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570684763] HOST ALERT: ieatnetsimv7035-20;DOWN;SOFT;1;PING NOK
[1570684825] HOST ALERT: ieatnetsimv7035-20;UP;SOFT;1;PING OK
[1570687355] SERVICE ALERT: ieatnetsimv5116-38;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570691465] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570691650] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570691836] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570693097] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1570693412] SERVICE FLAPPING ALERT: ieatlms3901;check vcs cluster hc;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570693701] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570694503] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570694632] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570694761] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570695377] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs hc;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1570695861] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2400) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570695885] SERVICE ALERT: ieatwlvm5116;check workload profiles;CRITICAL;HARD;3;CRITICAL- THERE ARE WORKLOAD PROFILES WHICH ARE IN DEAD STATE
[1570695933] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570695990] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2400) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570696118] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2400) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570697366] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570697552] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570697738] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570697932] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570697939] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570697944] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570697948] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570698087] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 21.89%
[1570698094] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 10.88%
[1570698101] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.58%
[1570698103] SERVICE ALERT: ieatnetsimv7035-86;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 51.58%
[1570698244] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.86%
[1570698375] SERVICE ALERT: ieatnetsimv5116-42;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570698395] SERVICE ALERT: ieatnetsimv7035-82;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 9.01%
[1570698484] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (0) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570698895] SERVICE ALERT: ieatwlvm5116;check workload profiles;WARNING;HARD;3;WARNING- THERE ARE WORKLOAD PROFILES WHICH ARE IN ERROR STATE
[1570699202] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570699603] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1570699603] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570699746] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1570699746] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2539) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570700607] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570700738] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570700789] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570700869] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570700920] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570701052] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570701838] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.84%
[1570701988] SERVICE ALERT: ieatnetsimv7035-80;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 74.71%
[1570702179] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.95%
[1570702330] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 42.96%
[1570702539] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (308) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1570702577] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570702596] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.98%
[1570702733] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1570702747] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 89.66%
[1570702802] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.82%
[1570702873] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.78%
[1570702884] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 63.13%
[1570702898] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570702928] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.99%
[1570702931] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570702936] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570702937] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570702954] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570703025] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570703049] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.03%
[1570703079] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.16%
[1570703082] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.96%
[1570703087] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 72.92%
[1570703088] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.23%
[1570703105] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;OK;HARD;3;OK- CPU IS USED 88.34%
[1570703132] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1570703176] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;OK;HARD;3;OK- CPU IS USED 62%
[1570703185] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.93%
[1570703230] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 90.17%
[1570703233] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 95.21%
[1570703237] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.77%
[1570703238] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.77%
[1570703336] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.38%
[1570703361] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2178) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570703388] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570703477] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570703487] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570703492] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;2;WARNING- NUMBER OF PM FILES (2270) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570703540] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570703622] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2270) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570703880] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570704031] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.42%
[1570704182] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;OK;HARD;3;OK- CPU IS USED 65%
[1570704266] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized comecim;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1570704394] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1570704674] SERVICE ALERT: ieatnetsimv5116-40;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570704717] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570706145] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (608) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1570706271] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (608) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1570706397] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (608) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1570706636] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570706736] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570706787] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 66.48%
[1570706886] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 95.8%
[1570707037] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 93.18%
[1570707230] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2480) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570707931] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1570707931] SERVICE FLAPPING ALERT: ieatlms5735-1;check nas hc;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1570707931] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1570708123] SERVICE FLAPPING ALERT: ieatlms3901;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1570708123] SERVICE FLAPPING ALERT: ieatlms5735-1;check nas hc;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1570708123] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1570709296] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570709484] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570709672] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570709895] SERVICE ALERT: ieatnetsimv5116-39;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570710372] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570710408] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570710460] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;OK;HARD;3;OK- CPU IS USED 63.57%
[1570710524] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570710539] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570710562] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570710676] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570710692] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570710716] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570710845] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570711414] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570711566] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570711718] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570712652] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.89%
[1570712788] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6048) IS FROM 80% TO 95% OF ENM CPP NODES (6589)
[1570712802] SERVICE ALERT: ieatnetsimv7035-84;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 45.24%
[1570713282] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570713287] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.38%
[1570713339] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570713433] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 98.07%
[1570713438] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.33%
[1570713583] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 81.79%
[1570713589] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570714000] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570714022] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570714026] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570714091] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570714152] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570714175] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570714177] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570714244] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570714304] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570714326] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570714329] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570714396] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570714640] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (6564) IS HIGHER THAN 95% OF ENM CPP NODES (6589)
[1570715349] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;OK;HARD;3;OK- CPU IS USED 47.84%
[1570717219] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 45.67%
[1570717632] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;(Service check timed out after 800.86 seconds)
[1570717664] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;OK;HARD;3;OK- CPU IS USED 62.44%
[1570717668] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570717669] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;OK;HARD;3;OK- CPU IS USED 43.89%
[1570717808] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570717853] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570717868] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;OK;HARD;3;OK- CPU IS USED 53.48%
[1570717907] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.69%
[1570717918] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 47.71%
[1570717934] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.66%
[1570717957] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.05%
[1570717979] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570718015] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1570718026] SERVICE FLAPPING ALERT: ieatnetsimv7035-61;check netsim cpu;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1570718026] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.6%
[1570718039] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570718049] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;WARNING;SOFT;1;WARNING- NUMBER OF PM FILES (2276) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570718077] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.83%
[1570718118] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570718150] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;SOFT;2;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13983) IS HIGHER THAN 95% OF ENM COMECIM NODES (38)
[1570718190] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570718219] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1570718243] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical MGW;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1570718243] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570718319] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570718447] SERVICE ALERT: ieatlms3901;check pm collection statistical ERBS;OK;SOFT;2;OK- NUMBER OF PM FILES (2560) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570718615] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1570718801] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1570718809] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570718986] SERVICE ALERT: ieatlms3901;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1570718999] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570719188] SERVICE ALERT: ieatlms3901;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570719210] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570719340] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;2;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570719466] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;HARD;3;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570719905] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570720097] SERVICE ALERT: ieatlms3901;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570720104] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;1;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (6914) IS LOWER THAN 80% OF ENM COMECIM NODES (14097)
[1570720253] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;SOFT;2;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (6371) IS LOWER THAN 80% OF ENM COMECIM NODES (14097)
[1570720401] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;CRITICAL;HARD;3;CRITICAL- NUMBER OF COMECIM NODES SYNCHRONIZED (7354) IS LOWER THAN 80% OF ENM COMECIM NODES (14097)
[1570720860] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.74%
[1570721010] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.99%
[1570721161] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.92%
[1570721295] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570721300] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570721447] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570721452] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570721498] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.57%
[1570721538] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570721550] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570721566] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570721588] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.97%
[1570721598] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570721603] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570721650] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.15%
[1570721658] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570721689] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570721701] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570721707] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.94%
[1570721717] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570721739] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570721770] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570721800] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;OK;HARD;3;OK- CPU IS USED 68.59%
[1570721809] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.91%
[1570721840] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 92.83%
[1570721849] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570721852] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570721859] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570721869] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570721891] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570721955] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570721960] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 97.7%
[1570722010] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.23%
[1570722229] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13729) IS HIGHER THAN 95% OF ENM COMECIM NODES (14097)
[1570722261] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570722580] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570722649] SERVICE ALERT: ieatlms3901;check vcs cluster hc;OK;HARD;3;OK- vcs_cluster_healthcheck SUCCESSFULLY PASSED
[1570722731] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570722856] SERVICE FLAPPING ALERT: ieatlms3901;check vcs service group hc;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1570722856] SERVICE ALERT: ieatlms3901;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570722881] SERVICE ALERT: ieatnetsimv7035-50;check netsim cpu;OK;HARD;3;OK- CPU IS USED 35.21%
[1570722992] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570723173] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570723354] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570724126] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570724847] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570725214] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1570725468] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570725522] SERVICE FLAPPING ALERT: ieatnetsimv7035-62;check netsim cpu;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1570725522] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;OK;HARD;3;OK- CPU IS USED 65.29%
[1570726116] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6409) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570726671] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1570727655] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13383) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570727798] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13142) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570727941] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13777) IS HIGHER THAN 95% OF ENM COMECIM NODES (14097)
[1570728014] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570728049] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 95.06%
[1570728165] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570728200] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.48%
[1570728316] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 85.06%
[1570728351] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570729031] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570729153] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.31%
[1570729182] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570729241] SERVICE FLAPPING ALERT: ieatnetsimv7035-52;check netsim cpu;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1570729241] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED 84.69%
[1570729305] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570729333] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.99%
[1570729456] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570729791] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (180) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570730207] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2710) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570730896] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570731947] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570731982] SERVICE FLAPPING ALERT: ieatnetsimv7035-51;check netsim cpu;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1570731982] SERVICE ALERT: ieatnetsimv7035-51;check netsim cpu;OK;HARD;3;OK- CPU IS USED 25.83%
[1570732098] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.98%
[1570732249] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1570732347] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.05%
[1570732872] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570732964] SERVICE FLAPPING ALERT: ieatnetsimv7035-54;check netsim cpu;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1570732964] SERVICE ALERT: ieatnetsimv7035-54;check netsim cpu;OK;HARD;3;OK- CPU IS USED 67.25%
[1570733023] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;SOFT;2;CRITICAL- CPU IS USED 100%!
[1570733174] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 99.97%
[1570735591] SERVICE ALERT: ieatnetsimv7035-53;check netsim cpu;OK;HARD;3;OK- CPU IS USED 31.5%
[1570735879] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;HARD;3;OK- CPU IS USED 43.66%
[1570739884] SERVICE FLAPPING ALERT: ieatnetsimv7035-59;check netsim cpu;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1570739884] SERVICE ALERT: ieatnetsimv7035-59;check netsim cpu;OK;HARD;3;OK- CPU IS USED 86.06%
[1570739991] SERVICE ALERT: ieatnetsimv7035-61;check netsim cpu;OK;HARD;3;OK- CPU IS USED 72.36%
[1570740362] SERVICE ALERT: ieatnetsimv7035-61;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570740362] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570740362] HOST ALERT: ieatnetsimv7035-61;DOWN;SOFT;1;PING NOK
[1570740364] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570740364] HOST ALERT: ieatnetsimv7035-71;DOWN;SOFT;1;PING NOK
[1570740372] HOST ALERT: ieatnetsimv7035-52;DOWN;SOFT;1;PING NOK
[1570740438] HOST ALERT: ieatnetsimv7035-55;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570740452] HOST ALERT: ieatnetsimv7035-61;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1570740455] HOST ALERT: ieatnetsimv7035-71;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1570740456] HOST ALERT: ieatnetsimv7035-63;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570740462] HOST ALERT: ieatnetsimv7035-54;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1570740463] HOST ALERT: ieatnetsimv7035-52;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1570740467] HOST ALERT: ieatnetsimv7035-58;DOWN;SOFT;1;(Host check timed out after 31.00 seconds)
[1570740469] HOST ALERT: ieatnetsimv7035-47;DOWN;SOFT;1;(Host check timed out after 30.93 seconds)
[1570740481] HOST ALERT: ieatnetsimv7035-59;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1570740486] HOST ALERT: ieatnetsimv7035-12;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1570740496] HOST ALERT: ieatnetsimv7035-60;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570740501] SERVICE ALERT: ieatnetsimv7035-52;check netsim cpu;OK;HARD;3;OK- CPU IS USED %
[1570740504] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570740526] HOST ALERT: ieatnetsimv7035-56;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1570740528] HOST ALERT: ieatnetsimv7035-55;DOWN;SOFT;2;(Host check timed out after 30.15 seconds)
[1570740531] HOST ALERT: ieatnetsimv7035-52;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1570740531] HOST ALERT: ieatnetsimv7035-57;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1570740534] HOST ALERT: ieatlms5735-1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570740540] SERVICE ALERT: ieatnetsimv7035-75;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570740542] HOST ALERT: ieatnetsimv7035-61;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1570740546] HOST ALERT: ieatnetsimv7035-71;DOWN;SOFT;3;(Host check timed out after 30.16 seconds)
[1570740546] HOST ALERT: ieatnetsimv7035-63;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1570740553] HOST ALERT: ieatnetsimv7035-54;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1570740557] HOST ALERT: ieatnetsimv7035-58;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1570740560] HOST ALERT: ieatnetsimv7035-47;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1570740562] HOST ALERT: ieatnetsimv7035-80;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570740566] HOST ALERT: ieatnetsimv7035-88;DOWN;SOFT;1;PING NOK
[1570740566] HOST ALERT: ieatnetsimv7035-09;DOWN;SOFT;1;PING NOK
[1570740566] HOST ALERT: ieatnetsimv7035-01;DOWN;SOFT;1;PING NOK
[1570740566] HOST ALERT: ieatnetsimv7035-46;DOWN;SOFT;1;PING NOK
[1570740566] HOST ALERT: ieatnetsimv7035-91;DOWN;SOFT;1;PING NOK
[1570740566] HOST ALERT: ieatnetsimv7035-95;DOWN;SOFT;1;PING NOK
[1570740566] HOST ALERT: ieatnetsimv7035-85;DOWN;SOFT;1;PING NOK
[1570740566] HOST ALERT: ieatnetsimv7035-66;DOWN;SOFT;1;PING NOK
[1570740566] HOST ALERT: ieatnetsimv7035-110;DOWN;SOFT;1;PING NOK
[1570740566] HOST ALERT: ieatnetsimv7035-24;DOWN;SOFT;1;PING NOK
[1570740566] HOST ALERT: ieatnetsimv7035-70;DOWN;SOFT;1;PING NOK
[1570740567] HOST ALERT: ieatnetsimv7035-67;DOWN;SOFT;1;(Host check timed out after 30.62 seconds)
[1570740567] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (0) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (0)
[1570740567] HOST ALERT: ieatnetsimv7035-19;DOWN;SOFT;1;PING NOK
[1570740567] HOST ALERT: ieatnetsimv7035-100;DOWN;SOFT;1;PING NOK
[1570740567] HOST ALERT: ieatnetsimv7035-28;DOWN;SOFT;1;PING NOK
[1570740567] HOST ALERT: ieatnetsimv7035-32;DOWN;SOFT;1;PING NOK
[1570740568] HOST ALERT: ieatnetsimv7035-93;DOWN;SOFT;1;PING NOK
[1570740568] HOST ALERT: ieatnetsimv7035-104;DOWN;SOFT;1;PING NOK
[1570740568] HOST ALERT: ieatnetsimv7035-45;DOWN;SOFT;1;PING NOK
[1570740568] HOST ALERT: ieatnetsimv7035-73;DOWN;SOFT;1;PING NOK
[1570740568] HOST ALERT: ieatnetsimv7035-40;DOWN;SOFT;1;PING NOK
[1570740568] HOST ALERT: ieatnetsimv7035-02;DOWN;SOFT;1;PING NOK
[1570740568] HOST ALERT: ieatnetsimv7035-62;DOWN;SOFT;1;PING NOK
[1570740569] HOST ALERT: ieatnetsimv7035-82;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570740569] HOST ALERT: ieatnetsimv7035-17;DOWN;SOFT;1;PING NOK
[1570740570] HOST ALERT: ieatnetsimv7035-101;DOWN;SOFT;1;PING NOK
[1570740571] HOST ALERT: ieatnetsimv7035-60;DOWN;SOFT;2;PING NOK
[1570740571] HOST ALERT: ieatnetsimv7035-75;DOWN;SOFT;1;(Host check timed out after 31.00 seconds)
[1570740571] HOST ALERT: ieatnetsimv7035-107;DOWN;SOFT;1;PING NOK
[1570740571] HOST ALERT: ieatnetsimv7035-38;DOWN;SOFT;1;PING NOK
[1570740571] HOST ALERT: ieatnetsimv7035-08;DOWN;SOFT;1;PING NOK
[1570740571] HOST ALERT: ieatnetsimv7035-18;DOWN;SOFT;1;PING NOK
[1570740571] HOST ALERT: ieatlms5735-1;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1570740572] HOST ALERT: ieatnetsimv7035-14;DOWN;SOFT;1;PING NOK
[1570740572] HOST ALERT: ieatnetsimv7035-59;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1570740574] HOST ALERT: ieatnetsimv7035-84;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570740576] HOST ALERT: ieatnetsimv7035-03;DOWN;SOFT;1;PING NOK
[1570740577] HOST ALERT: ieatnetsimv7035-12;DOWN;SOFT;2;(Host check timed out after 31.01 seconds)
[1570740578] HOST ALERT: ieatnetsimv7035-86;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1570740580] HOST ALERT: ieatnetsimv7035-50;DOWN;SOFT;1;PING NOK
[1570740581] SERVICE ALERT: ieatnetsimv7035-05;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570740581] HOST ALERT: ieatnetsimv7035-05;DOWN;SOFT;1;PING NOK
[1570740583] SERVICE ALERT: ieatnetsimv7035-81;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570740583] HOST ALERT: ieatnetsimv7035-81;DOWN;SOFT;1;PING NOK
[1570740584] SERVICE ALERT: ieatnetsimv7035-25;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570740584] HOST ALERT: ieatnetsimv7035-25;DOWN;SOFT;1;PING NOK
[1570740586] HOST ALERT: ieatnetsimv7035-56;DOWN;SOFT;2;PING NOK
[1570740586] HOST ALERT: ieatnetsimv7035-30;DOWN;SOFT;1;PING NOK
[1570740588] HOST ALERT: ieatnetsimv7035-55;DOWN;SOFT;3;PING NOK
[1570740593] HOST ALERT: ieatnetsimv7035-52;UP;SOFT;1;PING OK
[1570740594] HOST ALERT: ieatnetsimv7035-57;UP;SOFT;1;PING OK
[1570740594] SERVICE ALERT: ieatnetsimv7035-03;check netsim nodes;CRITICAL;HARD;1;CRITICAL- THERE ARE STOPPED NODES
[1570740604] HOST ALERT: ieatnetsimv7035-61;UP;SOFT;1;PING OK
[1570740608] HOST ALERT: ieatnetsimv7035-71;UP;SOFT;1;PING OK
[1570740608] HOST ALERT: ieatnetsimv7035-63;UP;SOFT;1;PING OK
[1570740609] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570740609] SERVICE ALERT: ieatnetsimv7035-61;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570740612] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570740615] HOST ALERT: ieatnetsimv7035-54;UP;SOFT;1;PING OK
[1570740620] HOST ALERT: ieatnetsimv7035-58;UP;SOFT;1;PING OK
[1570740622] HOST ALERT: ieatnetsimv7035-47;UP;SOFT;1;PING OK
[1570740624] HOST ALERT: ieatnetsimv7035-80;UP;SOFT;1;PING OK
[1570740626] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570740628] HOST ALERT: ieatnetsimv7035-91;UP;SOFT;1;PING OK
[1570740628] HOST ALERT: ieatnetsimv7035-88;UP;SOFT;1;PING OK
[1570740628] HOST ALERT: ieatnetsimv7035-95;UP;SOFT;1;PING OK
[1570740628] HOST ALERT: ieatnetsimv7035-46;UP;SOFT;1;PING OK
[1570740628] HOST ALERT: ieatnetsimv7035-01;UP;SOFT;1;PING OK
[1570740628] HOST ALERT: ieatnetsimv7035-09;UP;SOFT;1;PING OK
[1570740628] HOST ALERT: ieatnetsimv7035-24;UP;SOFT;1;PING OK
[1570740628] HOST ALERT: ieatnetsimv7035-85;UP;SOFT;1;PING OK
[1570740628] HOST ALERT: ieatnetsimv7035-110;UP;SOFT;1;PING OK
[1570740628] HOST ALERT: ieatnetsimv7035-66;UP;SOFT;1;PING OK
[1570740628] HOST ALERT: ieatnetsimv7035-70;UP;SOFT;1;PING OK
[1570740629] HOST ALERT: ieatnetsimv7035-67;UP;SOFT;1;PING OK
[1570740629] HOST ALERT: ieatnetsimv7035-19;UP;SOFT;1;PING OK
[1570740629] HOST ALERT: ieatnetsimv7035-32;UP;SOFT;1;PING OK
[1570740629] HOST ALERT: ieatnetsimv7035-28;UP;SOFT;1;PING OK
[1570740629] HOST ALERT: ieatnetsimv7035-100;UP;SOFT;1;PING OK
[1570740630] HOST ALERT: ieatnetsimv7035-73;UP;SOFT;1;PING OK
[1570740630] HOST ALERT: ieatnetsimv7035-93;UP;SOFT;1;PING OK
[1570740630] HOST ALERT: ieatnetsimv7035-02;UP;SOFT;1;PING OK
[1570740630] HOST ALERT: ieatnetsimv7035-45;UP;SOFT;1;PING OK
[1570740630] HOST ALERT: ieatnetsimv7035-40;UP;SOFT;1;PING OK
[1570740630] HOST ALERT: ieatnetsimv7035-104;UP;SOFT;1;PING OK
[1570740631] HOST ALERT: ieatnetsimv7035-62;UP;SOFT;1;PING OK
[1570740631] HOST ALERT: ieatnetsimv7035-82;UP;SOFT;1;PING OK
[1570740631] HOST ALERT: ieatnetsimv7035-17;UP;SOFT;1;PING OK
[1570740632] HOST ALERT: ieatnetsimv7035-56;UP;SOFT;1;PING OK
[1570740632] HOST ALERT: ieatnetsimv7035-101;UP;SOFT;1;PING OK
[1570740633] HOST ALERT: ieatnetsimv7035-60;UP;SOFT;1;PING OK
[1570740633] HOST ALERT: ieatnetsimv7035-75;UP;SOFT;1;PING OK
[1570740634] HOST ALERT: ieatnetsimv7035-107;UP;SOFT;1;PING OK
[1570740634] HOST ALERT: ieatnetsimv7035-08;UP;SOFT;1;PING OK
[1570740634] HOST ALERT: ieatnetsimv7035-18;UP;SOFT;1;PING OK
[1570740634] HOST ALERT: ieatnetsimv7035-38;UP;SOFT;1;PING OK
[1570740634] HOST ALERT: ieatnetsimv7035-14;UP;SOFT;1;PING OK
[1570740634] HOST ALERT: ieatnetsimv7035-59;UP;SOFT;1;PING OK
[1570740637] HOST ALERT: ieatnetsimv7035-84;UP;SOFT;1;PING OK
[1570740638] HOST ALERT: ieatnetsimv7035-03;UP;SOFT;1;PING OK
[1570740639] HOST ALERT: ieatnetsimv7035-12;UP;SOFT;1;PING OK
[1570740641] HOST ALERT: ieatnetsimv7035-86;UP;SOFT;1;PING OK
[1570740642] HOST ALERT: ieatnetsimv7035-50;UP;SOFT;1;PING OK
[1570740643] HOST ALERT: ieatnetsimv7035-05;UP;SOFT;1;PING OK
[1570740645] HOST ALERT: ieatnetsimv7035-81;UP;SOFT;1;PING OK
[1570740646] HOST ALERT: ieatnetsimv7035-25;UP;SOFT;1;PING OK
[1570740649] HOST ALERT: ieatnetsimv7035-30;UP;SOFT;1;PING OK
[1570740650] HOST ALERT: ieatnetsimv7035-55;UP;SOFT;1;PING OK
[1570740675] SERVICE ALERT: ieatnetsimv7035-75;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1570740682] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570740715] SERVICE ALERT: ieatnetsimv7035-05;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1570740717] SERVICE ALERT: ieatnetsimv7035-81;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1570740720] SERVICE ALERT: ieatnetsimv7035-25;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1570740732] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570740761] SERVICE ALERT: ieatnetsimv7035-71;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570740779] SERVICE ALERT: ieatnetsimv7035-61;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570740807] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570740876] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570741226] HOST ALERT: ieatnetsimv7035-110;DOWN;SOFT;1;PING NOK
[1570741288] HOST FLAPPING ALERT: ieatnetsimv7035-110;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1570741288] HOST ALERT: ieatnetsimv7035-110;UP;SOFT;1;PING OK
[1570741893] SERVICE FLAPPING ALERT: ieatlms5735-1;check nas hc;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1570742281] SERVICE ALERT: ieatnetsimv7035-03;check netsim nodes;OK;SOFT;1;OK- THERE ARE NO STOPPED NODES
[1570742597] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570743160] SERVICE ALERT: ieatnetsimv7035-63;check netsim cpu;OK;HARD;3;OK- CPU IS USED 37.13%
[1570743171] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;OK;HARD;3;OK- CPU IS USED 50.76%
[1570743229] SERVICE FLAPPING ALERT: ieatnetsimv7035-56;check netsim cpu;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1570743229] SERVICE ALERT: ieatnetsimv7035-56;check netsim cpu;OK;HARD;3;OK- CPU IS USED 44.29%
[1570743233] SERVICE FLAPPING ALERT: ieatnetsimv7035-57;check netsim cpu;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1570743233] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;OK;HARD;3;OK- CPU IS USED 27.91%
[1570743474] SERVICE FLAPPING ALERT: ieatnetsimv7035-47;check netsim cpu;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1570743474] SERVICE ALERT: ieatnetsimv7035-47;check netsim cpu;OK;HARD;3;OK- CPU IS USED 63.3%
[1570743500] SERVICE FLAPPING ALERT: ieatnetsimv7035-60;check netsim cpu;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1570743500] SERVICE ALERT: ieatnetsimv7035-60;check netsim cpu;OK;HARD;3;OK- CPU IS USED 22%
[1570743887] SERVICE ALERT: ieatnetsimv7035-62;check netsim cpu;OK;HARD;3;OK- CPU IS USED 16.07%
[1570744181] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1570744329] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1570744472] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (10000) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1570744669] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570744769] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.09%
[1570744865] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1570744920] SERVICE ALERT: ieatnetsimv7035-94;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 35.39%
[1570745058] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1570745127] HOST ALERT: ieatnetsimv7035-100;DOWN;SOFT;1;PING NOK
[1570745189] HOST ALERT: ieatnetsimv7035-100;UP;SOFT;1;PING OK
[1570745488] HOST FLAPPING ALERT: ieatnetsimv7035-110;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570746710] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570746779] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570746802] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570746861] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 35.14%
[1570746864] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1570746930] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 94.44%
[1570746934] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570746953] SERVICE ALERT: ieatnetsimv7035-58;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 40.61%
[1570747012] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 98.74%
[1570747014] SERVICE ALERT: ieatnetsimv7035-57;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 60.29%
[1570747081] SERVICE ALERT: ieatnetsimv7035-55;check netsim cpu;OK;HARD;3;OK- CPU IS USED 17.26%
[1570747162] SERVICE ALERT: ieatnetsimv7035-49;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 14.74%
[1570747763] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;1;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13364) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570747924] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;WARNING;SOFT;2;WARNING- NUMBER OF COMECIM NODES SYNCHRONIZED (13382) IS FROM 80% TO 95% OF ENM COMECIM NODES (14097)
[1570747998] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570748008] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570748068] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13829) IS HIGHER THAN 95% OF ENM COMECIM NODES (14097)
[1570748119] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570748128] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570748198] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.15%
[1570748239] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570748249] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1570748341] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1570748349] SERVICE ALERT: ieatnetsimv5116-32;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 43.58%