[1567552407] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.72%
[1567552484] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (1974) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1567552558] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 93.17%
[1567552635] SERVICE ALERT: ieatlms3901;check synchronized cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (3357) IS LOWER THAN 80% OF ENM CPP NODES (7153)
[1567552708] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;HARD;3;OK- CPU IS USED 89.27%
[1567552778] SERVICE ALERT: ieatlms3901;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (6466) IS FROM 80% TO 95% OF ENM CPP NODES (7153)
[1567554603] SERVICE ALERT: ieatlms3901;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (7141) IS HIGHER THAN 95% OF ENM CPP NODES (7153)
[1567555555] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 83.16%
[1567556339] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.56%
[1567556490] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 51.98%
[1567556641] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.72%
[1567556791] SERVICE ALERT: ieatnetsimv5116-78;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 76.44%
[1567558548] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.14%
[1567558699] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 39.33%
[1567558850] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.97%
[1567559000] SERVICE ALERT: ieatnetsimv7035-109;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 6.7%
[1567559186] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567559337] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.93%
[1567563088] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.84%
[1567563238] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.38%
[1567563389] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 98.05%
[1567567019] SERVICE FLAPPING ALERT: ieatnetsimv7035-09;check netsim cpu;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1567567019] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 79.97%
[1567574250] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.93%
[1567574401] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.26%
[1567574552] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.88%
[1567574703] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 80.48%
[1567575451] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567575645] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567575838] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567576585] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567576771] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;SOFT;2;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567576909] SERVICE ALERT: ieatnetsimv5116-72;check netsim disk;OK;HARD;3;OK
[1567577123] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1567577244] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1567577365] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1567577457] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 99.97%
[1567577608] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 99.98%
[1567577712] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;HARD;3;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1567577759] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;CRITICAL;HARD;3;CRITICAL- CPU IS USED 100%!
[1567578454] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567578604] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 85.21%
[1567578755] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567578906] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 82.96%
[1567580965] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;OK;HARD;3;OK
[1567581390] SERVICE ALERT: ieatnetsimv5116-14;check netsim cpu;OK;HARD;3;OK- CPU IS USED 53.49%
[1567582616] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1567582658] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 96.84%
[1567582809] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 81.6%
[1567582959] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567583104] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1567583110] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 84.73%
[1567583388] SERVICE ALERT: ieatnetsimv7035-14;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567583554] SERVICE ALERT: ieatnetsimv7035-14;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567584566] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1567584687] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1567584795] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567584807] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1567584938] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567585124] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567585370] SERVICE ALERT: ieatlms5735-1;check synchronized comecim;OK;HARD;3;OK- NUMBER OF COMECIM NODES SYNCHRONIZED (13472) IS HIGHER THAN 95% OF ENM COMECIM NODES (14049)
[1567586791] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RadioNode;OK;HARD;3;OK- NUMBER OF PM FILES (13369) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (13935)
[1567586861] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.55%
[1567587012] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 86.95%
[1567587211] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567587362] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567587510] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567587769] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567587954] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567588140] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567588408] SERVICE FLAPPING ALERT: ieatnetsimv5116-59;check netsim disk;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1567588408] SERVICE ALERT: ieatnetsimv5116-59;check netsim disk;OK;HARD;3;OK
[1567588736] SERVICE ALERT: ieatlms5735-1;check pm collection statistical RBS;OK;HARD;3;OK- NUMBER OF PM FILES (1514) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (1514)
[1567589151] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4030) IS HIGHER THAN 95% OF ENM CPP NODES (4224)
[1567590005] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567593110] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.16%
[1567593261] SERVICE ALERT: ieatnetsimv7035-13;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 78.41%
[1567594575] SERVICE ALERT: ieatlms3901;check pm supervision nodes;CRITICAL;HARD;3;(Service check timed out after 800.01 seconds)
[1567594663] SERVICE ALERT: ieatlms3901;check cm supervision nodes;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1567594670] SERVICE ALERT: ieatnetsimv5116-73;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1567594790] SERVICE ALERT: ieatnetsimv5116-73;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1567594888] SERVICE ALERT: ieatnetsimv5116-63;check netsim disk;CRITICAL;SOFT;1;WARNING- FILESYSTEM IS USED >90%
[1567594911] SERVICE ALERT: ieatnetsimv5116-73;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1567595009] SERVICE ALERT: ieatnetsimv5116-63;check netsim disk;CRITICAL;SOFT;2;WARNING- FILESYSTEM IS USED >90%
[1567595065] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.82%
[1567595130] SERVICE ALERT: ieatnetsimv5116-63;check netsim disk;CRITICAL;HARD;3;WARNING- FILESYSTEM IS USED >90%
[1567595147] SERVICE ALERT: ieatlms3901;check pm collection statistical;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (0) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (21331)
[1567595215] SERVICE ALERT: ieatnetsimv7035-89;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 28.08%
[1567595361] SERVICE ALERT: ieatlms3901;check cm supervision nodes;OK;SOFT;2;OK- NUMBER OF NODES WITH CM SUPERVISION (24198) IS HIGHER THAN 95% OF ENM NODES (25422)
[1567596268] SERVICE ALERT: ieatlms3901;check pm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH PM SUPERVISION (23838) IS BETWEEN 80% AND 95% OF ENM NODES (25422)
[1567597963] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 90.15%
[1567598114] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 83.4%
[1567598264] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;CRITICAL- CPU IS USED 100%!
[1567598415] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.56%
[1567598566] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;HARD;3;WARNING- CPU IS USED 91.45%
[1567599044] SERVICE ALERT: ieatlms3901;check pm collection statistical;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (18777) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (21331)
[1567600870] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567601055] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567601241] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1567602112] SERVICE ALERT: ieatnetsimv5116-73;check netsim disk;OK;HARD;3;OK
[1567602197] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 87.81%
[1567602334] SERVICE ALERT: ieatnetsimv5116-63;check netsim disk;OK;HARD;3;OK
[1567605069] SERVICE ALERT: ieatnetsimv7035-30;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.83%
[1567605221] SERVICE ALERT: ieatnetsimv7035-30;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 67.96%
[1567613786] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567613979] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567614173] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED
[1567615642] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1567615763] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1567615848] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567615884] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1567616041] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;1;CRITICAL- vcs_cluster_healthcheck FAILED
[1567616162] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;SOFT;2;CRITICAL- vcs_cluster_healthcheck FAILED
[1567616283] SERVICE ALERT: ieatlms5735-1;check vcs cluster hc;CRITICAL;HARD;3;CRITICAL- vcs_cluster_healthcheck FAILED
[1567616470] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;1;CRITICAL- hw_resources_healthcheck FAILED
[1567616591] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;SOFT;2;CRITICAL- hw_resources_healthcheck FAILED
[1567616712] SERVICE ALERT: ieatlms5735-1;check hw resources hc;CRITICAL;HARD;3;CRITICAL- hw_resources_healthcheck FAILED
[1567616817] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1567616938] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1567616961] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;1;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567617059] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1567617082] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;SOFT;2;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567617203] SERVICE ALERT: ieatlms5735-1;check vcs llt heartbeat hc;CRITICAL;HARD;3;CRITICAL- vcs_llt_heartbeat_healthcheck FAILED
[1567617378] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;1;CRITICAL- storagepool_healthcheck FAILED
[1567617498] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;SOFT;2;CRITICAL- storagepool_healthcheck FAILED
[1567617619] SERVICE ALERT: ieatlms5735-1;check storagepool hc;CRITICAL;HARD;3;CRITICAL- storagepool_healthcheck FAILED
[1567617842] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567617929] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;1;CRITICAL- mdt_healthcheck FAILED
[1567618025] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1567618050] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;SOFT;2;CRITICAL- mdt_healthcheck FAILED
[1567618131] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567618146] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1567618171] SERVICE ALERT: ieatlms5735-1;check mdt hc;CRITICAL;HARD;3;CRITICAL- mdt_healthcheck FAILED
[1567618267] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1567618421] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567624483] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.59%
[1567624633] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 37.32%
[1567624784] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.34%
[1567624935] SERVICE ALERT: ieatnetsimv5116-72;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 61.52%
[1567625763] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567627428] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.52%
[1567627578] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;CRITICAL;SOFT;2;WARNING- CPU IS USED 91.41%
[1567627729] SERVICE ALERT: ieatnetsimv7035-09;check netsim cpu;OK;HARD;3;OK- CPU IS USED 80.7%
[1567627774] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567628081] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567628370] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567630397] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567630401] SERVICE ALERT: ieatnetsimv5116-73;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 91.08%
[1567630552] SERVICE ALERT: ieatnetsimv5116-73;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 7.3%
[1567630729] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567630873] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567630935] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567632848] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567633110] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567633454] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567633763] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567633980] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1567634773] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567635091] SERVICE ALERT: ieatnetsimv7035-10;check netsim nodes;OK;SOFT;2;OK- THERE ARE NO STOPPED NODES
[1567635996] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1567636281] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1567636602] SERVICE ALERT: ieatnetsimv7035-08;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1567637524] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1567637721] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;SOFT;2;CRITICAL- system_service_healthcheck FAILED
[1567637919] SERVICE ALERT: ieatlms3901;check system service hc;CRITICAL;HARD;3;CRITICAL- system_service_healthcheck FAILED