[1570146035] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570146806] SERVICE FLAPPING ALERT: ieatlms3901;check fm hb failures cpp;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570146806] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570155454] SERVICE FLAPPING ALERT: ieatlms5735-1;check synchronized comecim;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570158121] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (3999) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570158544] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570159303] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6439) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570161642] SERVICE FLAPPING ALERT: ieatlms5735-1;check hw resources hc;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1570162433] SERVICE ALERT: ieatlms5735-1;check node fs hc;OK;HARD;3;OK- node_fs_healthcheck SUCCESSFULLY PASSED
[1570162984] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;OK
[1570163226] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570163979] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570166645] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570166826] SERVICE ALERT: ieatlms5735-1;check node fs;OK;SOFT;2;OK
[1570167297] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570167483] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570167668] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570168499] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;(Service check timed out after 800.01 seconds)
[1570169274] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570171165] HOST ALERT: ieatnetsimv7035-96;DOWN;SOFT;1;PING NOK
[1570171227] HOST ALERT: ieatnetsimv7035-96;UP;SOFT;1;PING OK
[1570171333] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570173199] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570173293] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;1;CRITICAL- node_fs_healthcheck FAILED
[1570173384] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570173474] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;SOFT;2;CRITICAL- node_fs_healthcheck FAILED
[1570173569] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570173655] SERVICE ALERT: ieatlms5735-1;check node fs hc;CRITICAL;HARD;3;CRITICAL- node_fs_healthcheck FAILED
[1570174207] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570174212] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570174366] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570174388] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570174516] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570174568] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570177234] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1570177234] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570177882] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570178019] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570178131] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;1;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3574) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1570178164] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570178268] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;SOFT;2;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3867) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1570178403] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4087) IS HIGHER THAN 95% OF ENM CPP NODES (4255)
[1570179100] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570179285] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570179470] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570179996] SERVICE ALERT: ieatnetsimv017-29;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570180034] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (1276) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570180160] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (1427) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570180216] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2928) IS LOWER THAN 80% OF ENM CPP NODES (4255)
[1570180286] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1930) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570180351] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (2202) IS LOWER THAN 80% OF ENM CPP NODES (4255)
[1570180483] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;CRITICAL;HARD;3;CRITICAL- NUMBER OF CPP NODES SYNCHRONIZED (1846) IS LOWER THAN 80% OF ENM CPP NODES (4255)
[1570180869] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (36) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570180892] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1570180995] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (36) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570181012] SERVICE ALERT: ieatnetsimv017-15;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570181119] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (50) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570181160] SERVICE ALERT: ieatnetsimv017-15;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570181253] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1570181301] SERVICE ALERT: ieatnetsimv017-15;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570181613] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1570182835] SERVICE ALERT: ieatnetsimv017-16;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570182975] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (113) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570182992] SERVICE ALERT: ieatnetsimv017-16;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570183100] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (113) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570183142] SERVICE ALERT: ieatnetsimv017-16;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570183225] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (74) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570183629] SERVICE ALERT: ieatlms5735-1;check system service hc;CRITICAL;SOFT;1;CRITICAL- system_service_healthcheck FAILED
[1570183629] SERVICE ALERT: ieatlms5735-1;check node fs;OK;HARD;3;(No output on stdout) stderr: Warning: Permanently added '131.160.170.194' (RSA) to the list of known hosts.
[1570183823] SERVICE ALERT: ieatlms5735-1;check system service hc;OK;SOFT;2;OK- system_service_healthcheck SUCCESSFULLY PASSED
[1570184171] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570184623] SERVICE ALERT: ieatlms3901;check fm supervision nodes;OK;HARD;3;OK- NUMBER OF NODES WITH FM SUPERVISION (22698) IS HIGHER THAN 95% OF ENM NODES (23118)
[1570184791] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;2;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (2740) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570184936] SERVICE ALERT: ieatnetsimv017-15;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570184959] SERVICE ALERT: ieatnetsimv017-14;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570184982] SERVICE ALERT: ieatnetsimv017-16;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570185490] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;1;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570185550] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;HARD;3;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570185656] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;SOFT;1;WARNING- NUMBER OF NODES WITH FM SUPERVISION (22698) IS BETWEEN 80% AND 95% OF ENM NODES (26967)
[1570185670] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;SOFT;2;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570185851] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1570185851] SERVICE ALERT: ieatlms5735-1;check node fs;CRITICAL;HARD;3;CRITICAL- THERE ARE FILESYSTEMS EXCEEDING USAGE
[1570185967] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;SOFT;2;WARNING- NUMBER OF NODES WITH FM SUPERVISION (22698) IS BETWEEN 80% AND 95% OF ENM NODES (26967)
[1570186190] SERVICE ALERT: ieatlms3901;check fm supervision nodes;WARNING;HARD;3;WARNING- NUMBER OF NODES WITH FM SUPERVISION (22698) IS BETWEEN 80% AND 95% OF ENM NODES (26967)
[1570186769] SERVICE ALERT: ieatnetsimv017-15;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570186797] SERVICE ALERT: ieatnetsimv017-14;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570186828] SERVICE ALERT: ieatnetsimv017-16;check netsim nodes;CRITICAL;SOFT;1;CRITICAL- THERE ARE STOPPED NODES
[1570186911] SERVICE ALERT: ieatnetsimv017-15;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570186949] SERVICE ALERT: ieatnetsimv017-14;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570186994] SERVICE ALERT: ieatnetsimv017-16;check netsim nodes;CRITICAL;SOFT;2;CRITICAL- THERE ARE STOPPED NODES
[1570187061] SERVICE ALERT: ieatnetsimv017-15;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570187099] SERVICE ALERT: ieatnetsimv017-14;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570187148] SERVICE ALERT: ieatnetsimv017-16;check netsim nodes;CRITICAL;HARD;3;CRITICAL- THERE ARE STOPPED NODES
[1570187495] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;WARNING;HARD;3;WARNING- NUMBER OF PM FILES (2270) IS BETWEEN 80% AND 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570188952] SERVICE ALERT: ieatnetsimv017-31;check netsim nodes;OK;HARD;3;OK- THERE ARE NO STOPPED NODES
[1570189616] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;CRITICAL;SOFT;1;CRITICAL- NUMBER OF CPP NODES W/O FM HB FAILURES (4029) IS LOWER THAN 80% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570190335] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;OK;HARD;3;OK- vcs_service_group_healthcheck SUCCESSFULLY PASSED
[1570190352] SERVICE ALERT: ieatlms3901;check fm hb failures cpp;OK;SOFT;2;OK- NUMBER OF CPP NODES W/O FM HB FAILURES (6589) IS HIGHER THAN 95% OF CPP NODES WITH FM SUPERVISION ENABLED (6589)
[1570191101] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1913) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570194029] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;OK;HARD;3;OK- NUMBER OF PM FILES (147) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570194707] SERVICE FLAPPING ALERT: ieatlms5735-1;check pm collection statistical ERBS;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1570194707] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;OK;HARD;3;OK- NUMBER OF PM FILES (2559) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570195524] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;OK;HARD;3;OK- NUMBER OF PM FILES (80) IS HIGHER THAN 95% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570195801] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;1;CRITICAL- vcs_service_group_healthcheck FAILED
[1570195996] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;SOFT;2;CRITICAL- vcs_service_group_healthcheck FAILED
[1570196181] SERVICE ALERT: ieatlms5735-1;check vcs service group hc;CRITICAL;HARD;3;CRITICAL- vcs_service_group_healthcheck FAILED
[1570196255] SERVICE ALERT: ieatlms5735-1;check nas hc;OK;HARD;3;OK- nas_healthcheck SUCCESSFULLY PASSED
[1570198343] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 92.34%
[1570198493] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 42.85%
[1570200032] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;1;CRITICAL- nas_healthcheck FAILED
[1570200394] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;SOFT;2;CRITICAL- nas_healthcheck FAILED
[1570200696] SERVICE ALERT: ieatlms5735-1;check nas hc;CRITICAL;HARD;3;CRITICAL- nas_healthcheck FAILED
[1570201235] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (70) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570201360] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (70) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570201486] SERVICE ALERT: ieatlms5735-1;check pm collection statistical MGW;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (70) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (150)
[1570201914] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (1911) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570202043] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (1911) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570202169] SERVICE ALERT: ieatlms5735-1;check pm collection statistical ERBS;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (1279) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (2560)
[1570202729] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;1;CRITICAL- NUMBER OF PM FILES (39) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570202853] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;SOFT;2;CRITICAL- NUMBER OF PM FILES (39) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570202978] SERVICE ALERT: ieatlms5735-1;check pm collection statistical SGSN-MME;CRITICAL;HARD;3;CRITICAL- NUMBER OF PM FILES (39) IS LOWER THAN 80% OF NODES WITH PM SUPERVISION ENABLED (80)
[1570205700] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;WARNING;HARD;3;WARNING- NUMBER OF CPP NODES SYNCHRONIZED (3640) IS FROM 80% TO 95% OF ENM CPP NODES (4255)
[1570209314] SERVICE ALERT: ieatlms5735-1;check synchronized cpp;OK;HARD;3;OK- NUMBER OF CPP NODES SYNCHRONIZED (4253) IS HIGHER THAN 95% OF ENM CPP NODES (4255)
[1570209484] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;HARD;3;OK
[1570213091] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;WARNING;SOFT;1;WARNING- THE FOLLOWING NODES ARE NOT SYNCHRONIZED:
[1570213217] SERVICE ALERT: ieatlms5735-1;check synchronized nodes cpp;OK;SOFT;2;OK
[1570220112] SERVICE FLAPPING ALERT: ieatlms5735-1;check node fs;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1570226846] SERVICE FLAPPING ALERT: ieatlms5735-1;check vcs service group hc;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1570227444] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 94.56%
[1570227594] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 66.17%
[1570227745] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.99%
[1570227896] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 54.45%
[1570228046] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;CRITICAL;SOFT;1;WARNING- CPU IS USED 93.02%
[1570228197] SERVICE ALERT: ieatnetsimv5116-35;check netsim cpu;OK;SOFT;2;OK- CPU IS USED 59.7%